## DESCRIPTION
## Calculating double integrals in polar coordinates  
## ENDDESCRIPTION

## KEYWORDS('Double Integrals','Iterated Integrals','Polar Coordinates')

## DBsubject('Calculus')
## DBchapter('Multiple Integrals')
## DBsection('Double Integrals in Polar Coordinates')
## Date('10/31/09')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('16.4')
## Section2('16.4')
## Problem1('18')
## Problem2('22')

DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl",
  "parserMultiAnswer.pl",
);


TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(r=>"Real",theta=>"Real",dr=>"Real",dtheta=>"Real");


#-----------------------
#  limits of integration

$a = random(3,5,1);
$asq = $a**2;


#-----------
#  integrand
#

$f = Formula("r**2 * cos(theta) * sin(theta) * r") -> reduce;

#-----------
#  volume
#

$V = Formula("$a**4 / 16") -> reduce;


#=================================
#  Multiple solutions
#
#  inner and outer differentials
#
#  arrays are (if r inner & theta outer integral, if theta inner & r outer integral)
# 
@id = (Formula("dr"),Formula("dtheta"));
@od = (Formula("dtheta"),Formula("dr"));

# inner limit high, inner limit low, outer limit high, outer limit low (case 0, case 1)
@ilh = (Formula("$a"),Formula("pi/4"));
@ill = (Formula("0"),Formula("0"));
@olh = (Formula("pi/4"),Formula("$a"));
@oll = (Formula("0"),Formula("0"));


$multipart = MultiAnswer( $f, $id[0], $od[0], $oll[0], $olh[0], $ill[0], $ilh[0] )->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $fstu, $idstu, $odstu, $ollstu, $olhstu, $illstu, $ilhstu ) = @{$student};
        if ( 
             ( 
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           return 1;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration.");
           return 0.94;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration and order of integration.");
           return 0.47;




        } elsif (
             ( 
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           $self->setMessage(1,"Check your integrand and dA.");
           return 0.47;
        } elsif (
             (
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration, your integrand, and dA.");
           return 0.47;
        } elsif (
             (
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration, order of integration, integrand, and dA.");
           return 0.24;
        } else {
             return 0;
        }
    }
  );



##############################################
#  Main text

Context()->texStrings;
Context()->flags->set(reduceConstants=>0);
BEGIN_TEXT

Convert the integral below to polar coordinates and evaluate the integral.
$BR
\( \displaystyle \int_{0}^{$a / \sqrt{2}} \int_{y}^{\sqrt{$asq - y^2}} xy \, dx \, dy \)
$BR
$BR
${BITALIC}Instructions:${EITALIC}
Please enter the integrand in the first answer box,
typing ${BITALIC}theta${EITALIC} for \(\theta\). 
Depending on the order of integration you choose,  
enter ${BITALIC}dr${EITALIC} and ${BITALIC}dtheta${EITALIC} 
in either order into the second and third answer boxes 
with only one ${BITALIC}dr${EITALIC} or 
${BITALIC}dtheta${EITALIC} in each box.  Then, enter the limits of 
integration and evaluate the integral to find the volume. 
$BR
$BR
\( \displaystyle \int_A^B \int_C^D \) 
\{ $multipart->ans_rule(40) \}
\{ $multipart->ans_rule(5) \}
\{ $multipart->ans_rule(5) \}
$BR
$BR
A = \{ $multipart->ans_rule(20) \} $BR
B = \{ $multipart->ans_rule(20) \} $BR
C = \{ $multipart->ans_rule(20) \} $BR
D = \{ $multipart->ans_rule(20) \}
$BR
$BR
Volume = \{ ans_rule(40) \}


END_TEXT
Context()->normalStrings;





##################################################
#  Answers

ANS( $multipart->cmp() );
ANS(($V)->cmp());

$showPartialCorrectAnswers = 1;

COMMENT('Allows integration in either order.');

ENDDOCUMENT();