# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Find the distance between the points (x1,y1) and (x2,y2).
#
# If you need to use the square root symbol, 
# as in x=sqrt{17}, type it like: *sqrt(17)*
#
# We use Pythagorean triples so that everything works out 
# integer.
#
#       a = m^2-n^2, b = 2mn, c = m^2+n^2
#
# We choose m and n so that m!=n which means that m-n>0.
#
# Last updated: Hughes 8/6/13
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','square root','subtract')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################
DOCUMENT();

loadMacros(
      "PGstandard.pl",
      "MathObjects.pl",
      "parserAssignment.pl",
      "answerHints.pl",
      "PGML.pl",
      "contextLimitedRadical.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
);

##############################################
    
Context("LimitedRadical");
Context()->flags->set(reduceConstants=>0, reduceConstantFunctions=>0, formatStudentAnswer=>parsed);
                               
# this gurantees that a, b, c are Pythagorean triples
$m = 1; $n = 1;
while($m == $n)
{
    $m = random(2,10,1);
    $n = random(2,10,1);
}
$a = $m**2-$n**2;
$b = 2*$m*$n;
$c = $m**2+$n**2;

# choose (x1,y1) first
$x1 = random(-20,20,1);
$y1 = random(-20,20,1);

# then (x2,y2) follow by using $a and $b
$x2 = $x1+$a;
$y2 = $y1+$b;

$ans = Formula("$c");

##############################################

TEXT(beginproblem());
BEGIN_PGML
Find the distance between the points [`([$x1],[$y1])`] and [`([$x2],[$y2])`].

[@KeyboardInstructions('If you need to use the square root symbol, as in [`x=\sqrt{17}`], type it like [|sqrt(17)|]*')@]**

      [____________________]
END_PGML

ANS( $ans-> cmp( checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 if $ans->{isPreview} || $correct != $student;
  
  $student = $ans->{original_student_ans};
  $correct = ($ans->{correct_ans}) ? $ans->{correct_ans} : $ans->{correct_value};

  Context()->flags->set(setSqrt=>exp(1)/ln(2));
  delete $correct->{test_values}, $student->{test_values};
  my $OK = ($correct == $student);
  Context()->flags->set(setSqrt => 0);

  Value::Error("Check to see if your answer is simplified.") unless $OK;
  return $OK;
  }, formatStudentAnswer=>"reduced"
  )
);


##############################################

$xdist = ($x2-$x1)**2;
$ydist = ($y2-$y1)**2;
BEGIN_PGML_SOLUTION
The distance between two points, [`d`], can be found by using the Distance Formula:

    [`
        \begin{aligned}
            d& = \sqrt{(x_2-x_1)^2 + (y_2-y_1)^2} \\
            & = \sqrt{([$x2]-[$x1])^2 + ([$y2]-[$y1])^2} \\
             & = \sqrt{([$x2-$x1])^2 + ([$y2-$y1])^2} \\
             & = \sqrt{[$xdist] + [$ydist]} \\
             & = \sqrt{[$xdist + $ydist]} \\
             & = [$ans]
        \end{aligned}
    `]

The distance between the two points is [`[$ans]`].
END_PGML_SOLUTION

##############################################

ENDDOCUMENT();
