# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## DBsubject(Calculus - single variable)
## DBchapter(Parametric)
## DBsection(Curves)
## Institution(UMich)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('4.8')
## Section2('4.8')
## Section3('4.8')
## Section4('4.8')
## Problem1('9')
## Problem2('9')
## Problem3('16')
## Problem4('16')
## KEYWORDS('parametric equations')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Interval");
Context()->variables->add(t=>"Real");
Context()->strings->add(None=>{},none=>{alias=>"None"});
$showPartialCorrectAnswers = 1;

# use one of two functions in the argument of sine/cosine
$whichF = list_random(0,1);
$fullCirc = new_multiple_choice();

if ( $whichF ) {
    $domain = "\((0,\infty)\)";
    $a = random(1,5,1);
    $g = Compute("ln($a*t)")->reduce();
    $gp = Compute("1/t");

    $grange = Compute("(-Inf, Inf)");

    # then g' > 0 when 1/t > 0, or t > 0
    $clock = List("None");
    $cclock = List("(0,Inf)");

    $clockans = "Any other values of \(t\), \(t \le 0\), are not in " .
      "the domain of the function, so the particle is never moving clockwise.";

    $fullCirc->qa("Does the entire unit circle get traced " .
    	          "by this parameterization?", "yes");
    $fullCirc->makeLast("no");

    # a point where the function goes through (1,0)
    $solnPt = Real("1/$a");

} else {
    $domain = "\((-\infty,\infty)\)";
    $a = random(1,5,1);
    $g = Compute("$a*t^2 - t")->reduce();
    $twoa = 2*$a;
    $gp = Compute("$twoa*t - 1");

    $gmin = $g->eval(t=>0.5);
    $grange = Compute("($gmin, Inf)");

    # then g' > 0 when 2at - 1 > 0, or t > 1/2a
    $clock = List("(-Inf,1/(2*$a))");
    $cclock = List("(1/(2*$a),Inf)");

    $clockans = "Conversely, the circle is traced out clockwise when " .
      "\(t\) is in \(" . $clock->TeX . "\).";

    $fullCirc->qa("Does the entire unit circle get traced " .
    	          "by this parameterization?", "yes");
    $fullCirc->makeLast("no");

    # a point where the function goes through (1,0)
    $solnPt = Real("0");
}

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider the parameterization of the unit circle given by 
\( x =\cos( $g )\), \( y = \sin( $g )\) for \( t \) in $domain.
Describe in words and sketch how the circle is traced out, and use this to
answer the following questions.

$PAR
$BBOLD(a)$EBOLD
When is the parameterization tracing the circle out in a
${BITALIC}clockwise$EITALIC direction?
$BR
\{ ans_rule( 35 ) \}
$BR
${BITALIC}(Give your answer as a comma-separated list of intervals, for
example, (0,1), (3,Inf)).  Enter the word ${BBOLD}None$EBOLD if there are
no such intervals.$EITALIC
$PAR

$BBOLD(b)$EBOLD
When is the parameterization tracing the circle out in a
${BITALIC}counter-clockwise$EITALIC direction?
$BR
\{ ans_rule( 35 ) \}
$BR
${BITALIC}(Give your answer as a comma-separated list of intervals, for
example, (0,1), (3,Inf)).  Enter the word ${BBOLD}None$EBOLD if there are
no such intervals.$EITALIC
$PAR

$BBOLD(c)$EBOLD
\{ $fullCirc->print_q() \}
\{ $fullCirc->print_a() \}

$PAR
$BBOLD(d)$EBOLD
Give a time \(t\) at which the point being traced out on the circle is at
\((1,0)\):
$BR
\(t = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($clock->cmp() );
ANS($cclock->cmp() );
ANS(radio_cmp( $fullCirc->correct_ans() ) );
ANS($solnPt->cmp( checker=>sub {
		       my ( $correct, $student, $ans ) = @_;
		       my $gs = $g->eval(t=>$student);
		       my $gc = $g->eval(t=>$correct);
		       return ( cos($gs) == cos($gc) &&
				sin($gs) == sin($gc) ) ? 1 : 0;
		   } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Let \(g(t) = $g\).  Then \(g'(t) = $gp\).  The particle is moving
counterclockwise when \(g'(t)>0\), that is, when \(t\) is in \($cclock\).
$clockans

$PAR
The full circle is traced out if \(g(t) = $g\) produces values over one 
period of the functions \(sin(t)\) and \(cos(t)\).  In this case \(g(t)\)
produces all values \($grange\), so the full circle must be produced.

$PAR
To find a point where the parameterization is tracing out the point
\((1,0)\), we are looking for a \(t\) such that \(g(t) = 0\).  Such
a value is \(t = $solnPt\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
