# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## DBsubject(Calculus - single variable)
## DBchapter(Parametric)
## DBsection(Tangents, velocity, and speed)
## Institution(UMich)
## Author(Gavin LaRose)
## MLT(tangentLine_ParametricForm)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('4.8')
## Section2('4.8')
## Section3('4.8')
## Section4('4.8')
## Problem1('15')
## Problem2('15')
## Problem3('34')
## Problem4('34')
## KEYWORDS('parametric equation')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "hhAdditionalMacros.pl",
  "PGcourse.pl"
);

# Context("Vector");
Context("Numeric");
Context()->variables->add(t=>"Real");
$showPartialCorrectAnswers = 1;

$a = random(1,8,1);
$b = random(1,8,1);
$c = random(0,7,1);
$xfunc = Compute("$a*t^2 + $c")->reduce();
$yfunc = Compute("$b*t^3");

$t0 = random(1,3,1);

$x0 = $xfunc->eval(t=>$t0);
$y0 = $yfunc->eval(t=>$t0);
$xd = $xfunc->D('t');
$yd = $yfunc->D('t');
$xsl = $xd->eval(t=>$t0);
$ysl = $yd->eval(t=>$t0);
$xline = Compute("$x0 + $xsl*(t - $t0)");
$yline = Compute("$y0 + $ysl*(t - $t0)");

# $vfunc = Formula("<$xfunc, $yfunc>");
# $point = $vfunc->eval(t=>$t0);
# $dvec  = $vfunc->D('t')->eval(t=>$t0);
# $line  = ParametricLine("$point + $dvec*t");

$mult = MultiAnswer($xline, $yline)->with(
    singleResult => 0,
    checker => sub {
	my ($correct, $student, $self) = @_;
	my ($xstu, $ystu) = @{$student};
	my ($xcor, $ycor) = @{$correct};

	$xstu = Formula($xstu);
	$ystu = Formula($ystu);

	## if we got exactly the same functions, we're ok
	if ( $xstu == $xcor && $ystu == $ycor ) {
	    return [1,1];
	} else {
	    ## otherwise, the two functions must be linear, the
	    ##    slope must be correct, and they must pass through
	    ##    the point in question
	    $xd_stu = $xstu->D('t');
	    $yd_stu = $ystu->D('t');

	    if ( $xd_stu->isConstant && $yd_stu->isConstant ) {
		$xpoint = $xcor->eval(t=>$t0);
		$ypoint = $ycor->eval(t=>$t0);

		$xsl_stu = $xd_stu->eval(t=>0);
		$xpt_stu = $xstu->eval(t=>0);
		$ysl_stu = $yd_stu->eval(t=>0);
		$ypt_stu = $ystu->eval(t=>0);

		if ( $xsl_stu != 0 ) {
		    $sl_stu = $ysl_stu/$xsl_stu;
		    if ( $sl_stu == $ysl/$xsl ) {
			## what t value do we need to get to the point?
			$stu_t = ($xpoint - $xpt_stu)/$xsl_stu;

			## and does this give the y point in the y equation?
			if ( $ystu->eval(t=>$stu_t) == $ypoint ) {
			    return [1,1];
			}
		    }
		    ## if these didn't match, then we just check to see
		    ##    if one or the other of these matches
		    if ( $xstu == $xcor ) {
			return [1,0];
		    } elsif ( $ystu == $ycor ) {
			return [0,1];
		    } else {
			return [0,0];
		    }
		} else {
		    ## we know that the x-equation is wrong; check
		    ##    to see if the y-equation is right (that is,
		    ##    exactly what we have)
		    if ( $ystu == $ycor ) {
			return [0,1];
		    } else {
			return [0,0];
		    }
		}

	    } elsif ( $xd_stu->isConstant && $xstu == $xcor ) {
		return [1,0];
	    } elsif ( $yd_stu->isConstant && $ystu == $ycor ) {
		return [0,1];
	    } else {
		return [0,0];
	    }
	}
    }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find parametric equations for the tangent line at 
\( t = $t0 \) for the motion of a particle given by 
\( x(t) = $xfunc\), \( y(t) = $yfunc\).
$PAR
For the line,$BR
\( x(t) = \) \{ $mult->ans_rule(35) \}$BR
\( y(t) = \) \{ $mult->ans_rule(35) \}$BR
${BITALIC}(Note that because the correctness of
a parametrically described line depends on both
\(x(t)\) and \(y(t)\), both of your answers may
be marked incorrect if there is an error in one
of them.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $mult->cmp() );
# ANS( $line->cmp() );

$m = reduced_frac($ysl,$xsl);

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

At \(t=$t0\), the position is \(($x0,$y0)\).
The velocity in the \(x\)-direction is \($xd\bigg|_{t=$t0} = $xsl\), and the
velocity in the \(y\)-direction is \($yd\bigg|_{t=$t0} = $ysl\).
So we want the line going through 
the point \(($x0,$y0)\), with the given
\(x\)- and \(y\)-velocities:
\[
x = $x0 + $xsl(t-$t0),\quad y = $y0 + $ysl(t-$t0).
\]

In general, any parametric line that goes through the point and
has the slope \({dy\over dx} = {$ysl\over $xsl} = $m\) in the
\(xy\)-plane will also be correct.

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
