# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(UMich)
## Author(Gavin LaRose)
## MLT(pc_lines_2d)
## Level(1)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('4.8')
## Section2('4.8')
## Section3('4.8')
## Section4('4.8')
## Problem1('17')
## Problem2('17')
## Problem3('6')
## Problem4('6')
## KEYWORDS('parametric equations')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiPart.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add(t=>"Real");

$x0 = random(-8,-1,1);
$y0 = random(-8,-1,1);
$xans = Compute($x0);
$yans = Compute("t");

## this is completely unnecessary.  but I was copying from Q15
$mult = MultiAnswer($xans, $yans)->with(
    singleResult => 0,
    checker => sub {
	my ($correct, $student, $self) = @_;
	my ($xstu, $ystu) = @{$student};
	my ($xcor, $ycor) = @{$correct};

	$xstu = Formula($xstu);
	$ystu = Formula($ystu);

	## if we got exactly the same functions, we're ok
	if ( $xstu == $xcor && $ystu == $ycor ) {
	    return [1,1];
	} else {
	    ## otherwise, x must have the right value and y must
	    ##    be linear
	    $xd_stu = $xstu->D('t');
	    $yd_stu = $ystu->D('t');

	    $xval = 0;
	    $yval = 0;

	    $xval = 1 if ( $xd_stu == 0 && $xstu == Formula("$x0") );
	    $yval = 1 if ( $yd_stu->isConstant && $yd_stu != 0 );

	    return [$xval, $yval];
	}
    }
);


TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Write a parameterization for the curve in the \( xy\)-plane that
is a vertical line through the point \( ($x0,$y0) \).

$PAR
\( x(t) = \) \{ $mult->ans_rule(35) \} $BR
\( y(t) = \) \{ $mult->ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($mult->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

One possible answer is \(x=$x0, y=t\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
