# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## DBsubject(Calculus - single variable)
## DBchapter(Parametric)
## DBsection(Curves)
## Institution(UMich)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('4.8')
## Section2('4.8')
## Section3('4.8')
## Section4('4.8')
## Problem1('27')
## Problem2('36')
## Problem3('38')
## Problem4('38')
## KEYWORDS('parametric equations')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Point");
Context()->variables->add(t=>"Real");
Context()->strings->add("Y"=>{},"N"=>{},"yes"=>{alias=>"Y"},"no"=>{alias=>"N"});
$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(2,8,1);
$c = random(2,8,1);
$d = random(2,8,1);
$linex = Formula("$a + $b*t");
$liney = Formula("$c + $d*t");

$t0 = random(-5,-1,1);
$t1 = $t0 + random(1,8,1);
$t1++ if ( $t1 == 0 );
$t2 = $t1 + random(1,3,1);
$t2 = random(1,3,1) if ( $t2 < 0 );
$point1 = Point($linex->eval(t=>$t0),$liney->eval(t=>$t0));
$point2 = Point($linex->eval(t=>$t1), $liney->eval(t=>$t1));
$point3 = Point($linex->eval(t=>$t2), $liney->eval(t=>$t2));
$ans1 = String("N");
$ans2 = ( $t1 > 0 ) ? String("Y") : String("N");
$ans3 = String("Y");

$ans4 = Point($linex->eval(t=>0), $liney->eval(t=>0));

$p1 = random(-5,-1,1);
$p2 = $p1 + random(1,3,1);

$ans5 = Point($linex->eval(t=>$p1), $liney->eval(t=>$p1));
$ans6 = Point($linex->eval(t=>$p2), $liney->eval(t=>$p2));

$whicha = list_random("x","y");
$whichd = list_random(0,1);
if ( $whicha eq 'x' ) {
    if ( $whichd ) {
	$direction = "above";
	$ans7 = Interval("(-$c/$d,Inf)");
    } else {
	$direction = "below";
	$ans7 = Interval("(-Inf,-$c/$d)");
    }
} else {
    if ( $whichd ) {
	$direction = "to the left of";
	$ans7 = Interval("(-Inf,-$a/$b)");
    } else {
	$direction = "to the right of";
	$ans7 = Interval("(-$a/$b,Inf)");
    }
}

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

A line is parameterized by \(  x= $linex \) and \( y = $liney \).

$PAR
${BBOLD}(a)$EBOLD
Which of the following points are on the section of the line
obtained by restricting \(t\) to nonnegative numbers (for each, enter
Y if the point is on the section, and N if not)?
$BR
\( $point1 \) : \{ ans_rule(3) \} $BR
\( $point2 \) : \{ ans_rule(3) \} $BR
\( $point3 \) : \{ ans_rule(3) \} $BR
Then, give one more point that ${BITALIC}is$EITALIC on the section
of the line obtained by this restriction: \{ ans_rule(7) \}

$PAR
${BBOLD}(b)$EBOLD
What are the endpoints of the line segment obtained by restricting
\( t \) to \( $p1\le t\le $p2\)?
$BR
left endpoint : \{ ans_rule(7) \} $BR
right endpoint : \{ ans_rule(7) \}

$PAR
${BBOLD}(c)$EBOLD
How should \( t \) be restricted to give the part of the line 
$direction the \($whicha\)-axis (give your answer as an interval for 
\(t\), for example, (3,8) or [-2,Inf))?
$BR
\(t\) must be in : \{ ans_rule(7) \}

END_TEXT
Context()->normalStrings;

ANS( $ans1->cmp() );
ANS( $ans2->cmp() );
ANS( $ans3->cmp() );
ANS( $ans4->cmp( checker=>sub {
	my ( $correct, $student, $ansHash ) = @_;
	if ( $student == $point3 || ($ans2 eq "Y" && $student == $point2) ) {
	    Value->Error("This is the same as one of the given points.");
	    # or $ans->{ans_message} = $ans->{error_message} = "message";
	} else {
	    my @stu_vals = $student->value;
	    my @cor_vals = $correct->value;
	    my $stu_t = ($stu_vals[0] - $a)/$b;
	    if ( $stu_t >= 0 && $liney->eval(t=>$stu_t) == $stu_vals[1] ) {
		return 1;
	    } else {
		return 0;
	    }
	} } ) );
ANS( $ans5->cmp() );
ANS( $ans6->cmp() );
Context("Interval");
ANS( $ans7->cmp() );
Context("Point");

$ispoint = ( $ans2 eq 'Y' ) ? "is" : "also is not";
$whichLine = ( $whicha eq 'x' ) ? $liney : $linex;
$inequ = ( $direction =~ /above/ or $direction =~ /right/ ) ? ">" : "<";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
If we have \(t\ge 0\), then \(x\ge $a\) and \(y\ge $c\), so we're 
looking at the part of the line to the right of and above \(($a,$c)\).
Thus \($point1\) is not on this section of the line, \($point2\) $ispoint,
and \($point3\) is.  Any additional points must have \(x\ge$a\) and
\(y\ge$c\), and must lie on the line.

$PAR
${BBOLD}(b)$EBOLD
We find the endpoints of the line segment obtained by restricting \(t\)
to \($p1\le t\le $p2\) we can plug \(t=$p1\) and \(t=$p2\) into the
equations of the line, getting \($ans5\) and \($ans6\).

$PAR
${BBOLD}(c)$EBOLD
To get the part of the line $direction the \($whicha\)-axis, we need
\($whichLine $inequ 0\).  This gives the interval \($ans7\).

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
