## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Distinct real eigenvalues)
## Date(03/19/2011)
## Institution(TCNJ and Fort Lewis College)
## Author(Anonymous and Paul Pearson)
## MLT(SDE-DRE-IVP)
## Level(4)
## MO(1)

## Original: Library/TCNJ/TCNJ_Eigenvalues/problem10.pg


############################
#  Initialization
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "MathObjects.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric")->variables->are(t=>"Real");

$a = new Matrix(2,2);
# define matrix
        $a11 = random(-1,1,2);
        $a12 = non_zero_random(-3,3,1);
        $mult = non_zero_random(-2,2,1);
        $a21 = $mult * $a11;
        $b1 = random(-1,1,2);
        $a22 = $mult * $a12 + $b1;

        $a->assign(1,1, $a11 );
        $a->assign(1,2, $a12 );
        $a->assign(2,1, $a21 );
        $a->assign(2,2, $a22 );
        $a_lr = $a->decompose_LR();
        $a_det = $a11 * $b1;

# define inverse matrix
	$b = $a_lr->invert_LR();
# define eigenvalues
    $e = new Matrix(2,2);
    $e->one();
    $e->assign(1,1, non_zero_random(-8,8,1));
    $e->assign(2,2, non_zero_random(-8,8,1));
# make sure that the eigenvalues are distinct
	while ($e->element(1,1) == $e->element(2,2)) { $e->assign(2,2, non_zero_random(-8,8,1)); }
    $e1 = $e->element(1,1);
    $e2 = $e->element(2,2);

# define final matrix
	$matrix = $a * $e *$b;
	$matrix_lr = $matrix->decompose_LR();
	$matrix_det = $matrix_lr->det_LR();


$m11 = $matrix->element(1,1);
$m12 = $matrix->element(1,2);
$m21 = $matrix->element(2,1);
$m22 = $matrix->element(2,2);

$evects = MultiAnswer($a11, $a21, $a12, $a22)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='0'>
    <tr>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  "\left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack, 
    \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
        my @s = @{$student};
        my $flag  = 0;
        my $flag1 = 0;
        my $flag2 = 0;

        if ( $c[0]->typeMatch($s[0]) && $c[1]->typeMatch($s[1]) ) {
           my $v1 = new Matrix(2, 1);
           if ( $s[0] != 0 || $s[1] != 0 ) {
              $v1-> assign(1,1, $s[0]);
              $v1-> assign(2,1, $s[1]);
              if ( $matrix * $v1 == $e1 * $v1 ) {
                $flag1 = 1;
              }
           }        
        }

        if ( $c[2]->typeMatch($s[2]) && $c[3]->typeMatch($s[3]) ) {
           my $v2 = new Matrix(2, 1);
           if ( $s[2] != 0 || $s[3] != 0 ) {
              $v2-> assign(1,1, $s[2]);
              $v2-> assign(2,1, $s[3]);
              if ( $matrix * $v2 == $e2 * $v2 ) {
                $flag2 = 1;
              }
           }        
        }

        if ($flag1==1 && $flag2 ==1) {
           $flag=1;
        }else{
           if ($flag1 == 1) {
              Value::Error("The vector on the left is correct");
           } elsif ($flag2 == 1) {
              Value::Error("The vector on the right is correct");
           }else{
              Value::Error("Neither of the eigenvectors is correct");
           }
        }
        return $flag; 
       
    }
);



#  initial conditions
$C1 = non_zero_random(-3,3,1);
$C2 = non_zero_random(-3,3,1);

$x1init = Real("($C1)*($a11)+($C2)*($a12)");
$x2init = Real("($C1)*($a21)+($C2)*($a22)");


##############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}
$ITEM
The matrix 
\( 
\displaystyle
\left\lbrack
\begin{array}{rr}
$m11 & $m12 \\
$m21 & $m22
\end{array}
\right\rbrack
\)
$SPACE
has eigenvalues \( \lambda_1 = $e1 \) and \( \lambda_2 = $e2 \). 
Find eigenvectors corresponding to these eigenvalues.
$BCENTER
\{ 
mbox(
[
'\( \boldsymbol{\vec{v}_1} =  \) ', 
display_matrix([[$evects->ans_rule(5)],[$evects->ans_rule(5)]]),
' and \( \boldsymbol{\vec{v}_2} =  \) ', 
display_matrix([[$evects->ans_rule(5)],[$evects->ans_rule(5)]]),
],
allowbreaks=>'yes'
) 
\}
$ECENTER

$ITEMSEP
$ITEM
Find the solution to the linear system of 
differential equations 
\( 
\displaystyle
\left\lbrack
\begin{array}{r}
x_1^{\,\prime} \\
x_2^{\,\prime}
\end{array}
\right\rbrack
=
\left\lbrack
\begin{array}{rr}
$m11 & $m12 \\
$m21 & $m22
\end{array}
\right\rbrack
\left\lbrack
\begin{array}{r}
x_1 \\
x_2
\end{array}
\right\rbrack
\)
$SPACE
satisfying the initial conditions
\( 
\displaystyle
\left\lbrack
\begin{array}{r}
x_1(0) 
\\
x_2(0)
\end{array}
\right\rbrack
=
\left\lbrack
\begin{array}{r}
$x1init 
\\
$x2init
\end{array}
\right\rbrack
\).
$BR
$BR
\( x_1(t) = \) \{ ans_rule(40) \} 
$BR
\( x_2(t) = \) \{ ans_rule(40) \} 

\{ EndList('OL') \}
END_TEXT
Context()->texStrings;


##############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $evects->cmp() );

ANS( Compute("$C1 * $a11 e^($e1 t) + $C2 * $a12 e^($e2 t)")->cmp() );
ANS( Compute("$C1 * $a21 e^($e1 t) + $C2 * $a22 e^($e2 t)")->cmp() );

COMMENT('Randomized with distinct real eigenvalues and integral answers.');

ENDDOCUMENT();
