## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Distinct real eigenvalues)
## Date(01/31/2012)
## Institution(Rochester and Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-DRE-evec_only)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.4')
## Problem1('13')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());


#####################################
#  Setup

Context("Numeric");

# dilation
$a = non_zero_random(-3,3,1);

$i = random(0,1,1);

if ($i == 0) {

$A11 =  1 * $a;  $A12 =  1 * $a;
$A21 = -4 * $a;  $A22 =  6 * $a;

$matrix = new Matrix(2,2);
$matrix->assign(1,1, $A11 );
$matrix->assign(1,2, $A12 );
$matrix->assign(2,1, $A21 );
$matrix->assign(2,2, $A22 );


$lambda1 =  2 * $a;
$v11 = Real("1");
$v21 = Real("1");

$lambda2 = 5 * $a;
$v12 = Real("1");
$v22 = Real("4");

} else {

# the "rotated" version

$A11 =  6 * $a;  $A12 = -1 * $a;
$A21 =  4 * $a;  $A22 =  1 * $a;

$matrix = new Matrix(2,2);
$matrix->assign(1,1, $A11 );
$matrix->assign(1,2, $A12 );
$matrix->assign(2,1, $A21 );
$matrix->assign(2,2, $A22 );


$lambda1 =  2 * $a;
$v11 = Real("1");
$v21 = Real("4");

$lambda2 = 5 * $a;
$v12 = Real("1");
$v22 = Real("1");

}


$multians = MultiAnswer($v11, $v21, $v12, $v22)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='0'>
    <tr>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  "\left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack, 
    \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $self ) = @_;
        my @c = @{$correct};
	my @s = @{$student};
        my $flag  = 0;
        my $flag1 = 0;
        my $flag2 = 0;

        if ($c[0]->typeMatch($s[0]) && $c[1]->typeMatch($s[1])) {
           my $v1 = new Matrix(2, 1);
           if($s[0] != 0 || $s[1] !=0){
             $v1-> assign(1,1, $s[0]);
             $v1-> assign(2,1, $s[1]);
             if( $matrix * $v1 == $lambda1 * $v1){
                $flag1=1;
             }
           }
        }

        if($c[2]->typeMatch($s[2]) && $c[3]->typeMatch($s[3])){
           my $v2 = new Matrix(2, 1);
           if($s[2] != 0 || $s[3] !=0){
             $v2-> assign(1,1, $s[2]);
             $v2-> assign(2,1, $s[3]);
             if( $matrix * $v2 == $lambda2 * $v2){
               $flag2=1;
             }
           }
        }

        if($flag1==1 && $flag2 ==1){
           $flag=1;
        }else{
           if($flag1 == 1){
              Value::Error("The vector on the left is correct");
           } elsif ($flag2 == 1) {
              Value::Error("The vector on the right is correct");
           }else{
              Value::Error("Neither of the vectors is correct");
           }
        }
        return $flag; 
       
    }
);



#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
The matrix 
\( 
\displaystyle
\left\lbrack
\begin{array}{rr}
$A11 & $A12 \\
$A21 & $A22
\end{array}
\right\rbrack
\)
$SPACE
has eigenvalues \( \lambda_1 = $lambda1 \) and \( \lambda_2 = $lambda2 \). 
Find eigenvectors corresponding to these eigenvalues.
$BCENTER
\{ 
mbox(
[
'\( \boldsymbol{\vec{v}_1} =  \) ', 
display_matrix([[$multians->ans_rule(5)],[$multians->ans_rule(5)]]),
' and \( \boldsymbol{\vec{v}_2} =  \) ', 
display_matrix([[$multians->ans_rule(5)],[$multians->ans_rule(5)]]),
],
allowbreaks=>'yes'
) 
\}
$ECENTER
END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

COMMENT('Randomized with distinct real eigenvalues and integral answers.');

ENDDOCUMENT(); 
