# DESCRIPTION
# Homework problem from A First Course in Probability, Ross;
# 9th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('probability,random variables,discrete random variables')
## Tagged by glr 13/09/25

## DBsubject('Probability')
## DBchapter('Random variables')
## DBsection('Expectation')
## Date('08/12/16')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('A First Course in Probability')
## EditionText1('9')
## AuthorText1('Sheldon Ross')
## Section1('7.2')
## Problem1('10')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGnumericalmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");

# normal_prob(a, b, mean=>0, deviation=>1);
$n = random(2,4,1);
$p = random(0.2, 0.5, 0.1);
$pans1 = Compute( $p );
$ex = $n*$p;
$pans1rule = $pans1->ans_rule();

# first set up some probabilities that work
$pans2 = Compute( 0 );
$pans2rule = $pans1->ans_rule();
@pvals = ();
push( @pvals, Compute( "1 - 2*$p" ) );
for ( my $i=1; $i<$n; $i++ ) {
  push( @pvals, Compute( "2*$p/($n-1)" ) );
}
push( @pvals );

# then we can set up the multianswer for the probabilities
$mp = MultiAnswer( @pvals )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ansHash ) = @_;
	my @c = @$c;
	my @s = @$s;

	my $tprob = 0;
	my $eval  = 0;
	for ( my $i=0; $i<@s; $i++ ) {
	    $tprob += $s[$i];
	    $eval  += $i*$s[$i];
	}
	if ( $tprob != 1 || Compute($ex) != $eval ) {
	    return [ (0)x$n ];
	} else {
	    return [ (1)x$n ];
	}
});

# now set up the text that we need
$otherprob = '';
for ( my $i=0; $i<$n; $i++ ) {
    $otherprob .= "\( P\lbrace X = $i\rbrace = \) " . $mp->ans_rule() .
        " $BR";
}

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider $n trials, each having the same probability of success.
Let \(X\) denote the total number of successes in these trials.
If \(E[X] = $ex\), find each of the following.
$PAR
${BBOLD}(a)$EBOLD
The largest possible value of \(P\lbrace X = $n\rbrace\):$BR
\(P\lbrace X = $n\rbrace \le \) $pans1rule
$PAR
${BBOLD}(b)$EBOLD
The smallest possible value of \(P\lbrace X = $n\rbrace\):$BR
\(P\lbrace X = $n\rbrace \ge \) $pans2rule $BR
In this case, give possible values for the remaining
probabilities:$BR
$otherprob

END_TEXT
Context()->normalStrings;

ANS( $pans1->cmp() );
ANS( $pans2->cmp() );
ANS( $mp->cmp() );

$p1 = 1 - $pvals[0];
$nm1 = $n - 1;
$nsum = ($nm1*$n)/2;

Context()->texStrings;
BEGIN_SOLUTION

Note that the $n trials are ${BBOLD}not$EBOLD given to be
independent, so this is not a binomial distribution.  In
either case
\[ E[X] = 0\cdot P\lbrace X = 0 \rbrace + \cdots +
$n \cdot P\lbrace X = $n\rbrace. \]
$PAR

${BBOLD}(a)$EBOLD
The maximum value for \(P\lbrace X = $n\rbrace\) will occur
when all probabilities \(P\lbrace X = j\rbrace\) (for \(0 < j < $n\))
are zero.  Then
\[ E[X] = 0\cdot P\lbrace X = 0 \rbrace +
$n \cdot P\lbrace X = $n\rbrace = $n P\lbrace X = $n\rbrace = $ex, \]
and we see that \(P\lbrace X = $n\rbrace = $pvals[0]\).  (In this
case, \(P\lbrace X = 0\rbrace = $p1\).)

$PAR
${BBOLD}(b)$EBOLD
Reasoning similarly, we can take \(P\lbrace X = $n\rbrace = 0\) provided
the remaining probabilities are chosen so that they sum to one and give
the required expected value.  There are many possible solutions to this;
if we assign equal probabilities to each of \(X = j\), \(0 < j < $n\),
we have
\[ \sum\limits_{j=1}^{$nm1} j\cdot p = $nsum p = $ex, \]
so that \(p = $pvals[1]\).  Then \(P\lbrace X = 0\rbrace = $pvals[0]\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
