
DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "niceTables.pl",
  "AnswerFormatHelp.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl",
);

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup
############################################################
Context()->noreduce('(-x)+y', '(-x)-y');
$f = RandomVariableName(type=>'function');
do {
   $a = non_zero_random(-4,4,1);
   $b = non_zero_random(-19,19,1);
   $c = non_zero_random(-5,5,1);
} until (0 < -$b/(2*$a)  and  -$b/(2*$a) < 5);
$formula = Compute("$a x^2 + $b*x + $c")->reduce;
@defaultAnswers = ();
for my $i (0..6) {
  $x[$i] = $i-2;
  $y[$i] = $formula->eval(x=>$x[$i]);
  push(@defaultAnswers,$x[$i]);
  push(@defaultAnswers,$y[$i]);
}
$multians = MultiAnswer(@defaultAnswers)->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my @sxy = @{$student};
        my @xs, @ys;
       for my $k (0..6)
         {$xs[$k] = $sxy[2*$k];
          $ys[$k] = $sxy[2*$k+1];}
       #
       #  Loop though the student answers
       ##
       for ($i = 0; $i <= 6; $i++) {
         my $ith = Value::List->NameForNumber($i+1);
         #
         #  Check that the student's answers are Reals
         #
         Value->Error("Your $ith x-value is not a number") if ($xs[$i]->type ne "Number");
         Value->Error("Your $ith y-value is not a number") if ($ys[$i]->type ne "Number");
         #
         #  Check that the point hasn't been given before
         #
         for ($j = 0, $used = 0; $j < $i; $j++) {
           my $jth = Value::List->NameForNumber($j+1);
           if ($xs[$j] == $xs[$i] and $ys[$j] == $ys[$i]) {
             Value->Error("Your $ith row in the table is the same as your $jth row") unless $ansHash->{isPreview};
             $used = 1; last;
           }
         }
         #
         #  If not already used, check that it satisfies the equation
         #    and increase the score if so.
         #
         if (!$used) {
           Value->Error("Your $ith row is not correct") unless $ansHash->{isPreview} or $ys[$i] == $formula->eval(x=>$xs[$i]);
             }
           }
         return 1;
    }
);

for my $i (0..6) {
Context()->flags->set(reduceConstants=>0,reduceConstantFunctions=>0);
$rs[$i] = $formula->substitute(x=>$x[$i]);
}

############################################################
# Body
############################################################

BEGIN_PGML
Use technology to make a table of values for the function.

[`[$f](x)=[$formula]`]

>> [@DataTable(
  [
    [[PGML('[`x`]'),
      midrule => '1',
      rowcss  => 'border-bottom: 2px solid; ',],
     [PGML('[`[$f](x)`]'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
    [[PGML('[________]{$multians}'),],
     [PGML('[________]{$multians}'),],
    ],
  ],
  align => 'cc',
  center => 0,
);@]* <<

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION

Go to desmos.com, click on "Start Graphing" button. In the first text box, enter [`[$f](x)=[$formula]`].

Then, above the text box (which has the function now), click on the gear icon. If you hover the cursor at the gear icon, it should say "Edit List".

Once you click on the gear icon, click on the table icon. If you hover the cursor at the table icon, it should say "Convert to Table".

Now the table is created with some simple [`x`]-values. You can click on the table to add more rows. The table may be completed as:

>> [@DataTable(
  [
    [[PGML('[`x`]'),
      midrule => '1',
      rowcss  => 'border-bottom: 2px solid; ',],
     [PGML('[`[$f](x)`]'),],
    ],
    [[PGML('[`[$x[0]]`]'),],
     [PGML('[`[$y[0]]`]'),],
    ],
    [[PGML('[`[$x[1]]`]'),],
     [PGML('[`[$y[1]]`]'),],
    ],
    [[PGML('[`[$x[2]]`]'),],
     [PGML('[`[$y[2]]`]'),],
    ],
    [[PGML('[`[$x[3]]`]'),],
     [PGML('[`[$y[3]]`]'),],
    ],
    [[PGML('[`[$x[4]]`]'),],
     [PGML('[`[$y[4]]`]'),],
    ],
    [[PGML('[`[$x[5]]`]'),],
     [PGML('[`[$y[5]]`]'),],
    ],
    [[PGML('[`[$x[6]]`]'),],
     [PGML('[`[$y[6]]`]'),],
    ],
  ],
  align => 'cc',
  center => 0,
);@]* <<

END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
