# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Linear Transformations')
## KEYWORDS('linear transformations, one-to-one, onto')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.1')
## Problem1('46')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserPopUp.pl",
# "compoundProblem.pl",
"parserMultiAnswer.pl",
);

# make sure we're in the context we want
Context("Matrix");
$showPartialCorrectAnswers = 0;

$tfPop1 = PopUp( [ '?', 'True', 'False' ], 'True' );
$tfPop2 = PopUp( [ '?', 'True', 'False' ], 'False' );

# $cp = new compoundProblem( parts=>2, weights=>[1,2], nextStyle=>'Forced' );

$amat1 = Matrix( [[1,0],[0,1]] );
$amat2 = Matrix( [[-1,0],[1,1]] );
$amat = new MultiAnswer( $amat1, $amat2 )->with(
    singleResult=>1,
    checker=>sub {
	my ( $cor, $stu, $ans ) = @_;
	my ( $sMat1, $sMat2 ) = @{$stu};
	# we need that the sum of the matrices to be non-invertible.
	my $sSum = $sMat1 + $sMat2;
	return $sMat1->det() && $sMat2->det() && ! $sSum->det();
    }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

${BITALIC}(Note that this is a compound problem, for which the
second part of the problem is shown after the true/false questions
in the first part are answered correctly.)$EITALIC
$PAR

${BBOLD}(a)$EBOLD
Indicate if each of the following statements is true or false:
$PAR
If \(T_1(\mathbf{x})\) and \(T_2(\mathbf{x})\) are onto linear
transformations from \(\mathbf{R}^n\) to \(\mathbf{R}^m\), then
so is \(W(\mathbf{x}) = T_1(\mathbf{x}) + T_2(\mathbf{x})\).
$PAR
\(W\) is a linear transformation \{ $tfPop1->menu() \}; and $BR
\(W\) is an onto transformation  \{ $tfPop2->menu() \}.
$PAR

END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 0;
ANS( $tfPop1->cmp() );
ANS( $tfPop2->cmp() );

$ans_hash1 = $tfPop1->cmp()->evaluate( $inputs_ref->{ANS_NUM_TO_NAME(1)} );
$ans_hash2 = $tfPop2->cmp()->evaluate( $inputs_ref->{ANS_NUM_TO_NAME(2)} );

if ( $ans_hash1->{score} == 1 && $ans_hash2->{score} == 1 ) {

Context()->texStrings;
BEGIN_TEXT

${BBOLD}(b)$EBOLD
Let \(T_1\) and \(T_2\) map from \(\mathbf{R}^2\) onto \(\mathbf{R}^2\),
and let \(T_1(\mathbf{x}) = A_1\mathbf{x}\) and
\(T_2(\mathbf{x}) = A_2\mathbf{x}\), where \(A_1\) and \(A_2\) are
\(2\times2\) matrices.  Give matrices \(A_1\) and \(A_2\) such that
\(W\) is not an onto transformation.
$BR
\( A_1 = \) \{ $amat->ans_array() \}
$BR
\( A_2 = \) \{ $amat->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $amat->cmp() );

}

Context()->texStrings;
BEGIN_SOLUTION

${BBOLD}(a)$EBOLD
\(W\) is a linear transformation, but not necessarily onto.

$PAR
${BBOLD}(b)$EBOLD
For \(W\) not to be onto, we need the sum of the matrices
to have columns which are not linearly independent.  One such
pair of matrices \(A_1 = $amat1\) and \(A_2 = $amat2\).  Note
that both of the matrices \(A_1\) and \(A_2\) must themselves
have linearly independent columns, or the original transformations
are not onto.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
