# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose, 
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Linear Transformations')
## KEYWORDS('linear transformations, R2')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.1')
## Problem1('37')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
);

# make sure we're in the context we want
Context("Matrix");

$amat = Matrix( [[ random(1,5,1), random(-5,-1,1) ],
		 [ random(1,5,1), random(1,5,1) ]] );

$in11 = -1*$amat->element(1,2);
$in12 = $amat->element(1,1);
$in1 = ColumnVector( $in11, $in12 );
$in21 = 1;
$in22 = random(1,5,1);
$in2 = ColumnVector( $in21, $in22 );

$out1 = $amat*$in1;
$out2 = $amat*$in2;

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find an example of a linear transformation
\(T:\mathbf{R}^2 \rightarrow \mathbf{R}^2\)
given by \( T(\mathbf{x}) = A\mathbf{x} \)
such that
\[
T\left( $in1 \right) = $out1 \quad\mbox{and}\quad
T\left( $in2 \right) = $out2.
\]
$PAR
\( A = \) \{ $amat->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $amat->cmp( checker=>sub {
		     my ( $cor, $stu, $ans ) = @_;
		     return $stu * $in1 == $out1 && $stu * $in2 == $out2; 
		 } ) );

$out11 = $out1->element(1,1);  $out12 = $out1->element(2,1);
$out21 = $out2->element(1,1);  $out22 = $out2->element(2,1);
$in11d = ( $in11 == 1 ) ? '' : $in11;  $in12d = ( $in12 == 1 ) ? '' : $in12;
$in21d = ( $in21 == 1 ) ? '' : $in21;  $in22d = ( $in22 == 1 ) ? '' : $in22;

Context()->texStrings;
BEGIN_SOLUTION

We need a matrix \(A\) such that \(A $in1 = $out1\) and
\(A $in2 = $out2\).  With \(A = \begin{bmatrix} a_{11} & a_{12}\\
a_{21} & a_{22}\end{bmatrix}\), this requires that
\[\begin{aligned}
    $in11d a_{11} + $in12d a_{12} &= $out11,\\
    $in11d a_{21} + $in12d a_{22} &= $out12,\\
    $in21d a_{11} + $in22d a_{12} &= $out21, \quad\mbox{and}\\
    $in21d a_{21} + $in22d a_{22} &= $out22.
  \end{aligned}\]
Solving, we get \(A = $amat\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
