# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Linear Transformations')
## KEYWORDS('linear transformations, one-to-one, onto')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.1')
## Problem1('22')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

# make sure we're in the context we want
Context("Matrix");
Context()->variables->add( a=>'Real' );
Context()->strings->add(none=>{});

# decide if we're going to be one-to-one and onto or not,
#    favoring the latter
$whichoo = random(0,3,1);
if ( $whichoo ) {
    $oneonto = 'neither one-to-one nor onto';
    $not = 'not';
    $val2 = random(1,3,1);
    $val1 = non_zero_random(-1,1)*($val2 + random(1,3,1));
    $mult = random(2,3,1);
    $amat = Matrix( [[$val1, $val2],[$mult*$val1, $mult*$val2]] );
} else {
    $oneonto = 'both one-to-one and onto';
    $not = '';
    $val2 = random(1,3,1);
    $val1 = non_zero_random(-1,1)*($val2 + random(1,3,1));
    $mult = random(2,3,1);
    $amat = Matrix( [[$val1, $val2],
		     [$mult*$val1, $mult*$val2 - 1]] );
}

# set up the pop-up
$char = PopUp( [ '?', 'one-to-one', 'onto',
		 'both one-to-one and onto',
		 'neither one-to-one nor onto' ],
	       $oneonto );

# then, a new matrix
$amat2 = Matrix( [[$val1, $val2],[Compute("a"), $mult*$val2]] );
$mult2 = -1*random(2,3,1);
# possible values for a are
$aval1 = Compute( $mult*$val1 - 1 );
$aval2 = Compute( $mult*$val1 - 1 );
$aval3 = Compute( $mult*$val1 - 1 );
$aval4 = Compute( $mult*$val1 );
# the first column vector is
$v1 = ColumnVector( $val2, $mult*$val1 );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Let \(T(\mathbf{x}) = A\mathbf{x}\).
$PAR
${BBOLD}(a)$EBOLD
If \(A = $amat\), \(T\) is \{ $char->menu() \}.

$PAR
${BBOLD}(b)$EBOLD
Now consider \( A = $amat2 \), where \(a\) is a constant.
Give a value of \(a\) that makes \(T\) each of the following
${BITALIC}(if it is not possible to find such a value of \(a\),
enter ${BBOLD}none$EBOLD)$EITALIC:$BR
one-to-one:          \( a = \) \{ $aval1->ans_rule() \} $BR
onto:                \( a = \) \{ $aval2->ans_rule() \} $BR
one-to-one and onto: \( a = \) \{ $aval3->ans_rule() \} $BR
neither one-to-one nor onto: \( a = \) \{ $aval4->ans_rule() \}

END_TEXT
Context()->normalStrings;

ANS( $char->cmp() );
ANS( $aval1->cmp( checker=>sub {
		      my ($cor, $stu, $ans) = @_;
		      my $v2 = ColumnVector( $val2, $stu );
		      return (! $v1->isParallel($v2));
		  } ) );
ANS( $aval2->cmp( checker=>sub {
		      my ($cor, $stu, $ans) = @_;
		      my $v2 = ColumnVector( $val2, $stu );
		      return (! $v1->isParallel($v2));
		  } ) );
ANS( $aval3->cmp( checker=>sub {
		      my ($cor, $stu, $ans) = @_;
		      my $v2 = ColumnVector( $val2, $stu );
		      return (! $v1->isParallel($v2));
		  } ) );
ANS( $aval4->cmp() );

Context()->texStrings;
BEGIN_SOLUTION

${BBOLD}(a)$EBOLD
Because the matrix \(A\) is \(2\times2\) we know that the
transformation is one-to-one and onto if the columns of
\(A\) are linearly independent.  For \(A = $amat\), the
columns are $not linearly independent, so it is $oneonto.
$PAR

${BBOLD}(b)$EBOLD
In each of the first three cases we need to find a value for
\(a\) such that the columns of \(A = $amat2\) are linearly
independent; there are many possible values, one of which
is \(a = $aval1\).  For the last, we need to make the
columns linearly dependent, which we can do by making them
a constant multiple of each other, so take \(a = $aval4\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
