# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Linear Transformations')
## KEYWORDS('linear transformations')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.1')
## Problem1('34')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
);

# make sure we're in the context we want
Context("Matrix");

$amat = Matrix( [[ list_random(0,0,1), random(1,5,1) ],
		 [ list_random(0,0,1), random(1,5,1) ],
		 [ list_random(0,0,1), random(1,5,1) ]] );
$inVec = ColumnVector( list_random( 0, 0, 1 ), list_random( 1, 1, 2 ) );
$outVec = $amat * $inVec;

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find an example of a linear transformation
\(T:\mathbf{R}^2 \rightarrow \mathbf{R}^3\)
given by \( T(\mathbf{x}) = A\mathbf{x} \)
such that
\[
T\left( $inVec \right) = $outVec.
\]
$PAR
\( A = \) \{ $amat->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $amat->cmp( checker=>sub {
		    my ( $cor, $stu, $ans ) = @_;
		    return $stu * $inVec == $outVec; } ) );

Context()->texStrings;
BEGIN_SOLUTION

We need any matrix \(A\) such that \(A $inVec = $outVec\).  One
such is \(A = $amat\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
