# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Matrix Operations')
## KEYWORDS('matrix multiplication')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.2')
## Problem1('37')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);

# make sure we're in the context we want
Context("Matrix");

# create some reference matrices
$amat = Matrix( [ [ 1, 2 ], [ 2, 1 ] ] );
$bmat = Matrix( [ [ 2, 1 ], [ 1, 2 ] ] );
$cmat = Matrix( [ [ 1, 1 ], [ 1, 1 ] ] );
$zero = Matrix( [ [ 0, 0 ], [ 0, 0 ] ] );

$abacChecker = MultiAnswer( $amat, $bmat, $cmat )->with(
    singleResult => 1,
    checker => sub {
	my ( $cor, $stu, $ans ) = @_;
	@c = @{$cor};
	@s = @{$stu};
	return $zero != $s[0] && $zero != $s[1] && $zero != $s[2] &&
	    $s[0] != $s[1] && $s[0]*$s[2] == $s[1]*$s[2];
} );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find three nonzero \(2\times 2\) matrices \(A\), \(B\) and \(C\)
where \(A\ne B\) but \(AC = BC\).  Note that this demonstrates that
we cannot cancel terms from matrix multiples whose products are 
equal.
$BR
\( A = \) \{ $abacChecker->ans_array() \} $BR
\( B = \) \{ $abacChecker->ans_array() \} $BR
\( C = \) \{ $abacChecker->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $abacChecker->cmp() );

$ac = $amat*$cmat;

Context()->texStrings;
BEGIN_SOLUTION

There are many solutions; one is
\[ A = $amat, B = $bmat, C = $cmat; \]
then
\[ AC = BC = $ac. \]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
