# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Matrix Operations')
## KEYWORDS('matrix algebra')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.2')
## Problem1('40')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
# "compoundProblem.pl",
"parserPopUp.pl",
"parserMultiAnswer.pl",
);

# make sure we're in the context we want
Context("Matrix");

# we ask if these are diagonal for diagonal matrices A & B
@qlist = ( 'A - B', 'A + B', 'AB', 'BA',
	   'AB^T',  'A^TB' );
$item1 = list_random(0,1);
$item2 = $item1 + random(2,4,1);
# $tfpop = PopUp( [ '?', 'true', 'false' ], ($item2 > 3) ? 'false' : 'true' );
$tfpop = PopUp( [ '?', 'true', 'false' ], 'true' );

# $cp = new compoundProblem( parts=>2, weights=>[1,3], nextStyle=>'Forced' );

$amat = Matrix( [ [ 2, 1 ], [ 0, 3 ] ] );
$bmat = ( $item1 ) ? Matrix( [ [ 1, -1 ], [ 0, 1 ] ] ) :
    Matrix( [ [ 1, 1 ], [ 0, 1 ] ] );
$zero = Matrix( [ [ 0, 0 ], [ 0, 0 ] ] );
$dchk = new MultiAnswer( $amat, $bmat )->with(
    singleResult=>1,
    checker=>sub{
	my ( $cor, $stu, $ans ) = @_;
        @s = @{$stu};
	$ssum = ( $item1 ) ? $s[0] + $s[1] : $s[0] - $s[1];
	return $zero != $s[0] && $zero != $s[1] &&
	    ( $s[0]->element(1,2) != 0 || $s[0]->element(2,1) != 0 ) &&
	    ( $s[1]->element(1,2) != 0 || $s[1]->element(2,1) != 0 ) &&
	    $ssum->element(1,2) == 0 && $ssum->element(2,1) == 0;
} );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
${BITALIC}(Note that this is a compound problem, for which the
second part of the problem is shown after the first is answered
correctly.)$EITALIC
$PAR

Is the following statement true or false?
$PAR
If \(A\) and \(B\) are \(n\times n\) diagonal matrices, then
\($qlist[$item1]\) and \($qlist[$item2]\) are diagonal matrices also.
$BR
\{ $tfpop->menu() \}

END_TEXT
Context()->normalStrings;

ANS( $tfpop->cmp() );

$ans_hash1 = $tfpop->cmp()->evaluate( $inputs_ref->{ANS_NUM_TO_NAME(1)} );

if ( $ans_hash1->{score} == 1 ) {

Context()->texStrings;
BEGIN_TEXT

$PAR
Next, note that if \($qlist[$item1]\) is a diagonal matrix, it is
not necessarily the case that \(A\) and \(B\) are diagonal.  Give
nonzero \(2\times 2\) matrices \(A\) and \(B\) that are not diagonal
but for which \($qlist[$item1]\) is diagonal.
$BR
\( A = \) \{ $dchk->ans_array() \} $BR
\( B = \) \{ $dchk->ans_array() \}

END_TEXT

ANS( $dchk->cmp() );

} # end if ans_hash

if ( $item2 > 3 ) {
    $not = "not";
    $nstr = $qlist[$item2];
    $nstr =~ s/~~^T//;
    $expl = "while \($nstr\) is diagonal, \($qlist[$item2]\) will be " .
        "have entries along the backward diagonal of the product, which " .
	"is not our definition of a diagonal matrix.";
} else {
    $not = 'also';
    $expl = 'working out the terms in the product, we see that there are ' .
        'only terms on the diagonal of the product.';
}

Context()->texStrings;
BEGIN_SOLUTION

If \(A\) and \(B\) are diagonal, then \($qlist[$item1]\) is diagonal: if
\( i\ne j \), \(A_{ij} = B_{ij} = 0\), so \( ($qlist[$item1])_{ij} = 0\),
and \($qlist[$item1]\) is diagonal.  Then \($qlist[$item2]\) is $not
diagonal: $expl.

$PAR

There are many possible examples of matrices for which
\($qlist[$item1]\) is diagonal; \(A = $amat\) and \(B = $bmat\) constitute
one such.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
