# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('The Inverse of a Matrix')
## KEYWORDS('inverse matrix')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.3')
## Problem1('39')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);

# make sure we're in the context we want
Context("Matrix");

# create some random matrices that do what we want
$a11 = random(3,6,1);
$a13 = $a11 - 1;
$a22 = list_random( -1, 1 );
$b22 = $a22;

$amat = Matrix( [ [ $a11, 0, $a13 ], [ 0, $a22, 0 ] ] );
$bmat = Matrix( [ [ 1 , 0 ], [ 0, $b22 ], [ -1, 0 ] ] );

$matcmp = MultiAnswer( $amat, $bmat )->with(
    singleResult => 1,
    checker => sub {
	my ( $cor, $stu, $self ) = @_;
	return ( Value::Matrix->I(2) == $stu->[0]*$stu->[1] &&
		 Value::Matrix->I(3) != $stu->[1]*$stu->[0] );
    } );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find a \(2\times3\) matrix \(A\) and a \(3\times2\) matrix \(B\)
such that \(AB = I_2\) but \(BA \ne I_3\).
$PAR
\( A = \) \{ $matcmp->ans_array() \}
$PAR
\( B = \) \{ $matcmp->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $matcmp->cmp() );

$abmat = $amat * $bmat;
$bamat = $bmat * $amat;

Context()->texStrings;
BEGIN_SOLUTION

There are many answers.  One is \(A = $amat\) and \(B = $bmat\),
for which \(AB = $amat $bmat = $abmat \) while
\(BA = $bmat $amat = $bamat \).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
