# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Matrix Operations')
## KEYWORDS('matrix multiplication')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.2')
## Problem1('34')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);

# make sure we're in the context we want
Context("Matrix");

# create two 3x3 matrices as reference
$amat = Matrix( [ [ 1, 0, 1 ], [ 1, 1, 0 ], [ 1, 0, 1 ] ] );
$bmat = Matrix( [ [ 1, 1, 1 ], [ -1, -1, -1 ], [ -1, -1, -1 ] ] );
$zero = Matrix( [ [ 0, 0, 0 ], [ 0, 0, 0 ], [ 0, 0, 0 ] ] );

$abChecker = MultiAnswer( $amat, $bmat )->with(
    singleResult => 1,
    checker => sub {
	my ( $cor, $stu, $ans ) = @_;
	@c = @{$cor};
	@s = @{$stu};
	return $zero != $s[0] && $zero != $s[1] && $zero == $s[0]*$s[1];
} );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find two nonzero \(3\times 3\) matrices \(A\) and \(B\) such that
\(A B = 0_{33}\).
$BR
\( A = \) \{ $abChecker->ans_array() \} $BR
\( B = \) \{ $abChecker->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $abChecker->cmp() );

Context()->texStrings;
BEGIN_SOLUTION

There are many solutions; one is
\[ A = $amat, B = $bmat. \]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
