# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('The Inverse of a Matrix')
## KEYWORDS('inverse matrix, singular matrix')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.3')
## Problem1('36')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
);

# make sure we're in the context we want
Context("Matrix");

# create a random matrix as an example
$a11 = random(1,4,1);  $a21 = random(1,4,1);
$a12 = random(1,4,1);  $a22 = random(1,4,1);
$a13 = random(1,4,1);  $a23 = random(1,4,1);
# make the third row linearly dependent
$k1 = random(1,2,1);
$k2 = random(1,3,1);
$a31 = $k1*$a11 + $k2*$a21;
$a32 = $k1*$a12 + $k2*$a22;
$a33 = $k1*$a13 + $k2*$a23;

$amat = Matrix( [ [ $a11, $a12, $a13 ], [ $a21, $a22, $a23 ],
		  [ $a31, $a32, $a33 ] ] );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find a singular \(3\times 3\) matrix \(A\) with no zero entries.
$PAR
\( A = \) \{ $amat->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $amat->cmp( checker=>sub {
    my ( $cor, $stu, $self ) = @_;
    foreach my $i ( 1,2,3 ) {
	foreach my $j ( 1,2,3 ) {
	    if ( Compute(0) == $stu->element($i,$j) ) {
		$self->{ans_message} = 'Your matrix should contain ' .
		    'no zero entries.';
		return 0;
	    }
	}
    }
    return $stu->det() == 0;
   } ) );

Context()->texStrings;
BEGIN_SOLUTION

There are many answers.  One is \(A = $amat\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
