#########################################################
# DESCRIPTION
## DBsubject(Statistics)
## DBchapter(Hypothesis tests)
## DBsection(Contingency Tables)
## level(4)

#########################################################
DOCUMENT();
## Initializations: (Required)
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "unionTables.pl",
  "answerCustom.pl",
  "PGcourse.pl"
);

#########################################################
## Setup: this is where we use Perl and PG objects (Required)
# this is to create the Contraceptives.csv file for ASSGT 4
# the data will have "OC" (random) users of oral contraceptives
# the data will have "noOC" (random) non-users of oral contraceptives
# the data will have "OC.MI" (random) users of oral contraceptives
# who develop myocardial infarction
# the data will have "noOC.MI" (random) non-users of oral contraceptives
# who develop myocardial infarction
# these will then be put into a dataframe, saved as "Contraceptives.csv"

## GENERATE THE DATA 
Context("Numeric");
# warn("Student name: $studentName, login: $studentLogin, id: $studentID");
my $hash = crypt($studentLogin, 'a1');
# warn("crypt: $hash");
@nums = (0..9,'a'..'z','A'..'Z');
my %nums = map { $nums[$_] => $_ } 0..$#nums;
my $seed = 0;
$seed = $seed * 62 + $nums{$_} foreach split(//, substr($hash, -5, 5));
# warn("seed: " . $seed);

rserve_eval("set.seed($seed)");
## fix the name of the downloadable data files
rserve_eval("filename <- file.path(tempdir(), 'Contraceptives_${studentLogin}.csv')");

my @remote_file = rserve_eval('
# randomly generate the number of OC users, and non-OC users
OC <- sample(size=1, 5000:6000)
noOC <- sample(size=1, 10000:11000)
# randomly generate the number of OC users and non-OC users
# who experience myocardial infarction
OC.MI <- sample(size=1, 11:17)
noOC.MI <- sample(size=1, 5:11)
# store the data is a matrix, and save as a .csv file
dat <- matrix( c(rep(c(1,1), times=OC.MI), 
                 rep(c(1,0), times=(OC-OC.MI)), 
                 rep(c(0,1), times=noOC.MI), 
                 rep(c(0,0), times=(noOC-noOC.MI))), 
               ncol=2, nrow=(OC+noOC), byrow=T)
# give it column names
colnames(dat) <- c("OC.use", "MI.occur")
# make it into a dataframe
dat <- as.data.frame(dat)
write.csv(dat, filename, row.names=FALSE)
filename
');
my $local_file = rserve_get_file($remote_file[0]);
$local_url = $local_file;
$local_url =~ s|$tempDirectory|$tempURL|;


rserve_eval("set.seed($seed)");
@data = rserve_eval('
# randomly generate the number of OC users, and non-OC users
OC <- sample(size=1, 5000:6000)
noOC <- sample(size=1, 10000:11000)
# randomly generate the number of OC users and non-OC users
# who experience myocardial infarction
OC.MI <- sample(size=1, 11:17)
noOC.MI <- sample(size=1, 5:11)
# store the data is a matrix, and save as a .csv file
dat <- matrix( c(rep(c(1,1), times=OC.MI), 
                 rep(c(1,0), times=(OC-OC.MI)), 
                 rep(c(0,1), times=noOC.MI), 
                 rep(c(0,0), times=(noOC-noOC.MI))), 
               ncol=2, nrow=(OC+noOC), byrow=T )
# give it column names
colnames(dat) <- c("OC.use", "MI.occur")
# make it into a dataframe
dat <- as.data.frame(dat)

### This is the code to calculate and store all the possible p-values
## for part(a)
q1 <- round((OC.MI/OC)*100, 2)

## for part(c)
q3 <- round((OC.MI/(OC.MI+noOC.MI))*100, 2)

## for part (h)
answers <- round( c( chisq.test(table(dat$OC.use, dat$MI.occur), correct=T)$p.value,
                     chisq.test(table(dat$OC.use, dat$MI.occur), correct=F)$p.value,
                     fisher.test(table(dat$OC.use, dat$MI.occur))$p.value,
                     prop.test(x=c(OC.MI, noOC.MI), n=c(OC, noOC))$p.value), 4)


p <- 0.05

num_sig <- sum(answers < p)
num_sig <- ifelse(num_sig > 0, 1, 0)

## answers for (j)
p1 <- (OC.MI/OC)
p2 <- (noOC.MI/noOC)
RD <- round((p1-p2)*100, 2)

## answer for (k)
p1 <- (OC.MI/OC)
p2 <- (noOC.MI/noOC)
RR <- round(p1/p2, 1)

## answer for (l)
p1 <- (OC.MI/OC)
p2 <- (noOC.MI/noOC)
OR <- round(((p1/(1-p1))/(p2/(1-p2))), 1)

x <- c(q1, q3, answers, RD, RR, OR, num_sig)
');

$dat = join(",  ", @data);

$q1ans = $data[0];

$q2mc = new_multiple_choice();
$q2mc -> qa(
"Do you believe the percentage calculated in part (a) is representative of the population percentage of OC users who would develop an MI?",
"YES, it is representative"
);
$q2mc -> extra(
"NO, it is not representative"
);

$q3ans = $data[1];

$q4mc = new_multiple_choice();
$q4mc -> qa(
"Do you believe the percentage calculated in part (c) is representative of the population percentage of MI occurrences who use OCs?",
"NO, it is not representative"
);
$q4mc -> extra(
"YES, it is representative"
);

$q5mc = new_checkbox_multiple_choice();
$q5mc -> qa(
"Which plot(s) allow one to examine the relationship between two categorical variables, to examine if there is any association between the two? Choose all that apply.",
"A mosaic plot.", 
"Side-by-side bar charts."
);
$q5mc -> extra(
"A bar chart.", "Side-by-side boxplots.", "A scatterplot."
);

$q6mc = new_multiple_choice();
$q6mc -> qa(
"Produce a plot to examine the relationship between OC use and MI occurrence (you do not need to include this plot with your assignment submission). Using only this plot, are you able to determine if there is an association between OC use and MI occurrence?",
"NO, it is not possible to say, from this plot alone"
);
$q6mc -> extra(
"YES, they are clearly associated"
);

$q7mc = new_checkbox_multiple_choice();
$q7mc -> qa(
"Which hypothesis test(s) would be appropriate for testing if there
is a statistically significant relationship between OC use and MI
occurrence? Choose all that apply.",
"Two proportions t-test",
"Fisher's exact test",
"Pearsons Chi-square test of independence",
"Hypothesis test for OR = 1"
);
$q7mc -> extra(
"Two means t-test"
);


$q81ans = Real($data[2]);
$q82ans = Real($data[3]);
$q83ans = Real($data[4]);
$q84ans = Real($data[5]);
@q8ans = ($q81ans, $q82ans, $q83ans, $q84ans);
$q8tol = Real("0.0002");

$p = Real("0.05");
$q9mc_que = "Based on the test conducted in part (h), what do you conclude about the relationship between OC use and MI occurrence?";
$q9mc_ans1 = "1. There is a statistically significant relationship between OC use and MI occurrence.";
$q9mc_ans2 = "2. There is not a statistically significant relationship between OC use and MI occurrence.";
if ($data[9] == 1){
$q9mc = RadioButtons([$q9mc_ans1, $q9mc_ans2],  $q9mc_ans1, labels=>[1, 2],  order=>[$q9mc_ans1, $q9mc_ans2] );
$q9_cor = 1;
} else {
$q9mc = RadioButtons([$q9mc_ans1, $q9mc_ans2],  $q9mc_ans2, labels=>[1, 2],  order=>[$q9mc_ans1, $q9mc_ans2] );
$q9_cor = 2;
}


$multi_q8q9 = MultiAnswer($q82ans, Real($q9_cor))->with(
  singleResult => 0,
  checker => sub {
		my ( $correct, $student, $self ) = @_;
		my ( $stu_q8, $stu_q9 ) = @{$student};
		if ((grep {abs($stu_q8 - $_) < $q8tol} @q8ans)){
                     if ($stu_q8 < $p) {
                         if ($stu_q9 == 1){
                             return [1, 1];
                         } else { return [1, 0];}
                    } else {
                                if ($stu_q9 == 1){
                                    return [1, 0]; 
                                } else {return [1, 1];}
                               }
                 } else{ 
                            if ($stu_q8 < $p){
                               if ($stu_q9 == 1){
                                   return [0, 1];
                               } else {return [0, 0];}
                           } else {
                                        if ($stu_q9 == 1){
                                             return [0, 0]; 
                                         } else {return [0, 1];}
                                    }
                 }
});


$q10ans = $data[6];

$q11ans = $data[7];

$q12ans = $data[8];

#########################################################
## Main Text: where all text goes (Required)
Context()->texStrings;
BEGIN_TEXT
This assignment will be submitted in two parts.  Some answers will be submitted directly to the online homework system, and these questions will include a note saying $BBOLD (WebWorkiR) $EBOLD at the start.  Other questions will be submitted on paper, in class.  These questions will include a note saying $BBOLD (On-Paper) $EBOLD at the start.
$BR
$BR
A study looked at the effect of oral contraceptive (OC) use on heart disease
(MI=myocardial infarction) in women aged 40-44. Researchers followed a
cohort of women who used oral contraceptives as well as a group who did
not use oral contraceptives, following them for a 3-year period. Use of oral
contraceptives was assumed to not change over the duration of study. The file $BBOLD Contraceptives.csv  (\{ htmlLink($local_url, "Download the csv file here") \}) $EBOLD contains data for this study. The variable OC.use is an indicator for oral contraceptive use, with 1 indicating the subject used oral contraceptives. The variable MI.occur is an indicator for evidence of myocardial infarction, with 1 indicating occurrence of myocardial infarction.
$PAR

Use the data to answer the following questions.
$PAR
$BBOLD Part (a) (WebWorkiR) $EBOLD
What percentage of women in the study who use OCs developed MI within a 3-year period? Provide your answer as a percentage, rounded to two decimal places.
\{ ans_rule(7) \}\(%\)
$BR
$BR
$BBOLD Part (b) (WebWorkiR) $EBOLD
\{ $q2mc -> print_q() \}
$BR
\{ $q2mc -> print_a() \}
$BR
$BR
$BBOLD Part (c) (WebWorkiR) $EBOLD
What percentage of women who developed an MI within a 3-year period used OCs? Give your answer as a percentage, rounded to two decimal places.
\{ ans_rule(7) \}\(%\)
$BR
$BR
$BBOLD Part (d) (WebWorkiR) $EBOLD
\{ $q4mc -> print_q() \}
$BR
\{ $q4mc -> print_a() \}
$BR
$BR
$BBOLD Part (e) (WebWorkiR) $EBOLD
\{ $q5mc -> print_q() \}
$BR
\{ $q5mc -> print_a() \}
$BR
$BR
$BBOLD Part (f) (WebWorkiR) $EBOLD
\{ $q6mc -> print_q() \}
$BR
\{ $q6mc -> print_a() \}
$BR
$BR
$BBOLD Part (g) (WebWorkiR) $EBOLD
\{ $q7mc -> print_q() \}
$BR
\{ $q7mc -> print_a() \}
$BR
$BR
$BBOLD Part (h) (WebWorkiR) $EBOLD
Conduct a hypothesis test to test if there is a significant relationship between OC use and MI occurrence. Provide the p-value for the test. Round the p-value to four decimal places.
\{ ans_rule(7) \}
$BR
$BR
$BBOLD Part (i) (WebWorkiR) $EBOLD
$q9mc_que
$BR
\{ $q9mc ->buttons() \}
$BR
$BR
$BBOLD Part (j) (WebWorkiR) $EBOLD
Calculate the Risk Difference/Attributable Risk (RD/AR), considering OC use as a risk factor for MI. Provide the RD as a percentage, rounded to two decimal places.
\{ ans_rule(7) \}\(%\)
$BR
$BR
$BBOLD Part (k) (WebWorkiR) $EBOLD
Calculate the Risk Ratio (RR) for this study, considering OC use as a risk factor for MI. Round the RR to one decimal place.
\{ ans_rule(7) \}
$BR
$BR
$BBOLD Part (l) (WebWorkiR) $EBOLD
Calculate the Odds Ratio (OR) for this study, considering OC use as a risk factor for MI. Round the OR to one decimal place.
\{ ans_rule(7) \}
$BR
$BR
$BBOLD Part (m) (On-Paper) $EBOLD
Write up a report no more than one page long including any plots/tables/summaries, summarizing the relationship between OC use and MI occurrence, based on your analysis of this data. You may want to comment on the significance of any association, as well as comment on some measures of the strength and direction of association. Include any information you think is relevant, and exclude anything you think is not relevant.
$BR
$BR
END_TEXT

#########################################################
## Answers evaluation (Required)
ANS( num_cmp($q1ans, tol=> 0.01, format =>'%0.2f') );
ANS( radio_cmp( $q2mc -> correct_ans() ) );
ANS( num_cmp($q3ans, tol=> 0.01, format =>'%0.2f') );
ANS( radio_cmp( $q4mc -> correct_ans() ) );
ANS( checkbox_cmp( $q5mc->correct_ans() ) );
ANS( radio_cmp( $q6mc -> correct_ans() ) );
ANS( checkbox_cmp( $q7mc->correct_ans() ) );
$showPartialCorrectAnswers = 1;
ANS( $multi_q8q9->cmp() ); 
ANS( num_cmp($q10ans, tol=> 0.01, format =>'%0.2f') );
ANS( num_cmp($q11ans, tol=> 0.1, format =>'%0.1f') );
ANS( num_cmp($q12ans, tol=> 0.1, format =>'%0.1f') );


#########################################################
## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
$BBOLD Part (a) $EBOLD
The required percentage is $q1ans%. 
$BR
$BBOLD Part (b) $EBOLD
The answer is YES.
$BR
$BBOLD Part (c) $EBOLD
The percentage of women who developed an MI within a 3-year period and used OCs is $q3ans%.
$BR
$BBOLD Part (d) $EBOLD
The answer is NO.
$BR
$BBOLD Part (e) $EBOLD
Both a mosaic plot and side-by-side bar charts are appropriate.
$BR
$BBOLD Part (f) $EBOLD
No, one cannot tell from the plot alone.
$BR
$BBOLD Part (g) $EBOLD
The two-proportion t-test, Pearsons Chi-squared test, Fisher's exact test, and the test for OR = 1 are all applicable.
$BR
$BBOLD Part (h) $EBOLD
Depending which test is chosen to perform the p-value will be one of: $q81ans (Chi-squared test with continuity correction), $q82ans (Chi-squared test with no continuity correction), $q83ans (Fisher's exact test), $q84ans (two-proportion t-test).
$BR
$BBOLD Part (i) $EBOLD
If the p-value in (h) is less than 0.05, then there is a statistically significant relationship; otherwise, no such conclusion could be drawn.
$BR
$BBOLD Part (j) $EBOLD
The RD value is $q10ans%. 
$BR
$BBOLD Part (k) $EBOLD
The risk ratio is $q11ans.
$BR
$BBOLD Part (l) $EBOLD
The odds ratio is $q12ans.
$BR
$BBOLD Part (m) $EBOLD
Answer submitted ON-PAPER.

END_SOLUTION
ENDDOCUMENT();
