##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Calculus')
## DBchapter('Infinite sequences and series')
## DBsection('Geometric')
## Date('8/3/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "niceTables.pl",
   "contextFraction.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Fraction");
Context()->flags->set(
    reduceConstants=>0,
    reduceConstantFunctions=>0,
);
Context()->variables->are(n=>["Real", resolution=>1, limits=>[0,25]]);
Parser::Number::NoDecimals();
Context()->{error}{msg}{"You are not allowed to type decimal numbers in this problem"} 
  = "Decimals are not allowed for this problem, please use exact forms (such as fractions) instead of decimal approximations.";
    
@n = ( random(9,18,1), random(20,30,1), random(35,50,1), random(70,95,1) );

foreach my $i (0..3) {
    my $plusminus = ($i>1)?1:0;
    my $fraction = ($i%2 == 1)?1:0;
    my $r = random(2,5,1)*(-1)**$plusminus;
    $ratio[$i] = ($fraction==0)?Real($r):Fraction(1,$r);
    $init[$i] = random(2,9,1);
    $seq[$i][0] = Fraction($init[$i]);
    foreach my $j (1..4) {
        $seq[$i][$j] = $seq[$i][$j-1] * $ratio[$i];
    }
    $seqDisp[$i] = List($seq[$i][0],$seq[$i][1],$seq[$i][2],$seq[$i][3],$seq[$i][4]);
    if ($n[$i] % 10 == 1 && $n[$i] > 20) { $th[$i] = "st" }
    elsif ($n[$i] % 10 == 2 && $n[$i] > 20) { $th[$i] = "nd" }    
    elsif ($n[$i] % 10 == 3 && $n[$i] > 20) { $th[$i] = "rd" }
    else { $th[$i] = "th" }
    $nth[$i] = Formula("$init[$i]*($ratio[$i])^($n[$i]-1)");
    $closedForm[$i] = Formula("$init[$i]*($ratio[$i])^(n-1)");
    $closedFormTeX[$i] = $closedForm[$i]->TeX;
    $sn[$i] = Formula("$init[$i] * ( ( 1 - ($ratio[$i])^($n[$i]) )/( 1 - $ratio[$i] ) )");
    $infSum[$i] = Formula("$init[$i] * ( 1 / (1-$ratio[$i]) )") if ( abs($ratio[$i]) < 1 );
    $infSum[$i] = String("DNE")->with(typeMatch=>'Value::Real') if ( abs($ratio[$i]) >= 1 );
}

Context("LimitedNumeric");
Context()->{error}{msg}{"Can't use '+' in this context"} 
  = "Addition is disabled for this problem. <BR> Compute a decimal approximation for your answer instead.";
Context()->{error}{msg}{"Can't use '-' in this context"} 
  = "Subtraction is disabled for this problem. <BR> Compute a decimal approximation for your answer instead.";
Context()->{error}{msg}{"Can't use '*' in this context"} 
  = "Multiplication is disabled for this problem. <BR> Compute a decimal approximation for your answer instead.";
Context()->{error}{msg}{"Can't use '/' in this context"} 
  = "Division is disabled for this problem. <BR> Compute a decimal approximation for your answer instead.";
Context()->{error}{msg}{"Can't use '^' in this context"} 
  = "Exponents are disabled for this problem. Compute a decimal approximation for your answer instead.";
Context()->{error}{msg}{"Can't use '**' in this context"} 
  = "Exponents are disabled for this problem. Compute a decimal approximation for your answer instead.";

$multiAns = MultiAnswer(Real(0.5),Real(0.5**10),Real(0.5**100))->with(
    singleResult=>1,
    allowBlankAnswers=>1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $rStu, $r10Stu, $r100Stu ) = @{$student};
        Value::Error("Please enter a common ratio between -1 and 1") if ( $rStu eq "" );
        $self->setMessage(1,"Using zero as a common ratio makes for a very boring sequence. \(\{0,0,0,0,\dots\}\)") if ( $rStu == 0 );
        $self->setMessage(1,"You must select a common ratio between -1 and 1") if ( abs($rStu) >= 1 );
        $self->setMessage(2,"You have incorrectly computed \(($rStu)^{10}\)") if ( $r10Stu != ($rStu)**10 );
        $self->setMessage(3,"You have incorrectly computed \(($rStu)^{100}\)") if ( $r100Stu != ($rStu)**100 );
        return ( abs($rStu) < 1 && $r10Stu == ($rStu)**10 && $r100Stu == ($rStu)**100 && $rStu != 0 );
    });

##############################################################
#
#  Text
#
#

BEGIN_PGML

>> ### Infinite Geometric Sums ### <<

Under certain conditions, we can actually compute the sum of *all* the terms in a geometric sequence. For a geometric sequence such as [`a = \left\{ 2, 4, 8, 16, \dots \right\}`], the sum of _all_ the terms will be impossible - think about it, we'll be adding larger and larger quantities as we get farther along in the sequence. As a result, our _partial_ sums will continue to grow faster and faster as we consider longer and longer sums.

However, for _some_ geometric seqences, our terms *shrink* as we go farther into the sequence. Geometric sequences have _exponential_ behavior (as their closed form is always [`a_n = a_1 \cdot r^{n-1} `]) so the sequences that shrink are precisely the ones that have an exponential base between -1 and 1. In other words, we can compute the sum of *all* terms for _any_ geometric sequence with [`|r|<1`].

The *partial* sum, [`s_n`] is _always_ computable for a geometric sequence [`a_n = a_1 \cdot r^{n-1}`]:  
>> [`` s_n = \sum_{i=1}^n \left( a_1 \cdot r^{i-1} \right) = a_1 \left( \frac{1-r^n}{1-r} \right) ``] <<

Now, whenever [`|r|<1`], we should recognize that [`r^n`] will become smaller and smaller as we consider larger values of [`n`]. 

_Use only decimal approximations for the following questions:_  
* Select a value for [`r`] that satisfies the condition [`|r|<1`]: [__________]{$multiAns}
* Compute a decimal approximation for [`r^{10}`]: [__________]{$multiAns}
* Now compute a decimal approximation for [`r^{100}`]: [__________]{$multiAns}

In other words, for very *large* values of [`n`], we get that [`\color{red}{r^n \approx 0}`]. Let's look at how that can affect our computation of the partial sum:  
>> For large [`n`], when [`|r|<1`], [`` s_n = a_1 \left( \frac{1-\color{red}{r^n}}{1-r} \right) \color{red}{\approx} a_1 \left( \frac{1-\color{red}{0}}{1-r} \right) ``] <<

As a result, when we want to consider the sum of *all* terms in a geometric sequence:  
* First, we must satisfy the condition that [`|r|<1`], or else the sum of all the terms is impossible.
* Once we've confirmed that the common ratio is between -1 and 1, we can use the formula we approximated above:  
>> [`` \sum_{i=1}^{\infty} \left( a_1 \cdot r^{i-1} \right) = a_1 \left( \frac{1}{1-r} \right) ``] <<

>> #### Practice #### <<

Find the requested sums:  

* Use "DNE" if the requested sum _does not exist_.

1. [``\sum_{n=1}^{\infty} \left( [$closedFormTeX[0]] \right)``]  
    a. The first term appearing in this sum is [__________]{$init[0]}  
    b. The common ratio for our sequence is [__________]{$ratio[0]}  
    c. The sum is [____________________]{$infSum[0]}  

1. [``\sum_{n=1}^{\infty} \left( [$closedFormTeX[1]] \right)``]  
    a. The first term of the sequence [`a`] is [__________]{$init[1]}  
    b. The common ratio for the sequence [`a`] is [__________]{$ratio[1]}  
    c. The sum is [____________________]{$infSum[1]}  

1. [``\sum_{n=1}^{\infty} \left( [$closedFormTeX[2]] \right)``]  
    a. The first term of the sequence [`a`] is [__________]{$init[2]}  
    b. The common ratio for the sequence [`a`] is [__________]{$ratio[2]}  
    c. The sum is [____________________]{$infSum[2]}  

1. [``\sum_{n=1}^{\infty} \left( [$closedFormTeX[3]] \right)``]  
    a. The first term of the sequence [`a`] is [__________]{$init[3]}  
    b. The common ratio for the sequence [`a`] is [__________]{$ratio[3]}  
    c. The sum is [____________________]{$infSum[3]}  

END_PGML

##############################################################
#
#  Answers
#
#


ENDDOCUMENT();        
