# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('7.1')
## Problem1('3')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
"parserMultiAnswer.pl",
);

Context("Vector");

$v1 = ColumnVector( random(1,3,1), random(1,3,1), random(1,3,1) );
$v2 = ColumnVector( -1*random(1,3,1), random(1,3,1), -1*random(1,3,1) );
$v3 = ColumnVector( random(1,3,1), 0, -1*random(1,3,1) );
$e1 = random(-3,3,1);
$e2 = random(-3,3,1);
$e3 = random(-3,3,1);

@v1 = $v1->value();
@v2 = $v2->value();
@v3 = $v3->value();
$smat = Matrix( [[ $v1[0], $v2[0], $v3[0] ],
		 [ $v1[1], $v2[1], $v3[1] ],
		 [ $v1[2], $v2[2], $v3[2] ]] );
$dmat = Matrix( [[ $e1, 0, 0 ], [ 0, $e2, 0 ], [ 0, 0, $e3 ]] );

$mc = MultiAnswer( $smat, $dmat )->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checker => sub {
	my ( $c, $s, $a ) = @_;
	($cs, $cd) = @$c;
	($ss, $sd) = @$s;
	if ( $ss && $cs == $ss && $sd && $cd == $sd ) {
	    return [1,1];
	} elsif ( ! $cs ) {
	    return [0,0];
	} else {
	    @cscol = ( $cs->column(1), $cs->column(2), $cs->column(3) );
	    @cdcol = ( $cd->column(1), $cd->column(2), $cd->column(3) );
	    if ( $ss ) {
		@sscol = ( $ss->column(1), $ss->column(2), $ss->column(3) );
	    } else {
		@sscol = ( );
	    }
	    if ( $sd ) {
		@sdcol = ( $sd->column(1), $sd->column(2), $sd->column(3) );
	    } else {
		@sdcol = ( );
	    }

	    @eigval = ( $e1, $e2, $e3 );
	    @sscolCorrect = ( 0, 0, 0 );
	    @sdcolCorrect = ( 0, 0, 0 );
	    for ( my $i=0; $i<3; $i++ ) {
		if ( @sscol && $cscol[$i] == $sscol[0] ) {
		    $sscolCorrect[0] = 1;
		    if ( @sdcol &&
			 Matrix( [[$eigval[$i]], [0], [0]] ) == $sdcol[0] ) {
			$sdcolCorrect[0] = 1;
		    }
		} elsif ( @sscol && $cscol[$i] == $sscol[1] ) {
		    $sscolCorrect[1] = 1;
		    if ( @sdcol && 
			 Matrix( [[0], [$eigval[$i]], [0]] ) == $sdcol[1] ) {
			$sdcolCorrect[1] = 1;
		    }
		} elsif ( @sscol && $cscol[$i] == $sscol[2] ) {
		    $sscolCorrect[2] = 1;
		    if ( @sdcol &&
			 Matrix( [[0], [0], [$eigval[$i]]] ) == $sdcol[2] ) {
			$sdcolCorrect[2] = 1;
		    }
		}
	    }
	    return [ $sscolCorrect[0] && $sscolCorrect[1] && $sscolCorrect[2],
		     $sdcolCorrect[0] && $sdcolCorrect[1] && $sdcolCorrect[2] ];
	}
} );

$isDiag = PopUp( [ '?', 'yes', 'no' ], 'yes' );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Suppose that the eigenvectors of a \(3\times 3\) matrix \(A\) are
\( \vec v_1 = $v1 \), \( \vec v_2 = $v2 \) and \( \vec v_3 = $v3 \)
with corresponding eigenvalues \( \lambda_1 = $e1 \),
\( \lambda_2 = $e2 \), and \( \lambda_3 = $e3 \).
$PAR
Is \(A\) diagonalizable? \{ $isDiag->menu() \}
$PAR
If it is, what invertible matrix \(S\) and diagonal matrix \(B\)
diagonalize \(A\)?
$BR
\( S = \) \{ $mc->ans_array() \} $BR
\( B = \) \{ $mc->ans_array() \}
$BR
${BITALIC}(Enter both \(S\) and \(B\) before submitting.
If no such \(S\) and \(B\) exist, enter zeros for all entries
in both matrices.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $isDiag->cmp() );
ANS( $mc->cmp() );

ENDDOCUMENT();
