# WeBWorK problem written by Carl Yao, 2014
# Portland Community College
#
# Template:
# Change a quadratic function from standard form to vertex form. 
# Fractions are involved.
#
#       f(x) = x^2 + ax + b
#
# a is even.
#
# Last updated: Carl Yao, 11/13/2014
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','completing the square')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "contextFraction.pl",
  "PGcourse.pl",
);

##############################################

TEXT(beginproblem());
Context("Fraction");

$b = random(-5,5,2);
$c = non_zero_random(-10,10,1);

$vertexX = Fraction($b,2);
$vertexY = Fraction(4*$c-$b**2,4);

Context("Numeric");
Context()->flags->set(reduceConstants=>0,);
$ans = Compute("(x+$vertexX)**2+$vertexY");

$func = Compute("x**2+$b*x+$c")->reduce;


##############################################

BEGIN_PGML

With the skill of "completing the square," convert this quadratic function from standard form to vertex form, and then find this parabola's vertex.

Use fractions in your answer. Don't use decimals.

     [` f(x) = [$func] `]

     [`f(x)=`][____________________]{$ans}

     The parabola's vertex is [_____________].


END_PGML

##############################################

Context("Fraction-NoDecimals");
Context()->{precedence}{Fraction} = 2.5;
Context()->flags->set(reduceFractions=>0,
                      allowMixedNumbers=>1,
                      requirePureFractions=>1,
                      showExtraParens=>0, 
                      reduceConstants=>0,
                      );

Context()->parens->set('('=>{type=>'Point'});
$vertex = List("(-$vertexX,$vertexY)");
ANS($vertex->cmp(
      checker => sub {
        my ($correct, $student, $ans, $nth, $value) = @_;
        $student->context->setError("Your $nth $value contains an unreduced fraction","",undef,undef,$Value::CMP_WARNING)
          if $student->{notReduced};
        return 0 unless $correct == $student;
        my $reduced = 1;
        foreach $x ($student->value) {
          if (Value::classMatch($x,'Fraction') && !$x->isReduced) {
            $reduced = 0;
            break;
          }
        }
        $student->{notReduced} = 1 unless $reduced || $ans->{isPreview};
        return $reduced;
      }
    )
);

##############################################

$func1 = Compute("x**2+$b*x")->reduce;

$absB = abs($b);
$output1 = $b>0 ? "\frac{$b}{2}" : "-\frac{$absB}{2}";
$output2 = $b>0 ? "+\frac{$b}{2}" : "-\frac{$absB}{2}";

$absC = abs(4*$c-$b**2);
$output3 = 4*$c-$b**2>0 ? "+\frac{$absC}{4}" : "-\frac{$absC}{4}";

##############################################

BEGIN_PGML_SOLUTION

One way to change a quadratic function from standard form to vertex form is to use the skill of "completing the square."

To use the skill of "completing the square," we need to use the perfect square formula:

    [`` x^2+2ax+a^2 = (x+a)^2 ``]

Step 1: By the formula [`x^2+2ax+a^2 = (x+a)^2`], we have [`[$func1]+...= (x [$output2])^{2} `].

Step 2: By the formula [`x^2+2ax+a^2 = (x+a)^2`], we have [`[$func1]+\frac{[$b**2]}{4}=(x[$output2])^{2} `].

The full solution is:

    [``
\begin{aligned}
   f(x) &= [$func] \\
         &= [$func1]+\frac{[$b**2]}{4} -\frac{[$b**2]}{4} +[$c] \\
         &= (x[$output2])^{2} -\frac{[$b**2]}{4} +[$c] \\
         &= (x[$output2])^{2} [$output3] \\
\end{aligned}
    ``]

Once we change the function's equation to vertex form, we can see the parabola's vertex is [`[$vertex]`].

END_PGML_SOLUTION

ENDDOCUMENT();




