# WeBWorK problem written by Carl Yao, 2015
# Portland Community College
#
# Template:
# Change a quadratic function from standard form to vertex form. 
# Fractions are involved.
#
#       f(x) = ax^2 + bx + c
#
# a is even.
#
# Last updated: Carl Yao, 09/30/2015
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','completing the square')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "contextFraction.pl",
  "PGcourse.pl",
);

##############################################

TEXT(beginproblem());
Context("Fraction");

Context()->variables->are(x=>"Real",y=>"Real");
$var = 'x';

do {
   $a = random(2,5,1);
   $b = random(1,5,1);
} until ( gcd($a,$b)==1 );
$b = $b*random(-1,1,2);
$c = random(1,5,1)*random(-1,1,2);

$vertexX = Fraction($b,2*$a);
$vertexY = Fraction(4*$a*$c-$b**2,4*$a);

Context("Numeric");
Context()->flags->set(reduceConstants=>0,);
$ans = Compute("$a*(x+$vertexX)**2+$vertexY");

$func = Compute("$a*x**2+$b*x+$c")->reduce;
$func1 = Compute("$a*x**2+$b*x")->reduce;

##############################################

BEGIN_PGML

With the skill of "completing the square," convert this quadratic function from standard form to vertex form, and then find this parabola's vertex.

Use fractions in your answer. Don't use decimals.

     [` f(x) = [$func] `]

     [`f(x)=`][____________________]{$ans}

     The parabola's vertex is [_____________].


END_PGML

##############################################

Context("Fraction-NoDecimals");
Context()->{precedence}{Fraction} = 2.5;
Context()->flags->set(reduceFractions=>0,
                      allowMixedNumbers=>1,
                      requirePureFractions=>1,
                      showExtraParens=>0, 
                      reduceConstants=>0,
                      );

Context()->parens->set('('=>{type=>'Point'});
$vertex = List("(-$vertexX,$vertexY)");
ANS($vertex->cmp(
      checker => sub {
        my ($correct, $student, $ans, $nth, $value) = @_;
        $student->context->setError("Your $nth $value contains an unreduced fraction","",undef,undef,$Value::CMP_WARNING)
          if $student->{notReduced};
        return 0 unless $correct == $student;
        my $reduced = 1;
        foreach $x ($student->value) {
          if (Value::classMatch($x,'Fraction') && !$x->isReduced) {
            $reduced = 0;
            break;
          }
        }
        $student->{notReduced} = 1 unless $reduced || $ans->{isPreview};
        return $reduced;
      }
    )
);

##############################################

$output1 = Compute("$a*(x**2+$b/$a*x)")->reduce;
$output2 = Compute("x**2+$b/$a*x")->reduce;
$output3 = Fraction($b,2*$a)->reduce;
$output4 = Compute("(x+$output3)**2")->reduce;
$output5 = Compute("(x+$output3)(x+$output3)")->reduce;
$frac1 = Fraction($b**2,4*$a**2)->reduce;
$output6 = Compute("x**2+$b/$a x+$frac1")->reduce;
$output7 = Compute("$a*(x**2+$b/$a*x)+$c")->reduce;
$output8 = Compute("$a*(x**2+$b/$a*x+$frac1-$frac1)+$c")->reduce;
$output9 = Compute("$a*((x+$output3)**2-$frac1)+$c")->reduce;
$frac2 = $a*$frac1->reduce;
$output10 = Compute("$a*(x+$output3)**2-$frac2+$c")->reduce;
$frac3 = $c-$frac2->reduce;
$output11 = Compute("$a*(x+$output3)**2+$frac3")->reduce;

##############################################

BEGIN_PGML_SOLUTION

First, we need to change this quadratic function from standard form, [`f(x)=a x^2+bx+c`], to vertex form, [`f(x)=a(x+d)^2+e`].

To use the skill of "completing the square," we need to observe a pattern first:

    [``
\begin{aligned}
   (x+1)^{2} &= x^{2}+2x+1 \\
   (x+2)^{2} &= x^{2}+4x+4 \\
   (x+3)^{2} &= x^{2}+6x+9 \\
   (x+4)^{2} &= x^{2}+8x+16 \\
   ...
\end{aligned}
    ``]

The rule is:

    [`` (x+a)^{2} = x^2+2ax+a^2 ``]

For this problem, we need to factor out the coefficient in front of [` x^{2} `] first:

    [`` [$func1] = [$output1] ``]

and then complete the square for the part [` [$output2] `]. This implies the value of [`a`] in the above formula is [`[$output3]`]. Now we have [`[$output2]+...= [$output4] + ... `].

Note that [`[$output2]`] and [`[$output4]`] are not equivalent, because [`[$output4]=[$output5]=[$output6]`]. Now we have [`[$output2]=[$output4] - [$frac1]`].

The full solution is:

    [``
\begin{aligned}
   f(x) &= [$func] \\
   &= [$output7] \\
   &= [$output8] \\
   &= [$output9] \\
   &= [$output10] \\
   &= [$output11]
\end{aligned}
    ``]

Once we change the function's equation to vertex form, we can see the parabola's vertex is [`[$vertex]`].

END_PGML_SOLUTION

ENDDOCUMENT();




