# WeBWorK problem written by Carl Yao, 2014
# Portland Community College
#
# Template:
# Solve the following quadratic equation by completing the square
#
# 2x^2+x-3=0
#
# Last updated: Carl Yao, 11/13/2014
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','square root','fraction','complete the square')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################
DOCUMENT();
loadMacros(
     "PGstandard.pl",
     "MathObjects.pl",
     "parserAssignment.pl",
     "answerHints.pl",
     "PGML.pl",
      "PCCmacros.pl",
     "contextFraction.pl",
  "PGcourse.pl",
);

##############################################

Context("Fraction")->flags->set(reduceFractions => 0,showMixedNumbers=>0,showExtraParens=>0);
parser::Assignment->Allow;
Context()->operators->redefine(',',using=>',',from=>'Numeric');
Context()->operators->redefine('or',using=>',',from=>'Numeric');
Context()->operators->set(
  ','=>{string=>' or ',TeX=>'\hbox{ or }'},
  'or'=>{string=>' or ',TeX=>'\hbox{ or }'}
);
Context()->lists->set(List => {separator => " or "});
Context()->{error}{msg}{"Function 'sqrt' is not allowed in this context"} 
  = "Please simplify your answer further";

# add solution strings to context- this means that if
# students enter these (and they are not correct), then 
# WW will not give a Context warning
Context()->strings->add("no real solutions"=>{},
  "no real solution"=>{alias=>'no real solutions'}, 
  "none"=>{alias=>'no real solutions'}, 
  );

$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

do {
   do {
      $a = random(2,5,1);
      $b = random(1,5,1);
   } until ( gcd($a,$b)==1 );
   $b = $b*random(-1,1,2);
   $c = random(1,5,1)*random(-1,1,2);

   $soln1 = Fraction(-$b,$a)->reduce;
   $soln2 = Fraction(-$c,1)->reduce;
   $ans = Compute("$var = $soln1, $var = $soln2");

   $d = $a*$c+$b;
   $e = $b*$c;
} until ( (abs($d)<5) );
$func = Compute("$a*$var**2+$d*$var+$e")->reduce;
$func1 = Compute("$a*$var**2+$d*$var")->reduce;

##############################################

TEXT(beginproblem());
BEGIN_PGML
Solve the quadratic equation by the skill of *completing the square*:

* Enter multiple answers separated by commas, such as [`x=1,x=-1`]
* If you need to use the square root symbol, 
as in [`x=\sqrt{17}`], type it using: *sqrt(17)*
* If there are no real solutions, enter *no real solutions*

    [`  [$func] = 0 `]

    [______________________]


END_PGML

##############################################

$showPartialCorrectAnswers = 1;
ANS($ans->cmp(
  entry_type => "a solution",
  checker => sub {
    my ($correct,$student,$ans,$nth,$value) = @_;
    if ($correct->type eq "Assignment") {
      my ($svar,$sfrac) = $student->value; # get the variable and fraction
      if(Value::classMatch($sfrac,'Fraction'))
      {
        return 0 unless $sfrac->isReduced;
      }
    }
    return $correct == $student;
  },
  extra => sub {
    my ($student,$ansHash,$nth,$value) = @_;
    if($student eq "no real solutions")
    {
         $student->context->setError("This equation does have some solutions- try using the square root method","",undef,undef,$Value::CMP_WARNING)
         unless $ans->{isPreview};
         return;
    }
    if ($student->type ne "Assignment" && $ansHash->{student_formula}->type ne "Assignment") {
      $student->context->setError("Your $nth solution should be written $var = $US$US$US","",undef,undef,$Value::CMP_WARNING)
         unless $ans->{isPreview};
      return;
    }
    my ($svar,$sfrac) = $student->value; # get the variable and fraction
    if (Value::classMatch($sfrac,'Fraction') && !$sfrac->isReduced) {
      $student->context->setError("Your $nth $value is not reduced","",undef,undef,$Value::CMP_WARNING)
         unless $ans->{isPreview};
      return;
    }
    return Value::Real->typeMatch($student);
  }
)->withPostFilter(AnswerHints(
  ["$var=$soln1","$var=$soln2"] => "Are you sure you have all the solutions?",
  [$soln1,$soln2] => ["Your solution is a correct one, but you should write $var = $US$US$US<br>Are you sure you have all the solutions?",replaceMessage=>1],
  ["$soln1,$soln2","$soln1,$soln2"] => ["Your solutions are correct, but you should write $var = $US$US$US",replaceMessage=>1],
)));

##############################################

$fracB = Fraction($d,$a)->reduce;
$fracC = Fraction($e,$a)->reduce;

$func2 = Compute("$var**2+$fracB*$var+$fracC")->reduce;
$func3 = Compute("$var**2+$fracB*$var")->reduce;
$frac4 = (-$fracC)->reduce;

$frac5 = ($fracB/2)->reduce;
$func4 = Compute("($var+$frac5)**2")->reduce;
$frac6 = ($frac5*$frac5)->reduce;
$func5 = Compute("$var**2+$fracB*$var+($frac5)**2")->reduce;

$func6 = Compute("$var+$frac5")->reduce;

$frac1 = -$e + Fraction($d**2,$a*4);
$frac2 = ($frac4+$frac6)->reduce;
$frac3 = sqrt($frac2)->reduce;

$frac7 = ($frac3-$frac5)->reduce;
$frac8 = (-$frac3-$frac5)->reduce;

BEGIN_PGML_SOLUTION

There are easier ways to solve this quadratic equation, but we will solve it by the skill of "completing the square."

For this problem, we need to get rid of the leading coefficient [`[$a]`] by:

    [``
\begin{aligned}
   [$func]  &= 0 \\
   \frac{1}{[$a]}([$func]) &= \frac{1}{[$a]} \cdot 0 \\
   [$func2] &= 0 \\
   [$func3] &= [$frac4]
\end{aligned}
    ``]

Next, we complete the square for the part [` [$func3] `].

To use the skill of "completing the square," we need to use the perfect square formula:

    [`` x^2+2ax+a^2 = (x+a)^2 ``]

Step 1: By the formula [`x^2+2ax+a^2 = (x+a)^2`], we have [`[$func3]+...= [$func4] `].

Step 2: By the formula [`x^2+2ax+a^2 = (x+a)^2`], we have [`[$func5] = [$func4] `].

The full solution is:

    [``
\begin{aligned}
   [$func]  &= 0 \\
   \frac{1}{[$a]}([$func]) &= \frac{1}{[$a]} \cdot 0 \\
   [$func2] &= 0 \\
   [$func3] &= [$frac4] \\
   [$func3] \mathbf{+[$frac6]} &= [$frac4] \mathbf{+[$frac6]} \\
   [$func4] &=  [$frac2] \\
   [$func6] &=  \pm \sqrt{[$frac2]} \\
   [$func6] &=  [$frac3] \: \text{ or } \: [$func6] =  -[$frac3] \\
   [$var] &= [$frac7] \: \text{ or } \: [$var] = [$frac8]
\end{aligned}
    ``]

END_PGML_SOLUTION

##############################################

ENDDOCUMENT();
