# WeBWorK problem written by Carl Yao, 2016
# Portland Community College
#
# Template:
# Change a quadratic function from standard form to vertex form. 
# Fractions are involved.
#
#       f(x) = x^2 + ax + b
#
# a is odd.
#
# Last updated: Carl Yao, 08/18/2016
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','completing the square')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "contextFraction.pl",
  "PGcourse.pl",
);

##############################################

TEXT(beginproblem());
Context("Fraction");

$a = 1;
$b = random(-5,5,2);
$c = non_zero_random(-10,10,1);

$h = -Fraction($b,2)->reduce;
($hNum, $hDen) = $h->value;
$k = Fraction(4*$c-$b**2,4)->reduce;
($kNum, $kDen) = $k->value;

Context("Numeric");
$ans = Compute("(x-$h)**2+$k");

$func = Compute("x**2+$b*x+$c")->reduce;

$nhNum = -$hNum;
$nkNum = -$kNum;
$s1 = ($h>0) ? "(x-\frac{$hNum}{$hDen})^2" : "(x+\frac{$nhNum}{$hDen})^2";
$s2 = ($k>0) ? "+\frac{$kNum}{$kDen}" : "-\frac{$nkNum}{$kDen}";
$ansShow = $s1.$s2;


##############################################

BEGIN_PGML

Convert this quadratic function from standard form to vertex form, and then find this parabola's vertex. There is no need to use the skill of "completing the square."

Use fractions in your answer. Don't use decimals.

     [` f(x) = [$func] `]

     [`f(x)=`][____________________]{$ans}

     The parabola's vertex is [_____________].


END_PGML

##############################################

Context("Fraction-NoDecimals");
Context()->{precedence}{Fraction} = 2.5;
Context()->flags->set(reduceFractions=>0,
                      allowMixedNumbers=>1,
                      requirePureFractions=>1,
                      showExtraParens=>0, 
                      reduceConstants=>0,
                      );

Context()->parens->set('('=>{type=>'Point'});
$vertex = List("($h,$k)");
ANS($vertex->cmp(
      checker => sub {
        my ($correct, $student, $ans, $nth, $value) = @_;
        $student->context->setError("Your $nth $value contains an unreduced fraction","",undef,undef,$Value::CMP_WARNING)
          if $student->{notReduced};
        return 0 unless $correct == $student;
        my $reduced = 1;
        foreach $x ($student->value) {
          if (Value::classMatch($x,'Fraction') && !$x->isReduced) {
            $reduced = 0;
            break;
          }
        }
        $student->{notReduced} = 1 unless $reduced || $ans->{isPreview};
        return $reduced;
      }
    )
);

##############################################

$h2 = $h*2;
$s4 = Compute("x**2-$h2*x")->reduce;

BEGIN_PGML_SOLUTION

A quadratic function's vertex form looks like [` f(x) = a(x-h)^2+k `].

Since there is no coefficient for the [`x^2`] term, the value of [`a`] is simply [`1`]. Now we have [` f(x) = (x-h)^2+k `].

We can use the formula [`x=-\frac{b}{2a}`] to find a parabola's axis. In this problem, we have:

    [`` x=-\frac{b}{2a} = -\frac{[$b]}{2([$a])} = [$h] ``]

This implies [` h=[$h] `]. Now we have: [` f(x) = [$s1]+k `].

Next, expand the right side, and we have:

    [``
\begin{aligned}
   f(x) &= [$s1]+k \\
         &= [$s4]+[$h*$h]+k \\
\end{aligned}
    ``]

Compare the result with the given function [`f(x)=[$func]`], we have:

    [``
\begin{aligned}
   [$h*$h]+k &= [$c] \\
    k &= [$k]
\end{aligned}
    ``]

The function's vertex form is [`f(x)=[$ansShow]`], and we can tell the vertex is [`[$vertex]`].

END_PGML_SOLUTION

ENDDOCUMENT();




