# WeBWorK problem written by Carl Yao, 2016
# Portland Community College
#
# Template:
# Given the vertex and the value of a, write the quadratic function's formula.
#
# Last updated: Yao 08/18/2016
# ENDDESCRIPTION

## DBCC('')
## DBsubject('Algebra')
## DBchapter('Quadratic equations and functions')
## DBsection('Forms: vertex, factored, general')
## KEYWORDS('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

TEXT(beginproblem()); 

Context("Numeric");
Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
'*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
'/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
'//' => {class => 'bizarro::BOP::divide', isCommand => 1},
);


$a = non_zero_random(-9,9,1);
$h = non_zero_random(-9,9,1);
$k = non_zero_random(-9,9,1);
$ans = Compute("$a*(x-$h)^2+$k");
if ($a == 1) {$ans = Compute("(x-$h)^2+$k");};
if ($a == -1) {$ans = Compute("-(x-$h)^2+$k");};

$evaluator = $ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroAdd=>1,bizarroSub=>1,bizarroMul=>1,bizarroDiv=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      my $OK = (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroAdd=>0,bizarroSub=>0,bizarroMul=>0,bizarroDiv=>0);
      Value::Error("Your answer is equivalent to the correct answer, but not in vertex form") unless $OK;
      return $OK;
});


##############################################

BEGIN_PGML

Write the vertex form for the quadratic function [`f`], whose vertex is [`([$h],[$k]) `] and has leading coefficient [`a=[$a]`].

    [`` f(x) =``] [_________________]{$evaluator}

END_PGML

##############################################


BEGIN_PGML_SOLUTION

The solution is [` f(x) = [$a](x-[$h])^2+[$k] `].

END_PGML_SOLUTION


ENDDOCUMENT(); 
