## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Applications)
## Date(03/19/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-A-springs_n_things)
## Level(5)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('3.6')
## Problem1('')



##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionLists.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##################################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");


$f = list_random(2,3,4);
$ff = $f * $f;

#  Differential equation in matrix form:
#
#  [ $m1  0   ] [ x_1'' ]     [ -($k1+$k2)        $k2  ] [ x_1 ]
#  [   0  $m2 ] [ x_2'' ]  =  [       $k2   -($k2+$k3) ] [ x_2 ]
# 

# matrix A = M^(-1) K listed by rows (not columns)
@A = ();
$A[0] = Real("-(10*$ff)"); # A11
$A[1] = Real("6*$ff");     # A12
$A[2] = Real("6*$ff");     # A21
$A[3] = Real("-(10*$ff)"); # A22

#
#  masses and spring constants
#
$m1 = random(1,5,1);
$m2 = $m1;

$k1 = 4 * $m1 * $ff;
$k2 = 6 * $m1 * $ff;
$k3 = $k1;



$multians1 = MultiAnswer( $A[0], $A[1], $A[2], $A[3] )->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='5'>
    <tr><td> %s </td><td> %s </td></tr>
    <tr><td> %s </td><td> %s </td></tr>
    </table>",
    tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\[2ex] %s & %s \end{array}\right\rbrack" ,
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($stuA11, $stuA12, $stuA21, $stuA22) = @{$student};

        if ( ($stuA11 == $A[0]) && ($stuA12 == $A[1]) && ($stuA21 == $A[2]) && ($stuA22 == $A[3]) ) 
        { return 1; } else { Value::Error('At least one entry is incorrect'); }

    }
);

$c0 = Real("1");
$c1 = Real("1");
$c2 = Formula("2*$f t");
$c3 = Formula("2*$f t");
$c4 = Real("-1");
$c5 = Real("1");
$c6 = Formula("4*$f t");
$c7 = Formula("4*$f t");


$multians2 = MultiAnswer( $c0, $c1, $c2, $c3, $c4, $c5, $c6, $c7 )->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "%s, %s, %s, %s, <br /> %s, %s, %s, %s",
    tex_format =>  "\left\lbrack\begin{array}{rr} %s \\[2ex] %s \end{array}\right\rbrack (a_1 \cos(%s) + b_1 \sin(%s)) + 
                    \left\lbrack\begin{array}{rr} %s \\[2ex] %s \end{array}\right\rbrack (a_2 \cos(%s) + b_2 \sin(%s))",
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($s0, $s1, $s2, $s3, $s4, $s5, $s6, $s7) = @{$student};

        if ( (
             ($s0 != 0 || $s1 != 0) && 
             ($s4 != 0 || $s5 != 0) && 
             ($c0 * $s1 == $c1 * $s0) && 
             ($c2 == $s2) && 
             ($c3 == $s3) && 
             ($c4 * $s5 == $c5 * $s4) && 
             ($c6 == $s6) && 
             ($c7 == $s7)  
             ) || (
             ($s0 != 0 || $s1 != 0) && 
             ($s4 != 0 || $s5 != 0) && 
             ($c4 * $s1 == $s0 * $c5) && 
             ($c6 == $s2) && 
             ($c7 == $s3) && 
             ($c0 * $s5 == $c1 * $s4) && 
             ($c2 == $s6) && 
             ($c3 == $s7)  
             )
           ) 
        { return 1; } else { Value::Error('At least one entry is incorrect'); }

    }
);



#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{
ColumnTable(
"Consider a system of two toy railway cars (i.e., frictionless masses)
connected to each other
by three springs, two of which are attached to walls, as shown in the
figure.  Let \( x_1 \) and \( x_2 \) be the displacement of the first 
and second masses from their equilibrium positions.  Suppose the 
masses are \( m_1 = $m1 \ \mathrm{kg} \) and \( m_2 = $m2 \ \mathrm{kg} \),
and the spring constants are \( k_1 = $k1 \ \mathrm{N/m} \), 
\( k_2 = $k2 \ \mathrm{N/m} \), and \( k_3 = $k3 \ \mathrm{N/m} \)."
,
$BCENTER.
image("2masses-3springs-new.png", width=>444, height=>122, tex_size=>800).
$ECENTER
,
indent => 0, separation => 30, valign => "TOP"
);
\}

\{ BeginList('OL',type=>'a') \}

$ITEM
Set up a system of second-order differential equations that models this
situation.
\{ 
mbox(
[
'\( \boldsymbol{\vec{x}^{\, \prime\prime}} = \) ', 
display_matrix([[$multians1->ans_rule(5),$multians1->ans_rule(5)],[$multians1->ans_rule(5),$multians1->ans_rule(5)]]),
' \( \boldsymbol{\vec{x}} \) '
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM
Find the general solution to this system of differential equations.
$BR
\{ 
mbox(
[
display_matrix([['\( x_1(t) \)'],['\( x_2(t) \)']]),
' = ',
display_matrix([[$multians2->ans_rule(5)],[$multians2->ans_rule(5)]]),
'\( \Big( a_1 \cos\big( \)' . $multians2->ans_rule(5) . '\( \big) + b_1 \sin\big( \)' . $multians2->ans_rule(5) . '\( \big) \Big) + \)',
display_matrix([[$multians2->ans_rule(5)],[$multians2->ans_rule(5)]]),
'\( \Big( a_2 \cos\big( \)' . $multians2->ans_rule(5) . '\( \big) + b_2 \sin\big( \)' . $multians2->ans_rule(5) . '\( \big) \Big) \)'
],
allowbreaks=>'yes'
) 
\}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

ANS( $multians2->cmp() );

;

ENDDOCUMENT();
