# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('parametric curve', 'multivariable', 'calculus')
## Tagged by glr 06/02/10


## DBsubject('Calculus')
## DBchapter('Parametric Equations and Polar Coordinates')
## DBsection('Curves Defined by Parametric Equations')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('17.1')
## Problem1('37')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')
## HHSection1('Parameterized Curves')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->are( t=>'Real' );
$showPartialCorrectAnswers = 1;

$r0 = random(2,9,1);
$twor0 = 2*$r0;

$pchk = MultiAnswer( Compute( "$r0 - $r0*cos(t)" ),
		     Compute( "$r0*sin(t)" ), Compute("0"),
		     Compute("pi") )->with(
    singleAnswer => 0,
    checker => sub {
	my ( $c, $s, $ans ) = @_;

	## make sure that the student answers are formulas
	my ( $xsf, $ysf ) = ( Formula( $s->[0] ), Formula( $s->[1] ) );

	my @ret = ( 0, 0, 0, 0 );
	if ( ($xsf - $r0)*($xsf - $r0) + $ysf*$ysf == $r0*$r0 ) {
	    $ret[0] = 1;
	    $ret[1] = 1;
	}
	if ( $xsf->eval(t=>$s->[2]) == 0 &&
	     $ysf->eval(t=>$s->[2]) == 0 ) {
	    $ret[2] = 1;
	}
	if ( $xsf->eval(t=>$s->[3]) == $twor0 &&
	     $ysf->eval(t=>$s->[3]) == 0 ) {
	    $ret[3] = 1;
	}
	return [@ret];
    }
);


Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find parametric equations for the
arc of a circle of radius $r0 from \( P=(0,0) \) to \( Q=($twor0,0) \).

$PAR
\( x(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( y(t) = \) \{ $pchk->ans_rule(35) \}, $BR
with
\{ $pchk->ans_rule(15) \} \( \le t\le \) \{ $pchk->ans_rule(15) \}

END_TEXT
Context()->normalStrings;

ANS( $pchk->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The line segment \(PQ\) has length $twor0, so it must be a diameter of
the circle.  The center of the circle is therefore the midpoint of
\(PQ\), which is the point \(($r0,0)\). The  upper arc of the circle between
\(P\) and \(Q\) can be parameterized as follows:
\[
\vec r (t) = $r0\vec i +  $r0(-\cos t \vec i + \sin t \vec j),
\quad 0 \le t \le \pi.
\]
The  lower arc can be parameterized as follows:
\[
\vec r (t) = $r0\vec i + $r0(\cos t \vec i + \sin t \vec j),
\quad \pi \le t \le 2\pi.
\]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
