# DESCRIPTION
# Skills building problem
# ENDDESCRIPTION

## DBsubject(Precalculus)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: point-slope form)
## Level(2)
## KEYWORDS('point-slope form', 'line')

DOCUMENT();
loadMacros("PGstandard.pl",
           "MathObjects.pl",
           "PGML.pl",
	   "parserPopUp.pl",
	   "parserMultiAnswer.pl",
# "PGchoicemacros.pl",
# "PGgraphmacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( y => 'Real' );

$a = random(-9,-5,1);
$b = random(-9,-5,1);
$c = random(-4,0,1);
$d = random(-4,0,1);

$ma = MultiAnswer( Compute("y-$b"), Compute( "($d-$b)*(x-$a)/($c-$a)" ) )->with(
    singleResult => 1,
    separator => ' = ',
    tex_separator => ' = ',
    checker => sub {
        my ( $cor, $stu, $self ) = @_;
	my ( $cleft, $cright ) = @{$cor};
	my ( $sleft, $sright ) = @{$stu};
	if ( ($cleft == $sleft && $cright == $sright) ||
	     ($cleft == $sright && $cright == $sleft) ) {
	    return (1,1);
	} else {
	    my $cleft2 = Compute( "y-$d" );
	    my $cright2 = Compute( "($d-$b)*(x-$c)/($c-$a)" );

	    if ( ( $cleft2 == $sleft && $cright2 == $sright ) ||
		 ( $cleft2 == $sright && $cright2 == $sleft ) ) {
		return (1,1);
	    } else {
		return (0,0);
	    }
	}
    }
);

TEXT(beginproblem());
BEGIN_PGML

Write the equation of the line that passes through the points
[`([$a], [$b])`] and [`([$c], [$d])`] in point-slope form.

    [___]{$ma}{5} [` = `] [___]{$ma}{15}

END_PGML

($n,$den) = reduce( $b-$d, $a-$c );
$m = ( $den == 1 ) ? "$n" : "\frac{$n}{$den}";

BEGIN_PGML_SOLUTION

The slope of the line is [`m = \frac{[$b]-[$d]}{[$a]-[$c]} = [$m]`].
Use either of the points [`([$a],[$b])`] or [`([$c],[$d])`] in the
point-slope form of the line, [`y - y_0 = m(x - x_0)`], to get
>> [`` y - [$b] = [$m] (x - [$a]), ``] <<
or
>> [`` y - [$d] = [$m] (x - [$c]). ``] <<

END_PGML_SOLUTION

ENDDOCUMENT();
