# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('chain rule', 'multivariable', 'functions', 'calculus', 'derivatives ')
## Tagged by glr 01/11/10

## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('The Chain Rule')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('14.6')
## Problem1('4')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')
## HHSection1('The Chain Rule')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( t=>'Real', y=>'Real', z=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,5,1);
$b = random(2,5,1);
$whichz = random(0,1,1);
if ( $whichz ) {
    $zfunc = Compute( "(x + y)*e^y" );
    $dzdt = Compute( "$a*e^y + (1 + x + y)*(e^y)*(-2*t)" );
    $dzdx = Compute( "e^y" );
    $dzdy = Compute( "(1 + x + y)*e^y" );
} else {
    $zfunc = Compute( "(x + y)*e^x" );
    $dzdt = Compute( "-2*t*e^x + (1 + x + y)*e^x*($a)" );
    $dzdx = Compute( "(1 + x + y)*e^x" );
    $dzdy = Compute( "e^x" );
}
$xfunc = Compute( "$a*t" );
$yfunc = Compute( "$b - t^2" );
$dxdt = Compute("$a");
$dydt = Compute("-2*t");

$dzdtsub = $dzdt->substitute( x=>$xfunc, y=>$yfunc )->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

If
\[ z = $zfunc,\qquad x = $xfunc,\qquad y = $yfunc, \]
find \( dz/dt \) using the chain rule. Assume the variables are
restricted to domains on which the functions are defined.

$PAR
\(\frac{dz}{dt} = \) \{ ans_rule(65) \}

END_TEXT
Context()->normalStrings;

ANS($dzdtsub->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( $correct == $student ) {
      return 1;
    } else {
      return $dzdt == $student ? 1 : 0;
    }
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Substituting into the chain rule gives
\[
\frac{dz}{dt} = \frac{\partial z}{\partial x}\frac{dx}{dt} +
\frac{\partial z}{\partial y}\frac{dy}{dt} \]
\[
  = ($dzdx)($dxdt) + ($dzdy)($dydt) \]
\[
  = $dzdtsub. \]

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
