# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 01/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Chain rule)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(chain3)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('14.6')
## Section2('14.6')
## Section3('14.6')
## Problem1('9')
## Problem2('11')
## Problem3('11')
## KEYWORDS('chain rule', 'multivariable', 'functions', 'calculus', 'derivatives ')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( u=>'Real', v=>'Real', y=>'Real' );
$showPartialCorrectAnswers = 1;

$r = random(2,3,1);
$whichz = random(0,1,1);
$whichz = 1;
if ( $whichz ) {
    $zfunc = Compute( "x*e^y" );
} else {
    $zfunc = Compute( "y*e^x" );
}
$whichu = random(0,1,1);
if ( $whichu ) {
    $xfunc = Compute( "u^$r + v^$r" );
    $yfunc = Compute( "u^$r - v^$r" );
} else {
    $xfunc = Compute( "u^$r - v^$r" );
    $yfunc = Compute( "u^$r + v^$r" );
}

$dzdx = $zfunc->D('x');
$dzdy = $zfunc->D('y');

$dxdu = $xfunc->D('u');
$dxdv = $xfunc->D('v');
$dydu = $yfunc->D('u');
$dydv = $yfunc->D('v');

$dzdxsub = $dzdx->substitute( x=>$xfunc, y=>$yfunc );
$dzdysub = $dzdy->substitute( x=>$xfunc, y=>$yfunc );
$dzdusub = Compute( "($dzdxsub)*($dxdu) + ($dzdysub)*($dydu)" )->reduce();
$dzdvsub = Compute( "($dzdxsub)*($dxdv) + ($dzdysub)*($dydv)" )->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

If
\[ z = $zfunc,\qquad x = $xfunc,\qquad y = $yfunc, \]
find \( \partial z/\partial u \) and \( \partial z/\partial v \).
The variables are restricted to domains on which
the functions are defined.
$PAR
\( \partial z/\partial u = \) \{ ans_rule(55) \}
$PAR
\( \partial z/\partial v = \) \{ ans_rule(55) \}

END_TEXT
Context()->normalStrings;

ANS($dzdusub->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( $correct == $student ) {
        return 1;
    } else {
        return ( Compute( "($dzdx)*($dxdu)+($dzdy)*($dydu)" ) == $student ) ? 1 : 0;
    } } ) );
ANS($dzdvsub->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( $correct == $student ) {
        return 1;
    } else {
        return ( Compute( "($dzdx)*($dxdv)+($dzdy)*($dydv)" ) == $student ) ? 1 : 0;
    } } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Since \(z\) is a function of two variables \(x\) and \(y\) which are
functions of two variables $u$ and $v$, the two chain rule identities
which apply are:
\[
\frac{\partial z}{\partial u}  = 
       \frac{\partial z}{\partial x}\frac{\partial x}{\partial u}+
       \frac{\partial z}{\partial y}\frac{\partial y}{\partial u} \]
\[ = ($dzdx) \left($dxdu\right) + ($dzdy)\cdot\left($dydu\right) \]
\[ = $dzdusub. \]
and
\[
\frac{\partial z}{\partial v} = 
       \frac{\partial z}{\partial x}\frac{\partial x}{\partial v}+
       \frac{\partial z}{\partial y}\frac{\partial y}{\partial v} \]
\[ = ($dzdx)\left($dxdv\right) + ($dzdy)\cdot \left($dydv\right) \]
\[ = $dzdvsub. \]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
