## MO(1)
# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPolynomial.pl",
  
);

TEXT(beginproblem());
##############################################

Context("Numeric");
Context()->variables->are(x=>'Real');
Context()->strings->add("cannot be factored"=>{},
    "cannot be factored."=>{alias=>"cannot be factored"});
Context()->flags->set(partialCredit=>0);

$a = non_zero_random(-20,20,1);

$exp = "x^2+$a x";

$fac1 = Compute("x");
$fac2 = Compute("(x+$a)")->reduce;

$fctd = Compute("$fac1  $fac2");
$fctd2 = Compute("$fac2  $fac1");

##############################################

Context()->texStrings;

BEGIN_TEXT

Factor the polynomial using the greatest common factor. 
$PAR


$SPACE $SPACE $SPACE \( $exp \)

$PAR


\{ans_rule(20)\}



 




END_TEXT

Context()->normalStrings;

##############################################
$showPartialCorrectAnswers = 1;


ANS($fctd->cmp(typeMatch=>Formula("x"), checker => sub {
      my ($correct,$student,$ans) = @_;
      if (($student eq $correct) || ($student eq $fctd2))
      {return 1};
    }));


ENDDOCUMENT();
