## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Kernel and image)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "MatrixReduce.pl",
  "MatrixUnits.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context('Matrix')->variables->are('w'=>'Real','x'=>'Real','y'=>'Real','z'=>'Real',);
Context()->noreduce('(-x)-y','(-x)+y');

@u = GL3Z_perl();
$U = Matrix(@u);

@areduced = (
[ 1, non_zero_random(-5,5,1), 0, non_zero_random(-5,5,1) ],
[ 0, 0, 1, non_zero_random(-5,5,1) ],
[ 0, 0, 0, 0 ],
);
$Areduced = Matrix(@areduced);

$A = $U * $Areduced;

($m,$n) = $A->dimensions;

$x = Matrix([
["-1*$areduced[0][3] * z -1*$areduced[0][1] * x"],
["x"],
["-1*$areduced[1][3] * z"],
["z"],
]);

Context()->texStrings;
BEGIN_TEXT
Let \( f : \mathbb{R}^{$n} \to \mathbb{R}^{$m} \) be the linear transformation defined by
\[
f(w,x,y,z) = $A 
\left\lbrack \begin{array}{c} w \\ x \\ y \\ z \end{array} \right\rbrack.
\]
Find the kernel of \( f \).  If necessary, parametrize your answer using any of the variables \( w \), \( x \), \( y \) and \( z \) that are ${LQ}free${RQ}.
$BR
$BR
\( \left\lbrack \begin{array}{c} w \\ x \\ y \\ z \end{array} \right\rbrack = \) 
\{ $x->ans_array(30) \}
END_TEXT
Context()->normalStrings;

$Zdomain = Matrix([[0],[0],[0],[0]]);
$Zrange = Matrix([[0],[0],[0]]);

ANS($x->cmp(checker => sub {
    my ( $correct, $student, $ansHash ) = @_;

    return 0 if ($ansHash->{student_formula} =~ /w/); # if it contains w, return 0
    return 0 if !($ansHash->{student_formula} =~ /x/); # if it does not contain x, return 0
    return 0 if ($ansHash->{student_formula} =~ /y/); # if it contains y, return 0
    return 0 if !($ansHash->{student_formula} =~ /z/); # if it does not contain z, return 0

    $S = Matrix($student);
    $Sx = Matrix($student)->eval(w=>0,x=>1,y=>0,z=>0);
    $Sz = Matrix($student)->eval(w=>0,x=>0,y=>0,z=>1);

    #warn $Sx;
    #warn $Sz;
    if (($Sx == $Zdomain) or ($Sz == $Zdomain)) {
      return 0;
    } elsif ( ($A * $S == $Zrange) and ($A * $Sx == $Zrange) and ($A * $Sz == $Zrange) and !( Vector($Sx)->isParallel(Vector($Sz)) ) ) {
      return 1;
    } else {
      return 0;
    }
  }
));

COMMENT('MathObject version.  Correct answer has two free variables (x and z).  Only one parametrization is marked correct (the standard parametrization with x = x and z = z that uses the reduced echelon form to solve for w and y).');

ENDDOCUMENT();
