# DESCRIPTION
# Gives a rational expression to simplify
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: multiply, divide)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(2)
## MLT(multiply_divide_a_b_two_fractions)
## KEYWORDS('rational','expressions')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","s","t","a","b","p","q");
$ch = random(0,6,2);
$v1 = $vars[$ch];
$v2 = $vars[$ch+1];
Context()->variables->are($v1=>"Real",$v2=>"Real");

$case = random(0,1,1);
$flip = random(0,1,1); # flips the fraction or not

if ($case == 0) {
  $a = random(2,4);
  do {$b = random(2,5)*random(-1,1,2)} until (gcd($a,$b) == 1);
  $c = random(2,5);
  $m = random(1,12);
  $k = random(2,4);
  do {$n = random(2,12)} until ($n != $m);
  $p = random(2,7);
  $q = random(1,$p-1);

  $e1 = $a*$c;
  $e2 = $b*$c;
  $e3 = $m*$k;
  $e4 = $c*$n;
  $e5 = $k*($a)**2;
  $e6 = $k*($b)**2;
  $p1 = $q+2;
  $num1 = "$a $v1^3 + $b $v1^2 $v2 + $e1 $v1 + $e2 $v2";
  $num2 = "$e3 $v1^{$p}";
  $den1 = ($q == 1) ? "$n $v1^3 + $e4 $v1" : "$n $v1^{$p1} + $e4 $v1^{$q}";
  $den2 = "$e5 $v1^2 - $e6 $v2^2";

  $numconst = $m/gcd($m,$n);
  $denconst = $n/gcd($m,$n);
  $powans = $p - $q;

  if ($flip == 0) {
    $disp = "\frac{".$num1."}{".$den1."}\div\frac{".$den2."}{".$num2."}";
    $numans = Formula("$numconst $v1^$powans")->reduce;
    $denans = Formula("$denconst*$a $v1 - $denconst*$b $v2")->reduce;
  } else {
    $disp = "\frac{".$den1."}{".$num1."}\div \frac{".$num2."}{".$den2."}";
    $denans = Formula("$numconst $v1^$powans")->reduce;
    $numans = Formula("$denconst*$a $v1 - $denconst*$b $v2")->reduce;
  }
} else {
  $m = random(2,12);
  do {$n = random(2,12)} until ($n != $m);
  $a = non_zero_random(-7,7);
  do {$b = non_zero_random(-7,7)} until (abs($b) != abs($a));
  do {$c = non_zero_random(-7,7)} until ($c != $a && $c != $b);

  $e1 = $m*$a;
  $e2 = $a*$b;
  $e3 = $n*$c;
  $e4 = 2*$b;
  $e5 = ($b)**2;
  $num1 = "$m $v2 + $e1";
  $num2 = "$v1 $v2 + $a $v1 - $b $v2 - $e2";
  $den1 = "$n $v1 + $e3";
  $den2 = "$v1^2 - $e4 $v1 + $e5";

  $numconst = $m/gcd($m,$n);
  $denconst = $n/gcd($m,$n);

  if ($flip == 0) {
    $disp = "\frac{".$num1."}{".$den1."}\div \frac{".$num2."}{".$den2."}";
    $numans = Formula("$numconst ($v1 - $b)")->reduce;
    $denans = Formula("$denconst ($v1+$c)")->reduce;
  } else {
    $disp = "\frac{".$den1."}{".$num1."}\div \frac{".$den2."}{".$num2."}";
    $denans = Formula("$numconst ($v1 - $b)")->reduce;
    $numans = Formula("$denconst ($v1+$c)")->reduce;
  }
}

$multians = MultiAnswer($numans, $denans)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"type an answer first");
          $self->setMessage(2,"type an answer first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

$dispfrac = ColumnTable(
  "\( \displaystyle $disp = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Perform the indicated operation and simplify if possible. Assume all cancelled factors are not zero.
$PAR
$BCENTER
$dispfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());

Context()->normalStrings;

ENDDOCUMENT();
