# DESCRIPTION
# Gives a rational expression to simplify
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: multiply, divide)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(2)
## KEYWORDS('rational','expressions')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","s","t","a","b","p","q");
$ch = random(0,7,1);
$v = $vars[$ch];
Context()->variables->are($v=>"Real");

$flip = random(0,1,1); # flips the fractions or not

$a = non_zero_random(-9,9,1);
do {$b = non_zero_random(-9,9,1)} until (abs($b) != abs($a));
$k = random(2,10)**2;

$e1 = ($a)**2;
$e2 = $b - $a;
$e3 = -$a*$b;
$e4 = -2*$a;
$e5 = $e1;
$num1 = "$v^2 - $e1";
$num2 = "$v^2 + $e2 $v + $e3";
if ($num2 == -1) {
  $num2 = "$v^2 - $v + $e3";
} elsif ($num2 == 1) {
  $num2 = "$v^2 + $v + $e3";
}
$den1 = "$k";
$den2 = "$v^2 + $e4 $v + $e5";

if ($flip == 0) {
  $disp = "\frac{".$num1."}{".$den1."}\cdot \frac{".$num2."}{".$den2."}";
  $numans = Formula("($v + $a)($v+$b)")->reduce;
  $denans = Formula("$k")->reduce;
} else {
  $disp = "\frac{".$den2."}{".$num2."}\cdot \frac{".$den1."}{".$num1."}";
  $denans = Formula("($v + $a)($v+$b)")->reduce;
  $numans = Formula("$k")->reduce;
}

$multians = MultiAnswer($numans, $denans)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"type an answer first");
          $self->setMessage(2,"type an answer first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

$dispfrac = ColumnTable(
  "\( \displaystyle $disp = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Perform the indicated operation and simplify if possible. Assume all cancelled factors are not zero.
$PAR
$BCENTER
$dispfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());

Context()->normalStrings;

ENDDOCUMENT();
