# DESCRIPTION
# Gives a rational expression to simplify
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: multiply, divide)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(2)
## KEYWORDS('rational','expressions')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","s","t","a","b","p","q");
$ch = random(0,6,2);
$v1 = $vars[$ch];
$v2 = $vars[$ch+1];
Context()->variables->are($v1=>"Real",$v2=>"Real");

$case = random(0,1,1);
$flip = random(0,1,1); # flips the signs or not

$b = random(2,4);
$c = random(2,4)*$b;
$a = random(2,4)*$c;
$p1 = random(1,2);
$p2 = $p1+random(1,3);
$p3 = $p2+random(1,3);

$sign1 = ($flip == 0) ? "\div" : "\cdot";
$sign2 = ($flip == 0) ? "\cdot" : "\div";

$disp = ($p1 == 1) ? "\frac{$a}{$v1} $sign1 \frac{$b $v1 $v2}{$v1^$p2}" : "\frac{$a}{$v1^$p1} $sign1 \frac{$b $v1^$p1 $v2^$p1}{$v1^$p2}";
$disp = $disp."$sign2 \frac{$c $v1^$p2}{$v1^$p3}";

if ($flip == 0) {
  $numconst = $a*$c/$b;
  $pow = 2*$p1+$p3-2*$p2;
  if ($pow >= 0) {
    $numans = Formula("$numconst")->reduce;
    $denans = Formula("$v1**$pow $v2^($p1)")->reduce;
  } else {
    $numans = Formula("$numconst $v1**(-$pow)")->reduce;
    $denans = Formula("$v2^($p1)")->reduce;
  }
} else {
  $numconst = $a*$b/$c;
  $pow = 2*$p2-$p3;
  if ($pow > 0) {
    $numans = Formula("$numconst $v2^($p1)")->reduce;
    $denans = Formula("$v1**($pow)")->reduce;
  } else {
    $numans = Formula("$numconst $v1^(-$pow) $v2^($p1) ")->reduce;
    $denans = Formula("1")->reduce;
  }
}

$multians = MultiAnswer($numans, $denans)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"type an answer first");
          $self->setMessage(2,"type an answer first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

$dispfrac = ColumnTable(
  "\( \displaystyle $disp = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Perform the indicated operation and simplify if possible. Assume all cancelled factors are not zero.
$PAR
$BCENTER
$dispfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());

Context()->normalStrings;
COMMENT("MathObject version");

ENDDOCUMENT();
