# DESCRIPTION
# Asks to find the difference of two rational expressions
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(2)
## KEYWORDS('least common denominator')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPowers.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");

@vars = ("x","y","z","a","b","c");
$ch = random(0,5,1);
$v = $vars[$ch];
Context()->variables->are($v=>"Real");

$a = random(1,5);
do {$b = random(1,5)} until ($b != $a);
$p1 = random(1,2);
$p2 = random(1,2);
$frac = ($a == 1) ? "$v^{-$p1} + " : "($a $v)^{-$p1}+";
if ($b == 1) {
  $frac = $frac."$v^{-$p2}";
} else {
  $frac = $frac."($b $v)^{-$p2}";
}

if ($p1 == $p2) {
  $e1 = $a**$p1 + $b**$p2;
  $e2 = ($a**$p1)*($b**$p2);
  $n = $e1/gcd($e1,$e2);
  $d = $e2/gcd($e1,$e2);
  $num = Formula("$n")->reduce;
  $den = Formula("$d*$v^$p1")->reduce;
} else {
  $d = gcd($a**$p1,$b**$p2);
  $e1 = $a**$p1/$d;
  $e2 = $b**$p2/$d;
  $e3 = $a**$p1*$b**$p2/$d;
  if ($p1 > $p2){
    $num = Formula("$e1*$v+$e2")->reduce;
    $den = Formula("$e3*$v^$p1")->reduce;
  } else {
    $num = Formula("$e2*$v+$e1")->reduce;
    $den = Formula("$e3*$v^$p2")->reduce;
  }
} 

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

LimitedPowers::OnlyNonNegativeIntegers();

###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Perform the indicated operation, and simplify if possible.  Write your answer without using negative exponents.
$BR
$BR
$BCENTER
$displayfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());

Context()->normalStrings;

ENDDOCUMENT();
