# DESCRIPTION
# Asks to find the difference of two rational expressions
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(SumThreeRationals)
## Level(2)
## KEYWORDS('least common denominator')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","z","a","b","c");
$ch = random(0,5,1);
$v = $vars[$ch];
Context()->variables->are($v=>"Real");

$case = random(0,1,1);

if ($case == 0) {
  $a = random(1,9);
  $b = non_zero_random(-9,9);
  $c = random(2,9);
  $m = random(2,9);
  do {$n = random(2,9)} until (gcd($c,$n) == 1);

  $num = Formula("$a*$n+$c - $n $v")->reduce;
  $den = Formula("$n($v +$b)")->reduce;
  $e1 = $m*$b;
  $e2 = $n*$b;
  $frac = "\frac{$a}{$v+$b} - \frac{$m $v}{$m $v + $e1} + \frac{$c}{$n $v + $e2}";
} else {
  do {
    $a = random(1,9);
    do {$m = random(2,9)} until (gcd($m,$a) == 1);
    $b = non_zero_random(-9,9);
    $c = random(1,9);
    do {$d = random(2,9)*random(-1,1,2)} until (gcd($c,$d) == 1);
    do {$n = random(2,9)} until (gcd($n,$c) == 1 || gcd($n,$d) == 1);
  } until (gcd($a*$n+$m*$c,$m*$n) == 1);
  
  $e0 = $m*$d - $m*$n;
  $num = Formula("$a*$n+$m*$c +$e0 $v")->reduce;
  $den = Formula("$m*$n($v +$b)")->reduce;
  $e1 = $m*$b;
  $e2 = $n*$b;
  $frac = "\frac{$a}{$m $v+$e1} - \frac{$v}{$v + $b} + \frac{$c+$d $v}{$n $v + $e2}";
}

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Both answer blanks must be filled");
          $self->setMessage(2,"Both answer blanks must be filled");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Perform the indicated operation, and simplify if possible.  Assume any factors you cancel are not zero.
$BR
$BR
$BCENTER
$displayfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());

Context()->normalStrings;

ENDDOCUMENT();
