## DESCRIPTION
## Higher order ODEs: constant coefficient higher order homogeneous linear ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear, constant coefficients, homogeneous (complex roots))
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('2.3')
## Problem1('6')
## KEYWORDS('differential equations','second order linear ODE','constant coefficients')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserAssignment.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup 1

SRAND($psvn);

$a = random(2,5,1);
do { $b = random(-1,1,2)*random(2,9,1); } until ($b != $a);

$c3 = 4*$a;
$c2 = 4*$a**2 + 2*($a**2 + 1);
$c1 = 4*$a*($a**2+1);
$c0 = ($a**2 + 1)**2;

Context("ImplicitEquation");
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add(
#"A"=>"Real","B"=>"Real","C"=>"Real","D"=>"Real",
#"y"=>"Real",
"y'"=>"Real","y''"=>"Real","y'''"=>"Real","y''''"=>"Real",
);

if (random(0,1,1)==0) {
  $term = "$b e^{$a x} x \cos(x)"; # tex
} else {
  $term = "$b e^{$a x} x \sin(x)"; # tex
}

$diffeq = ImplicitEquation("y'''' - $c3 y''' + $c2 y'' - $c1 y' + $c0 y = 0");


######################
#  Main text 1

Context()->texStrings;
BEGIN_TEXT
Suppose that a fourth order differential equation 
has a solution \( y = $term \).
$BR
$BR
(a) Find such a differential equation, assuming
it is homogeneous and has constant coefficients.
$BR
\{ ans_rule(40) \}
\{ AnswerFormatHelp("equations") \}
$BR
$BR
END_TEXT
Context()->normalStrings;

########################
#  Answer evaluation 1

$showPartialCorrectAnswers = 1;

ANS( $diffeq->cmp() );




########################
#  Setup 2

Context("Numeric");
Context()->variables->add(
"c1"=>"Real","c2"=>"Real","c3"=>"Real","c4"=>"Real","y"=>"Real",
);
Context()->variables->set(
'c1'=>{limits=>[2,4]},
'c2'=>{limits=>[2,4]},
'c3'=>{limits=>[2,4]},
'c4'=>{limits=>[2,4]}
);
Context()->flags->set(
    formatStudentAnswer=>'parsed',
    reduceConstants=>0,
    reduceConstantFunctions=>0,
);
parser::Assignment->Allow;


$answer = Compute("y = (c1 + c2 x) e^($a x) sin(x) + (c3 + c4 x) e^($a x) cos(x)");

#############################
#  Main text 2
Context()->texStrings;
BEGIN_TEXT
(b) Find the general solution to this 
differential equation.  
In your answer, use \( c_1, c_2, c_3 \) and \( c_4 \) to 
denote arbitrary constants and \( x \) 
the independent variable.  Enter
\( c_1 \) as ${BTT}c1${ETT}, 
\( c_2 \) as ${BTT}c2${ETT}, etc.
$BR
\{ ans_rule(60) \}
\{ AnswerFormatHelp("equations") \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation 2

ANS( $answer->cmp( # tolType=>'absolute', tolerance=>0.001, 
  checker => sub {

    my ( $correct, $student, $answerHash ) = @_;
    my $stu   = Formula($student->{tree}{rop});

    #################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('c1'))==Formula(0) || 
      Formula($stu->D('c2'))==Formula(0) || 
      Formula($stu->D('c3'))==Formula(0) ||
      Formula($stu->D('c4'))==Formula(0)
    );


    ##################################
    #  Linear independence (Wronskian)
    #
    my $x = Real(0.239);

    my $a11 = $stu->eval('c1'=>1,'c2'=>0,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a12 = $stu->eval('c1'=>0,'c2'=>1,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a13 = $stu->eval('c1'=>0,'c2'=>0,'c3'=>1,'c4'=>0,x=>$x,y=>0);
    my $a14 = $stu->eval('c1'=>0,'c2'=>0,'c3'=>0,'c4'=>1,x=>$x,y=>0);
    
    my $a21 = $stu->D('x')->eval('c1'=>1,'c2'=>0,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval('c1'=>0,'c2'=>1,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a23 = $stu->D('x')->eval('c1'=>0,'c2'=>0,'c3'=>1,'c4'=>0,x=>$x,y=>0);
    my $a24 = $stu->D('x')->eval('c1'=>0,'c2'=>0,'c3'=>0,'c4'=>1,x=>$x,y=>0);

    my $a31 = $stu->D('x','x')->eval('c1'=>1,'c2'=>0,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a32 = $stu->D('x','x')->eval('c1'=>0,'c2'=>1,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a33 = $stu->D('x','x')->eval('c1'=>0,'c2'=>0,'c3'=>1,'c4'=>0,x=>$x,y=>0);
    my $a34 = $stu->D('x','x')->eval('c1'=>0,'c2'=>0,'c3'=>0,'c4'=>1,x=>$x,y=>0);

    my $a41 = $stu->D('x','x','x')->eval('c1'=>1,'c2'=>0,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a42 = $stu->D('x','x','x')->eval('c1'=>0,'c2'=>1,'c3'=>0,'c4'=>0,x=>$x,y=>0);
    my $a43 = $stu->D('x','x','x')->eval('c1'=>0,'c2'=>0,'c3'=>1,'c4'=>0,x=>$x,y=>0);
    my $a44 = $stu->D('x','x','x')->eval('c1'=>0,'c2'=>0,'c3'=>0,'c4'=>1,x=>$x,y=>0);

    # my $wronskian = 
    #  $a11 * ( $a22*($a33*$a44-$a43*$a34) -
    #           $a23*($a32*$a44-$a42*$a34) +
    #           $a24*($a32*$a43-$a42*$a33) )
    #
    # -$a12 * ( $a21*($a33*$a44-$a43*$a34) -
    #           $a23*($a31*$a44-$a41*$a34) +
    #           $a24*($a31*$a43-$a41*$a33) )
    #
    # +$a13 * ( $a21*($a32*$a44-$a42*$a34) -
    #           $a22*($a31*$a44-$a41*$a34) +
    #           $a24*($a31*$a42-$a41*$a32) )
    #
    # -$a14 * ( $a21*($a32*$a43-$a42*$a33) -
    #           $a22*($a31*$a43-$a41*$a33) +
    #           $a23*($a31*$a42-$a41*$a32) );

    # Value->Error("Your functions are not linearly independent or your answer is not complete") 
    # if ($wronskian==Real(0));

    my $w1 = 
     $a11 * ( $a22*($a33*$a44-$a43*$a34) -
              $a23*($a32*$a44-$a42*$a34) +
              $a24*($a32*$a43-$a42*$a33) )
    +$a13 * ( $a21*($a32*$a44-$a42*$a34) -
              $a22*($a31*$a44-$a41*$a34) +
              $a24*($a31*$a42-$a41*$a32) );

    my $w2 = 
     $a12 * ( $a21*($a33*$a44-$a43*$a34) -
              $a23*($a31*$a44-$a41*$a34) +
              $a24*($a31*$a43-$a41*$a33) )
    +$a14 * ( $a21*($a32*$a43-$a42*$a33) -
              $a22*($a31*$a43-$a41*$a33) +
              $a23*($a31*$a42-$a41*$a32) );

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($w1 == $w2);




    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu1  = Formula($stu->D('x'));
    my $stu2  = Formula($stu->D('x','x'));
    my $stu3  = Formula($stu->D('x','x','x'));
    my $stu4  = Formula($stu->D('x','x','x','x'));
    return ($stu4 + $c2 * $stu2 + $c0 * $stu) == ($c3 * $stu3 + $c1 * $stu1);

    # my $stuDE = Formula($stu4 - $c3*$stu3 + $c2*$stu2 - $c1*$stu1 + $c0*$stu)
    # ->with(test_points=>[[0,0,0,0,0,0],[1,1,1,1,1,1],[1,2,-1,2,3,1]]);
    # return ($stuDE==Formula(0));


}));


;

ENDDOCUMENT();


