## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Date(6/3/2002)
## MLT(groupingSpecialF)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('59 67')
## KEYWORDS('factoring', 'identities')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

Context("Numeric");
Context()->variables->are(a => "Real", b => "Real");

$A = random(1,9,1);
$B = non_zero_random(-8,8,1);
while (gcd($A,$B) > 1 || $B==-$A) {$B = random(1,9,1);}

$C = random(1, 9,1);

$c1 = $A**2;
$c2 = 2*$A*$B;
$c3 = $B**2;
$c4 = -$C**2;

$poly=Formula("$c1*a^2 + $c2*ab + $c3*b^2 + $c4 ")->reduce;

$ans1 = Formula("$A*a+$B*b-$C");
$ans2 = Formula("$A*a+$B*b+$C");

$ans = MultiAnswer($ans1,$ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      my @score=(0,0);
      if ( $poly == $f1stu*$f2stu ){
        # check degrees
        if ($f1stu->D('a','a')->isZero && 
            $f1stu->D('a','b')->isZero &&
            $f1stu->D('b','b')->isZero ){
          $score[0]=1;
        }
        else { $self->setMessage(1,"The degree should be 1");}
        if ($f2stu->D('a','a')->isZero && 
            $f2stu->D('a','b')->isZero &&
            $f2stu->D('b','b')->isZero ){
          $score[1]=1;
        }
        else { $self->setMessage(2,"The degree should be 1");}        
      }
      return @score;
  }
);


BEGIN_PGML

The polynomial [`[$c1] a^2 + [$c2] ab + [$c3] b^2 + [$c4] `] can be factored into
a product of two polynomials,    
    [`[$poly]=A \cdot B `]    
where both [`A`] and [`B`] are polynomials of degree = 1.  

Find [`A`] and [`B`].
  
    [`A = `] [_]{$ans}{10}
  
    [`B = `] [_]{$ans}{10}
END_PGML



BEGIN_PGML_SOLUTION 
    [``\begin{aligned}
    & [$poly] \\
    &= \Big( [$c1] a^2 + [$c2] ab + [$c3] b^2 \Big) - [@ $C**2 @] \\    
    &= ([$A]a + [$B]b)^2 - [$C]^2 \end{aligned} ``]    
is the difference of two squares so it breaks down into a product   
    [``\begin{aligned}
    &([$A]a + [$B]b)^2 - [$C]^2 \\
    &=([$A]a + [$B]b + [$C])([$A]a + [$B]b - [$C]) \end{aligned} ``]
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.;