## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring polynomials: general)
## Date(6/3/2002)
## MLT(grouping)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('78')
## KEYWORDS('factoring', 'identities')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'MathObjects.pl','PGML.pl', 'parserMultiAnswer.pl','PGcourse.pl');

Context("Numeric");

$A = random(1,8,1);
$B = random(1,9,1);
while (gcd($A,$B) > 1) {$B = random(1,9,1);}

$D = random($A, 9,1);

$c1 = $B**2;
$c2 = ($B**2)*$D;
$c3 = -$A**2;
$c4 = -$A**2*$D;

$poly = Formula("$c1*x^3 + $c2*x^2 + $c3*x + $c4")->reduce;

$ans1 = Formula("$B*x - $A");
$ans2 = Formula("$B*x + $A");
$ans3 = Formula("x+$D");

$ans = MultiAnswer($ans1,$ans2,$ans3)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu, $f3stu ) = @{$student};
      my ( $f1, $f2, $f2 ) = @{$correct};
      my @score=(0,0,0);
      if ( $poly == $f1stu*$f2stu*$f3stu ){
        # check degrees
        if ($f1stu->D('x','x')->isZero && !($f1stu->D('x')->isZero) ){
          $score[0]=1;
        }
        else { $self->setMessage(1,"The degree should be 1");}
        if ($f2stu->D('x','x')->isZero && !($f2stu->D('x')->isZero) ){
          $score[1]=1;
        }
        else { $self->setMessage(2,"The degree should be 1");}        
        if ($f3stu->D('x','x')->isZero && !($f3stu->D('x')->isZero) ){
          $score[2]=1;
        }
        else { $self->setMessage(3,"The degree should be 1");}
      }
      return @score;
  }
);

BEGIN_PGML

The polynomial [`[$poly] `] can be factored into
the product of three polynomials,    
    [`[$poly]=A \cdot B \cdot C`]     
where each polynomial [`A`], [`B`], and [`C`] has degree = 1. 

Find [`A`], [`B`], and [`C`].
  
    [`A = `] [_]{$ans}{10}
  
    [`B = `] [_]{$ans}{10}
  
    [`C = `] [_]{$ans}{10}
END_PGML

BEGIN_PGML_SOLUTION
First, try factoring by grouping:    
    [``\begin{aligned}
    [$poly]&= \Big([$c1]x^3 + [$c2]x^2 \Big)+\Big([$c3]x + [$c4]) \\
    &= [$c1]x^2(x+[$D])-[@ $A**2 @](x+[$D])\\
    &= ([$c1]x^2-[@ $A**2 @])(x+[$D]) \end{aligned} ``]    
Then use the "difference of two squares" formula [`(P^2-Q^2)=(P-Q)(P+Q)`] to factor:    
    [``([$c1]x^2-[@ $A**2 @])=([$B]x-[$A])([$B]x+[$A])``]    
and plug that result into the previous equation to obtain    
    [``[$poly]=([$B]x-[$A])([$B]x+[$A])(x+[$D])``]. 

END_PGML_SOLUTION


ENDDOCUMENT();        # This should be the last executable line in the problem.;