## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring by grouping)
## Date(6/3/2002)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('74')
## KEYWORDS('factoring', 'identities')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'MathObjects.pl','PGML.pl', 'parserMultiAnswer.pl','PGcourse.pl');

Context("Numeric");

$A = random(1,9,1);
$B = random(1,9,1);
while (gcd($A,$B) > 1) {$B = random(1,9,1);}

$K = random(1,9,1);
$L = random(1,9,1);
while (gcd($K,$L)>1) {$L = random(1,9,1);}

$c1 = $K*$A;
$c2 = $K*$B;
$c3 = $L*$A;
$c4 = $L*$B;

$poly=Formula("$c1*x^3 + $c2*x^2 + $c3*x + $c4")->reduce;

$ans1 = Formula("$K*x**2 + $L")->reduce;
$ans2 = Formula("$A*x + $B")->reduce;

$ans = MultiAnswer($ans1,$ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $poly == $f1stu*$f2stu ){
        # check degrees
        if ($f1stu->D('x','x','x')==Formula(0) && $f2stu->D('x','x')==Formula(0)){
          return (1,1);
        }
        elsif ($f1stu->D('x','x')==Formula(0) && $f2stu->D('x','x','x')==Formula(0)){
           $self->setMessage(1,"Does this factor have the greatest degree?");              $self->setMessage(2,"Does this factor have the least degree?");
           return (0,0);
        }
        elsif ($f1stu->D('x')==Formula(0) || $f2stu->D('x')==Formula(0)){
          $self->setMessage(1,"Neither factor should be constant");
          $self->setMessage(2,"Neither factor should be constant");
          return (0,0)
        }
      }
      else { return (0,0); }
  }
);

BEGIN_PGML



The polynomial [`[$poly] `] can be factored into
the product of two polynomials,    
    [``[$poly]=A \cdot B``]    
where the degree of [`A`] is
greater than the degree of [`B`], and neither [`A`] nor [`B`] is a constant.  

Find [`A`] and [`B`].
  
    [`A = `] [_]{$ans}{10}
  
    [`B = `] [_]{$ans}{10}

END_PGML

BEGIN_PGML_SOLUTION
Try "factoring by grouping":

    [``\begin{aligned}
  &[$poly] \\
  &= \Big([$c1]x^3 + [$c2]x^2\Big) + \Big([$c3]x + [$c4]\Big) \\
  &=[$K]x^2\Big([$ans2]\big)+[$L]\Big([$ans2]\big)\\
  &=\big([$ans1]\Big)\Big([$ans2]\Big) \end{aligned} ``]
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.;