## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring by grouping)
## Date(6/3/2002)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## MLT(grouping_multi_variable)
## EditionText1('3')
## Section1('1.4')
## Problem1('59')
## KEYWORDS('factoring', 'identities')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'MathObjects.pl', 'PGML.pl', 'parserMultiAnswer.pl','PGcourse.pl');

Context("Numeric");
Context()->variables->add(y=>"Real");

$A = random(1,9,1);
$B = random(1,8,1);
while (gcd($A,$B) > 1) {$B = random(1,9,1);}

$C = random($B, 9,1);

$c1 = $A**2;
$c2 = -$B**2;
$c3 = $C*$A**2;
$c4 = -$B**2*$C;

$poly =Formula("$c1*x^3 + $c2*xy^2 + $c3*x^2 + $c4*y^2")->reduce;

$ans1 = Formula("$A*x - $B*y");
$ans2 = Formula("x+$C");
$ans3 = Formula("$A*x + $B*y");

$ans = MultiAnswer($ans1,$ans2,$ans3)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu, $f3stu ) = @{$student};
      my ( $f1, $f2, $f2 ) = @{$correct};
      my @score=(0,0,0);
      if ( $poly == $f1stu*$f2stu*$f3stu ){
        # check degrees
        if ($f1stu->D('x','x')->isZero && 
            $f1stu->D('x','y')->isZero &&
            $f1stu->D('y','y')->isZero ){
          $score[0]=1;
        }
        else { $self->setMessage(1,"The degree should be 1");}
        if ($f2stu->D('x','x')->isZero && 
            $f2stu->D('x','y')->isZero &&
            $f2stu->D('y','y')->isZero ){
          $score[1]=1;
        }
        else { $self->setMessage(2,"The degree should be 1");}        
        if ($f2stu->D('x','x')->isZero && 
            $f2stu->D('x','y')->isZero &&
            $f2stu->D('y','y')->isZero ){
          $score[2]=1;
        }
        else { $self->setMessage(3,"The degree should be 1");}
      }
      return @score;
  }
);

BEGIN_PGML

The polynomial [`[$c1] x^3 + [$c2] xy^2 + [$c3] x^2 + [$c4] y^2`] can be factored into a product of three polynomials,    
    [`[$poly]=A \cdot B \cdot C`]    
where each factor [`A`], [`B`], and [`C`] is a polynomial of degree = 1.

Find [`A`], [`B`] and [`C`]
  
    [`A = `] [_]{$ans}{10}
  
    [`B = `] [_]{$ans}{10}
  
    [`C = `] [_]{$ans}{10}
END_PGML

BEGIN_PGML_SOLUTION
    [``\begin{aligned}
      &[$poly]\\
      &= \Big([$c1] x^3 + [$c2] xy^2\Big) + \Big([$c3] x^2 + [$c4] y^2\Big) \\
      &= x ([$c1] x^2 + [$c2] y^2) + [$C]([$c1] x^2 + [$c2] y^2) \\
      &= (x+[$C])([$c1] x^2 + [$c2] y^2) \\
      &= (x+[$C])([$A]x-[$B]y)([$A]x+[$B]y) \end{aligned} ``]
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.;