## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Date(6/3/2002)
## MLT(simpleDiffSq)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('51')
## KEYWORDS('algebra, factoring, difference of squares')

DOCUMENT();     
loadMacros('PGstandard.pl', 'MathObjects.pl','PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

Context("Numeric");

$a1=Real(random(2,7,1));
$a2=Real(random(2,5,1));
$a3=$a1**2;
$a4=$a2**2;

$poly = Formula("$a3-$a4*x**2");

$ans = MultiAnswer($a1, $a2, $a1, $a2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $A, $B, $C, $D ) = @{$student};
      my $f1 = Formula("$A - $B*x");
      my $f2 = Formula("$A + $D*x");
      return $poly == $f1*$f2 ? 1 : 0;
  }
);


BEGIN_PGML

Factor [` [$a3]-[$a4] x^2 `]:    
    [` [$a3]-[$a4] x^2=( A-Bx)(C+ Dx) `]    
where     
    A=[_]{$a1}, B=[_]{$a2}, C=[_]{$a1}, D=[_]{$a2}
END_PGML

BEGIN_PGML_SOLUTION
[``[$poly]``] is the difference of two squares:    
    [``[$poly]=[$a1]^2-([$a2]x^2)=([$a1]-[$a2]x)([$a1]+[$a2]x)``]
END_PGML_SOLUTION
ENDDOCUMENT();