## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring polynomials: general)
## Date(6/3/2002)
## MLT(trinomials)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('69')
## KEYWORDS('algebra', 'factoring')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

Context("Numeric");

$m = Real(random(2,5,1));
$a = Real(random(1,4,1));
$b = Real(random(5,8,1));
$m1 = $m+1;
$m2 = $m+2;
$t = $b-$a;
$n = $a*$b;

$ab = MultiAnswer(-$a,$b)->with(
    singleResult => 0,
    checker      => sub {
        my ($correct, $student, $self) = @_;
        my ($f1stu, $f2stu) = @{$student};
        my ($f1,    $f2)    = @{$correct};
        if ($f1 == $f1stu && $f2 == $f2stu)
        {
            return [ 1, 1 ];
        }
        if ($f1 == $f2stu && $f2 == $f1stu)
        {  
            $self->setMessage(1,"Is A < B?");
        }
        return [0,0];
    }
);

BEGIN_PGML;
Factor the polynomial [`t^[$m2]+[$t] t^[$m1]-[$n] t^[$m]`]. The answer can be written as    
    [``t^[$m2]+[$t] t^[$m1]-[$n] t^[$m]=t^N(t+A)(t+B)``]     
where [`A < B`] and  

    [`N=`][_]{$m}, [`A=`][_]{$ab}, [`B=`][_]{$ab}
END_PGML

BEGIN_PGML_SOLUTION
    [``t^N(t+A)(t+B) = t^{N+2}+(A+B)t^{N+1}+ABt^N``].    
This equals [`t^[$m2]+[$t] t^[$m1]-[$n] t^[$m]`] if and only if    
    [`N=[$m]`] and [`A+B=[$t]`] and [`AB=-[$n]`].    
First let's try to find _integers_ [`A`] and [`B`] that satisfy the last two equations.  (Only a few pairs of integers [`A,B`] satisfy [`AB=-[$n]`]; if none of them also satisfy [`A+B=[$t]`] then we'll have to try something fancier.)  Fortunately [`A=-[$a]`] and [`B=[$b]`] is such a pair.  Also, with these values, [`A<B`] so the solution is    
    [`N=[$m]`] and [`A=-[$a]`] and [`B=[$b]`].
END_PGML_SOLUTION


ENDDOCUMENT();        # This should be the last executable line in the problem.