## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring polynomials: general)
## Date(6/3/2002)
## MLT(trinomials)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('73')
## KEYWORDS('algebra', 'factoring')

DOCUMENT();        # This should be the first executable line in the problem.    

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl','PGcourse.pl');

Context("Numeric");

$c = Real(random(3,9,1));
$t = -1+$c;

$ab = MultiAnswer(Real(-1),Real(1))->with(
    singleResult => 0,
    checker      => sub {
        my ($correct, $student, $self) = @_;
        my ($f1stu, $f2stu) = @{$student};
        my ($f1,    $f2)    = @{$correct};
        if ($f1 == $f1stu && $f2 == $f2stu)
        {
            return [ 1, 1 ];
        }
        if ($f1 == $f2stu && $f2 == $f1stu)
        {  
            $self->setMessage(1,"Is A < B?");
        }
        return [0,0];
    }
);
BEGIN_PGML;
Factor the polynomial [`x^4+[$t] x^2-[$c]`]. The answer can be written as   

    [`x^4+[$t] x^2-[$c]=(x+A)(x+B)(x^2+C)`]    

where [` A`], [`B`], [`C `] are integers and [`A<B`].

    [`A=`] [_]{$ab}, [`B=`] [_]{$ab}, [`C=`] [_]{$c} 
END_PGML

BEGIN_PGML_SOLUTION
First, try to factor it as a product of two factors
    [`x^4+[$t] x^2-[$c] = (x^2+M)(x^2+N)`]    
where [`M`] and [`N`] are integers.    
    [`(x^2+M)(x^2+N)=x^4+(M+N)x^2+MN`].    
This equals [`x^4+[$t] x^2-[$c]`] if and only if     
    [`M+N=[$t]`] and [`MN=-[$c]`].   
[`M=-1`] and [`N=[$c]`] satisfy these equations so
    [`x^4+[$t] x^2-[$c] = (x^2-1)(x^2+[$c])`].    
The first term on the right is the difference of two squares so it factors       
    [`x^2-1=(x-1)(x+1)`].    
Plug that result into the previous equation to obtain the answer:
    [`x^4+[$t] x^2-[$c] = (x-1)(x+1)(x^2+[$c])`].

END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.