## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: common factors)
## MO(1)
## Date(6/3/2002)
## Level(3)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('55')
## TitleText2('Precalculus')
## AuthorText2('Stewart, Redlin, Watson')
## EditionText2('3')
## Section2('1.3')
## Problem2('57')
## KEYWORDS('algebra', 'factoring')
## MLT(common_factors_clumsy)

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

$showPartialCorrectAnswers = 1;

Context("Numeric");

$a = Real(random(1,5,2));
$b = Real(random(2,6,2));

$multians = MultiAnswer(-$a ,$b)->with(
  singleResult => 0,
  checker => sub {
    my( $correct, $student, $self ) = @_;
    my( $corB, $corC ) = @{$correct};
    my( $stuB, $stuC ) = @{$student};
    if ( ($corB == $stuB and $corC == $stuC) or 
         ($corB == $stuC and $corC == $stuB)
       ){ 
         return [1,1];
    } else {
      if ( $corB == $stuB or $corC == $stuB ){
         return [1,0];
      } elsif  ( $corC == $stuC or $corB == $stuC ){
         return [0,1];
      } else {
        return [0,0];
      }
    }
    return [1,1];
  }
);

BEGIN_PGML
Factor the polynomial [`(x-[$a])(x+[$b])^2-(x-[$a])^2(x+[$b])`]. 
The answer can be written as    
    [`(x-[$a])(x+[$b])^2-(x-[$a])^2(x+[$b])=A(x+B)(x+C)`]     
with integers [` A`], [`B`], [`C `], where

    [`A=`][_]{-$a-$b}{10}, and    
    [`B=`][_]{$multians}{10}, and    
    [`C=`][_]{$multians}{10}
END_PGML

# ANS( ($a+$b)->cmp() );
# ANS( $multians->cmp() );

BEGIN_PGML_SOLUTION
Factor out the common terms    
    [``\begin{aligned} &(x-[$a])(x+[$b])^2-(x-[$a])^2(x+[$b]) \\
&= (x-[$a])(x+[$b])\Big((x-[$a])-(x+[$b])\Big) \\
&=(x-[$a])(x+[$b])(-[$a]-[$b])\\
&=([@ -$a-$b @])(x-[$a])(x+[$b]) \end{aligned} ``]
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.