## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: common factors)
## MO(1)
## Date(6/3/2002)
## Level(3)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('55')
## TitleText2('Precalculus')
## AuthorText2('Stewart, Redlin, Watson')
## EditionText2('3')
## Section2('1.3')
## Problem2('57')
## KEYWORDS('algebra', 'factoring')
## MLT(common_factors_clumsy)


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric");

$a = Real(random(1,5,2));
$b = Real(random(2,6,2));

$multians = MultiAnswer(-$a ,$b)->with(
  singleResult => 0,
  checker => sub {
    my( $correct, $student, $self ) = @_;
    my( $corB, $corC ) = @{$correct};
    my( $stuB, $stuC ) = @{$student};
    if ( ($corB == $stuB and $corC == $stuC) or 
         ($corB == $stuC and $corC == $stuB)
       ){ 
         return [1,1];
    } else {
      if ( $corB == $stuB or $corC == $stuB ){
         return [1,0];
      } elsif  ( $corC == $stuC or $corB == $stuC ){
         return [0,1];
      } else {
        return [0,0];
      }
    }
    return [1,1];
  }
);

Context()->texStrings;
BEGIN_TEXT
Factor the polynomial \((x-$a)(x+$b)^2-(x-$a)^2(x+$b)\). 
Your answer can be written as
\(A(x+B)(x+C)\) with integers \( A\), \(B\), \(C \) 
$PAR where \(A=\)\{ans_rule(10) \}, 
\(B=\)\{ $multians->ans_rule(10) \}, and 
\(C=\)\{ $multians->ans_rule(10) \}
END_TEXT

Context()->normalStrings;

ANS( ($a+$b)->cmp() );
ANS( $multians->cmp() );

ENDDOCUMENT();        # This should be the last executable line in the problem.
