## DESCRIPTION
## Algebra: Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Date(6/3/2002)
## MLT(diffCubes)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('95')
## KEYWORDS('factoring', 'identities')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl','PGcourse.pl');

Context("Numeric");
Context()->variables->add(y=>"Real");

$A = random(1,9,1);
$B = random(1,9,1);
while (gcd($A,$B) > 1) {$B = random(1,9,1);}

$K = random(1,9,1);
$L = random(1,9,1);
while (gcd($K,$L)>1) {$L = random(1,9,1);}

$c1 = $A**3;
$c2 = $B**3;
$e1 = $K*3;
$e2 = $L*3;

$poly = Formula("$c1*x^$e1 - $c2*y^$e2 ")->reduce;

$ans1 = Formula("$A**2 * x**(2*$K) + $A*$B*x**$K*y**$L +($B**2)*y**(2*$L)");
$ans2 = Formula("$A*x**$K  - $B*y**$L");

$AB=MultiAnswer($ans1,$ans2)->with(
  singleResult => 0,
  checker => sub {
     my ($correct, $student, $self) = @_;
     my ($f1stu, $f2stu) = @{$student};
     my ($f1,    $f2)    = @{$correct};
     if ($f1stu->D('x')->isZero && $f1stu->D('y')->isZero){
       $self->setMessage(1,"Constant polynomials aren't allowed here");
       return (0,0);
     }
     if ($f2stu->D('x')->isZero && $f2stu->D('y')->isZero){
       $self->setMessage(2,"Constant polynomials aren't allowed here");
       return (0,0);
     }
     if ($poly == $f1stu*$f2stu){
       return (1,1);
     }
     elsif ($f1==$f1stu or $f2==$f1stu){
       return (1,0);
     }
     elsif ($f1=$f2stu or $f2==$f2stu){
       return (0,1);
     }
     else { return (0,0); }
   }
);

BEGIN_PGML
The polynomial [`[$poly] `] can be factored into the product of two nonconstant polynomials, 
    [`[$poly] = A \cdot B`]    
Find [`A`] and [`B`].
  
    [`A = `] [_]{$AB}{30}
  
    [`B = `] [_]{$AB}{30}

END_PGML

BEGIN_PGML_HINT
This problem asks you to factor the difference of two terms.  Check to see if the terms are both perfect squares or both perfect cubes.  You know, or can look up, formulae to 
factor each of those cases.
END_PGML_HINT

$f=Formula("$A*x^$K")->reduce;
$g=Formula("$B*y^$L")->reduce;

BEGIN_PGML_SOLUTION
[``[$poly] =([$f])^3-([$g])^3``] is the difference of two cubes, so one may use the "difference of cubes" formula   
    [``P^3-Q^3=(P-Q)(P^2+PQ+Q^2)``]    
to factor it    
    [``\begin{aligned}
    [$poly]
    &=([$f])^3-([$g])^3\\
    &=([$f]-[$g])\Big(([$f])^2 + ([$f])([$g])+([$g])^2\Big) \end{aligned}``]
END_PGML_SOLUTION
ENDDOCUMENT();        # This should be the last executable line in the problem.;