##DESCRIPTION
##  Functions: Difference Quotient
##ENDDESCRIPTION

##KEYWORDS('algebra', 'functions', 'difference quotient')

## DBsubject('Algebra')
## DBchapter('Functions')
## DBsection('Combining Functions')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('2.2')

########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PeriodicRerandomization.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

#######################################
#  Set-up
Context()->variables->add( h=>"Real");

$c = random(-1,1,2)*random(2,5,1);

if ($c > 0) {
  $f = "\frac{$c}{x}";
} else {
  $f = "-\frac{".abs($c)."}{x}";
}

$num = Formula("-$c")->reduce;
$den = Compute(" x(x+h) ")->reduce;
if ($c > 0) {
  $texanswer = "\(\displaystyle -\frac{$c}{x(x+h)} \)";
} else {
  $texanswer = "\(\displaystyle \frac{".abs($c)."}{x(x+h)} \)";
}

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

ANS($multians->cmp() );

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ 
\frac{f(x + h) - f(x)}{h} = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle \frac{f(x+h)-f(x)}{h} = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


 

Context("Numeric");
################################################
#  Main text

BEGIN_TEXT
Find the difference quotient for the function \(f(x) = $f\).  Simplify your answer as much as possible.
$PAR
$displayfrac
END_TEXT


if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    $texanswer
    $PAR
    END_TEXT
}

Context()->normalStrings;

PeriodicStatus(); 


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;


COMMENT('MathObject version', 'Features PeriodicRerandomization');
ENDDOCUMENT();
