##DESCRIPTION
##  Solving Linear Equations
##ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Linear equations)
## Date(Spring 2012)
## Institution(University of Missouri)
## MLT(Variable_intermsof_variables)
## Level(3)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.5')
## Problem1('')
## KEYWORDS('linear', 'equations')
## MO(1)

##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


Context("Numeric");
Context()->variables->are(A=>"Real", a=>"Real", B=>"Real", b=>"Real", C=>"Real", c=>"Real", s=>"Real", S=>"Real");

$var = "A";
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(2,10,1);
$b = random(-1,1,2)*random(2,9,1);
do {$c = random(-1,1,2)*random(2,9,1)} until (gcd(gcd($a,$b),$c) ==1);

$equation = "S = $a AB + $b BC + $c AC";
$frac = "$var";

$num = Formula("S - $b BC")->reduce;
$den = Formula("$a B + $c C")->reduce;
$texanswer = "\( \dfrac{S - $b BC}{$a B + $c C}\)";

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Solve the equation for \($var\):
$PAR
\[ $equation \]
$PAR
$displayfrac
END_TEXT

    BEGIN_TEXT
    $PAR
    $BBOLD Note: $EBOLD
    Write your answer as a single fraction, and type the numerator and denominator separately in the answer blanks provided.
    END_TEXT

Context()->normalStrings;


###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp(studentsMustReduceFractions=>1));

ENDDOCUMENT();
