## DESCRIPTION
## Rational Expressions: Canceling Common Factors
## ENDDESCRIPTION

## KEYWORDS('Algebraic fractions')

## DBsubject('Algebra') 
## DBchapter('Basic Algebra') 
## DBsection('Algebraic Fractions') 
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.6')
## Problem1('')
## Adapted from a problem written by Paul Pearson (Fort Lewis)


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");


###################################
# Setup

Context("Numeric");

Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real");

$var = list_random('x','y','z','w','p','q','r','t','u','v');
$var1 = list_random('x','y','z',,'w','p','q');
$var2 = list_random('r','s','t','u','v');
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = list_random(2,3,4,5);
$b = list_random(-6,-5,-4,-3,-2,2,3,4,5,6);
$c = list_random(-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6,7,8);
do {$d = list_random(-5,-4,-3,-2,2,3,4,5)} until ($d != $b);
$asquared = $a**2;
$e1 = $a*($b + $c);
$e2 = $a*$b*$c;
$e3 = $asquared*($d + $c);
$e4 = $asquared*$d *$c;

if ($e1 == 0 && $e3 != 0) {
  $frac = "\frac{ $a $var^2 + $e2}{$asquared $var^2 + $e3 $var + $e4 }";
} elsif ($e1 != 0 && $e3 == 0) {
  $frac = "\frac{$a $var^2 + $e1 $var + $e2}{$asquared $var^2 + $e4 }";
} elsif ($e1 == 0 && $e3 == 0) {
  $frac = "\frac{ $a $var^2 + $e2}{$asquared $var^2 + $e4 }";
} else {
  $frac = "\frac{$a $var^2 +$e1 $var +$e2}{$asquared $var^2 +$e3 $var + $e4 }";
}

$num = Formula(" $var + $b ")->reduce;
$den = Formula(" $a($var + $d) ")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}



###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Simplify.  Assume any factors you cancel
are not zero.
$BR
$BR
$BCENTER
$displayfrac
$ECENTER

END_TEXT

if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    \(\dfrac{$num}{$den} \)
    $PAR
    END_TEXT
}

Context()->normalStrings;

PeriodicStatus();



###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

COMMENT('MathObject version');
ENDDOCUMENT();
