## DBsubject(Abstract algebra)
## DBchapter(Groups)
## DBsection(Subgroups)
## Institution(University of Massachusetts Amherst)
## Author(Siman Wong and Patrick O'Brien)
## Level(3)
## MO(1)
## KEYWORDS('subgroups')



DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

$r = random(0,26,1);
$f = random(0,26,1);
while($f == $r){
    $f = random(0,26,1);
}

if($r == 0){
    $n = 160;
    $b = 3;
    @seq = (1, 3, 9, 27, 81, 83, 89, 107);
}
elsif($r == 1){
    $n = 162;
    $b = 19;
    @seq = (1, 19, 37, 55, 73, 91, 109, 127, 145);
}
elsif($r == 2){
    $n = 164;
    $b = 3;
    @seq = (1, 3, 9, 27, 81, 79, 73, 55);
}
elsif($r == 3){
    $n = 165;
    $b = 4;
    @seq = (1, 4, 16, 64, 91, 34, 136, 49, 31, 124);
}
elsif($r == 4){
    $n = 170;
    $b = 9;
    @seq = (1, 9, 81, 49, 101, 59, 21, 19);
}
elsif($r == 5){
    $n = 171;
    $b = 4;
    @seq = (1, 4, 16, 64, 85, 169, 163, 139, 43);
}
elsif($r == 6){
    $n = 172;
    $b = 21;
    @seq = (1, 21, 97, 145, 121, 133, 41);
}
elsif($r == 7){
    $n = 174;
    $b = 7;
    @seq = (1, 7, 49, 169, 139, 103, 25);
}
elsif($r == 8){
    $n = 175;
    $b = 6;
    @seq = (1, 6, 36, 41, 71, 76, 106, 111, 141, 146);
}
elsif($r == 9){
    $n = 176;
    $b = 7;
    @seq = (1, 7, 49, 167, 113, 87, 81, 39, 97, 151);
}
elsif($r == 10){
    $n = 178;
    $b = 37;
    @seq = (1, 37, 123, 101, 177, 141, 55, 77);
}
elsif($r == 11){
    $n = 183;
    $b = 20;
    @seq = (1, 20, 34, 131, 58, 62, 142, 95, 70, 119);
}
elsif($r == 12){
    $n = 185;
    $b = 16;
    @seq = (1, 16, 71, 26, 46, 181, 121, 86, 81);
}
elsif($r == 13){
    $n = 186;
    $b = 23;
    @seq = (1, 23, 157, 77, 97, 185, 163, 29, 109, 89);
}
elsif($r == 14){
    $n = 187;
    $b = 16;
    @seq = (1, 16, 69, 169, 86, 67, 137, 135, 103, 152);
}
elsif($r == 15){
    $n = 189;
    $b = 4;
    @seq = (1, 4, 16, 64, 67, 79, 127, 130, 142);
}
elsif($r == 16){
    $n = 190;
    $b = 61;
    @seq = (1, 61, 111, 121, 161, 131, 11, 101, 81);
}
elsif($r == 17){
    $n = 192;
    $b = 7;
    @seq = (1, 7, 49, 151, 97, 103, 145, 55);
}
elsif($r == 18){
    $n = 194;
    $b = 33;
    @seq = (1, 33, 119, 47, 193, 161, 75, 147);
}
elsif($r == 19){
    $n = 196;
    $b = 29;
    @seq = (1, 29, 57, 85, 113, 141, 169);
}
elsif($r == 20){
    $n = 198;
    $b = 17;
    @seq = (1, 17, 91, 161, 163, 197, 181, 107, 37, 35);
}
elsif($r == 21){
    $n = 200;
    $b = 9;
    @seq = (1, 9, 81, 129, 161, 49, 41, 169, 121, 89);
}
elsif($r == 22){
    $n = 202;
    $b = 17;
    @seq = (1, 17, 87, 65, 95, 201, 185, 115, 137, 107);
}
elsif($r == 23){
    $n = 203;
    $b = 36;
    @seq = (1, 36, 78, 169, 197, 190, 141);
}
elsif($r == 24){
    $n = 204;
    $b = 19;
    @seq = (1, 19, 157, 127, 169, 151, 13, 43);
}
elsif($r == 25){
    $n = 205;
    $b = 3;
    @seq = (1, 3, 9, 27, 81, 38, 114, 137);
}
elsif($r == 26){
    $n = 209;
    $b = 18;
    @seq = (1, 18, 115, 189, 58, 208, 191, 94, 20, 151);
}

$answer = join(", ",@seq);
$answer = Compute("$answer");

if($f == 0){
    $nn = 160;
    $bb = 3;
    @seqq = (1, 3, 9, 27, 81, 83, 89, 107);
}
elsif($f == 1){
    $nn = 162;
    $bb = 19;
    @seqq = (1, 19, 37, 55, 73, 91, 109, 127, 145);
}
elsif($f == 2){
    $nn = 164;
    $bb = 3;
    @seqq = (1, 3, 9, 27, 81, 79, 73, 55);
}
elsif($f == 3){
    $nn = 165;
    $bb = 4;
    @seqq = (1, 4, 16, 64, 91, 34, 136, 49, 31, 124);
}
elsif($f == 4){
    $nn = 170;
    $bb = 9;
    @seqq = (1, 9, 81, 49, 101, 59, 21, 19);
}
elsif($f == 5){
    $nn = 171;
    $bb = 4;
    @seqq = (1, 4, 16, 64, 85, 169, 163, 139, 43);
}
elsif($f == 6){
    $nn = 172;
    $bb = 21;
    @seqq = (1, 21, 97, 145, 121, 133, 41);
}
elsif($f == 7){
    $nn = 174;
    $bb = 7;
    @seqq = (1, 7, 49, 169, 139, 103, 25);
}
elsif($f == 8){
    $nn = 175;
    $bb = 6;
    @seqq = (1, 6, 36, 41, 71, 76, 106, 111, 141, 146);
}
elsif($f == 9){
    $nn = 176;
    $bb = 7;
    @seqq = (1, 7, 49, 167, 113, 87, 81, 39, 97, 151);
}
elsif($f == 10){
    $nn = 178;
    $bb = 37;
    @seqq = (1, 37, 123, 101, 177, 141, 55, 77);
}
elsif($f == 11){
    $nn = 183;
    $bb = 20;
    @seqq = (1, 20, 34, 131, 58, 62, 142, 95, 70, 119);
}
elsif($f == 12){
    $nn = 185;
    $bb = 16;
    @seqq = (1, 16, 71, 26, 46, 181, 121, 86, 81);
}
elsif($f == 13){
    $nn = 186;
    $bb = 23;
    @seqq = (1, 23, 157, 77, 97, 185, 163, 29, 109, 89);
}
elsif($f == 14){
    $nn = 187;
    $bb = 16;
    @seqq = (1, 16, 69, 169, 86, 67, 137, 135, 103, 152);
}
elsif($f == 15){
    $nn = 189;
    $bb = 4;
    @seqq = (1, 4, 16, 64, 67, 79, 127, 130, 142);
}
elsif($f == 16){
    $nn = 190;
    $bb = 61;
    @seqq = (1, 61, 111, 121, 161, 131, 11, 101, 81);
}
elsif($f == 17){
    $nn = 192;
    $bb = 7;
    @seqq = (1, 7, 49, 151, 97, 103, 145, 55);
}
elsif($f == 18){
    $nn = 194;
    $bb = 33;
    @seqq = (1, 33, 119, 47, 193, 161, 75, 147);
}
elsif($f == 19){
    $nn = 196;
    $bb = 29;
    @seqq = (1, 29, 57, 85, 113, 141, 169);
}
elsif($f == 20){
    $nn = 198;
    $bb = 17;
    @seqq = (1, 17, 91, 161, 163, 197, 181, 107, 37, 35);
}
elsif($f == 21){
    $nn = 200;
    $bb = 9;
    @seqq = (1, 9, 81, 129, 161, 49, 41, 169, 121, 89);
}
elsif($f == 22){
    $nn = 202;
    $bb = 17;
    @seqq = (1, 17, 87, 65, 95, 201, 185, 115, 137, 107);
}
elsif($f == 23){
    $nn = 203;
    $bb = 36;
    @seqq = (1, 36, 78, 169, 197, 190, 141);
}
elsif($f == 24){
    $nn = 204;
    $bb = 19;
    @seqq = (1, 19, 157, 127, 169, 151, 13, 43);
}
elsif($f == 25){
    $nn = 205;
    $bb = 3;
    @seqq = (1, 3, 9, 27, 81, 38, 114, 137);
}
elsif($f == 26){
    $nn = 209;
    $bb = 18;
    @seqq = (1, 18, 115, 189, 58, 208, 191, 94, 20, 151);
}

$answerr = join(", ",@seqq);
$answerr = Compute("$answerr");

BEGIN_TEXT

Find all elements of the subgroup \(\langle $b \rangle\) in \(U($n)\): $BR $BR

\{ans_rule(40)\} $BR $BR

Find all elements of the subgroup \(\langle $bb \rangle\) in \(U($nn)\): $BR $BR

\{ans_rule(40)\}

$BR$BR

For both parts, enter your answer as a comma-separated list.

END_TEXT


ANS( $answer->cmp( checker => sub {
		my ( $correct, $student, $ansHash ) = @_;  # get correct and student MathObjects
		return ( $student % $n == $correct % $n ? 1 : 0 );
	}
) );

ANS( $answerr->cmp( checker => sub {
		my ( $correct, $student, $ansHash ) = @_;  # get correct and student MathObjects
		return ( $student % $nn == $correct % $nn ? 1 : 0 );
	}
) );







ENDDOCUMENT();       
