## DBsubject(Abstract algebra)
## DBchapter(Groups)
## DBsection(Subgroups)
## Institution(University of Massachusetts Amherst)
## Author(Siman Wong and Patrick O'Brien)
## Level(4)
## MO(1)
## KEYWORDS('subgroups, generators')



DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 0;

TEXT(beginproblem());

@chooseModulus = ( 7, 8, 9, 12, 14, 18 );

# for each modulus, we have precomputed the valid pairs of elements that satisfy the problem
%validPairs = (	7 =>	[ [ 2, 6 ], [ 4, 6 ] ],
				8 => 	[ [ 3, 5 ], [ 3, 7 ], [ 5, 7 ] ],
				9 => 	[ [ 4, 8 ], [ 7, 8 ] ],
				12 =>	[ [ 5, 7 ], [ 5, 11 ], [ 7, 11 ] ],
				14 =>	[ [ 9, 13 ], [ 11, 13 ] ],
				18 =>	[ [ 13, 17 ], [ 7, 17 ] ]
			);
			
# choose a random modulus
$n = $chooseModulus[ random( 0, @chooseModulus - 1, 1 ) ];

# make a hash that will be used to check and see if the student's answer is on the list of correct pairs. $modulusPairs{ x }{ y } = 1 iff <x>,<y> are proper subgroups of U(n) and <x,y> = U(n).
$modulusPairs{ $$_[ 0 ] }{ $$_[ 1 ] } = $modulusPairs{ $$_[ 1 ] }{ $$_[ 0 ] } = 1 foreach @{ $validPairs{ $n } };

# multiAnswer that checks the student's pair against the hash
$xyAnswer = MultiAnswer( map{ Compute( $_ ) } @{ $validPairs{ $n }[ 0 ] } )->with(
	singleResult => 1,
	checker => sub
	{
		my( $correct, $student, $self ) = @_;
		my $studentPair = [ map{ $_->value } @$student ];
		
		return $modulusPairs{ $$studentPair[ 0 ] }{ $$studentPair[ 1 ] } == 1 ? 1 : 0;
	}
);

Context()->texStrings;



BEGIN_TEXT

Find one pair of elements \(x,\) \(y\) in \(U($n)\) such that \(\langle x \rangle\) and \(\langle y \rangle\) are proper subgroups of \(U($n)\) and that \(\langle x,y \rangle = U($n)\). Be sure that \(x<y\) with \(1 \leq x, y<n\).$BR $BR

\( (x,y) = ( \) \{ $xyAnswer->ans_rule(2) \} \( , \) \{ $xyAnswer->ans_rule( 2 ) \} \( ) \) $BR

END_TEXT

Context()->normalStrings;



ANS( $xyAnswer->cmp() );

ENDDOCUMENT();       
