## DESCRIPTION
## Higher order ODEs: the Wronskian
## ENDDESCRIPTION

## KEYWORDS('differential equations','second order','Wronskian')

## DBsubject('Differential Equations')
## DBchapter('Higher Order Differential Equations')
## DBsection('Second Order Linear')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Notes on Diffy Qs')
## EditionText1('December 9, 2010')
## AuthorText1('Jiri Lebl')
## Section1('2.1')
## Problem1('1')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserPopUp.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

install_weighted_grader();


#############################
#  Setup

Context("Numeric");

$a = random(2,6,1);
do { $b = random(2,6,1); } until ($b != $a);

$a11 = Formula("sin($a x)");     $a12 = Formula("cos($b x)");
$a21 = Formula("$a cos($a x)");  $a22 = Formula("-$b sin($b x)");

$wronskian = $a11 * $a22 - $a12 * $a21; 

$pop = PopUp(["Choose","some","all"],"some");


$multians1 = MultiAnswer($a11, $a12, $a21, $a22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Use the Wronskian to determine whether the functions
\( y_1 = $a11 \) and \( y_2 = $a12 \)
are linearly independent.
$PAR
$BCENTER
\{ 
mbox(
  "Wronskian = \( \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  " = ".$SPACE.ans_rule(20)
);
\}
$ECENTER
$PAR
These functions are linearly independent because
the Wronskian is nonzero for \{ $pop->menu() \} value(s) of \(x\).
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 50 );

WEIGHTED_ANS( $wronskian->cmp(), 40 );

WEIGHTED_ANS( $pop->cmp(), 10 );

COMMENT("MathObject version.  Uses parserMultiAnswer.pl to format the matrix nicely.");

ENDDOCUMENT();
