## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Inverses)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.3')
## Problem1('')
## KEYWORDS('linear algebra','matrix','inverse')


DOCUMENT() ;        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 1 ;

############

TEXT(beginproblem());

Context('Matrix');

do { 

$A = Matrix([
[random(-6,6,1),random(-6,6,1)],
[random(-6,6,1),random(-6,6,1)],
]);

} until ( ($A->det()) != 0);


$I = Matrix([[1,0],[0,1]]);

$answer = $A->inverse();


Context()->texStrings;
BEGIN_TEXT
Find a \( 2 \times 2 \) matrix such that 
$BR
$BR
$BCENTER
\( $A \) \{ $answer->ans_array() \} \( = $I\).
$ECENTER
END_TEXT
Context()->normalStrings;


ANS( $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      $stu = Matrix($student);
      return $A * $stu == $I;
    }
) );

;

ENDDOCUMENT();
