## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Matrix algebra)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(4)
## Static(1)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.2')
## Problem1('')
## KEYWORDS('linear algebra','matrix','linear system')

DOCUMENT() ;        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 1 ;

############

TEXT(beginproblem());

Context('Matrix');

$answer = Matrix([
[0,1],
[0,0]
]);

$Z = Matrix([[0,0],[0,0]]);

Context()->texStrings;
BEGIN_TEXT
Find a non-zero \( 2 \times 2 \) matrix \( A \) such that \( A^2 = $Z \).
$BR
$BR
$BCENTER
\( A = \) \{ $answer->ans_array() \}
$ECENTER
END_TEXT
Context()->normalStrings;


ANS( $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      $stu = Matrix($student);
      return 0 if $stu == $Z;
      return $stu * $stu == $Z;
    }
) );

;

ENDDOCUMENT();
