# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Write an one-variable linear inequality to solve word problems.
#
# Last updated: Jordan 8/29/13;  Kling, 7/24/13; Hughes 7/15/2013, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBCCSS('7.EE.4','A-CED.1','A-REI.3')
## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Linear inequalities)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MLT(LinearIneq_applic)
## MLTleader(1)
## Level(4)
## MO(1)
## KEYWORDS('inequality','linear','set notation','interval notation','multiply','subtract')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

##############################################
#package MultiAnswer;
#our $separator = ' ';
#package main;

Context("Numeric");

$a=random(70,75,1);
$b=random(85,90,1);
$m=random(76,84,1);
$ans2=Compute($m*3-$a-$b);
while ($ans2==($a+$b)/2) {
   $m=random(76,84,1);
   $ans2=Compute($m*3-$a-$b);
}

Context("Interval");
$interval = Interval("[$ans2,100]");

$left = Formula("($a+$b+x)/3");
$right = Formula("$m");

Context()->strings->add("<"=>{}, ">"=>{}, "="=>{}, "<="=>{}, ">="=>{}, "!="=>{});

$comparitor = String(">=");
$opposite = String("<=");


$multians = MultiAnswer($left, $comparitor, $right)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stuLeft, $stuComp, $stuRight ) = @{$student};
      my ( $corLeft, $corComp, $corRight ) = @{$correct};
      my $stuExpr = Formula("$stuLeft - $stuRight");
      my $coeff = $stuExpr->D('x');
      if (!($coeff -> isConstant) or $coeff == Formula(0)){ return 0; }
      $stuExpr = $stuExpr/$coeff; 
      if ($stuExpr == 3*($corLeft - $corRight)){
         if ( ($coeff > 0 && $stuComp == $corComp) or 
              ($coeff<0 && $stuComp == $opposite) ){
            return 1;
         }
      }
      else { return 0;}
  }
);

##############################################

TEXT(beginproblem());

BEGIN_PGML

[%Not using percentages to avoid confusion%]
Your grade in a class is determined by the average of three test scores. You scored [`[$a]`] and [`[$b]`] points on the first two tests. To get an average score of at least [`[$m]`] points on the three tests, how many points would you need on the third test?

Let [`x`] be your score on the third test. Write an inequality to express the requirement that your average score must be at least [`[$m]`]. 

Your inequality has three parts, a left side, a right side, and a comparison operator in the middle: [`<`], [`>`], [`\leq`], [`\geq`], [`=`], or [`\neq`].  To type the comparison operator, enter one of these symbols:  *<*, *>*, *<=*, *>=*, *=*, or *!=*.

>> [__________]{$multians} [__]{$multians} [__________]{$multians} <<

Solve this inequality. What is the minimum score that you must earn on the third test to earn a [`[$m]`] average for the three tests? 

    [___________]{$ans2}

You cannot score over [`100`] on the third test. Use interval notation to represent the range of scores you could earn on the third test to average at least [`[$m]`] in this course.

    [______________]{$interval} 

END_PGML

##############################################

$showPartialCorrectAnswers=1;


$s1=3*$m;
$s2=$a+$b;

BEGIN_PGML_SOLUTION

Let [`x`] be the score you will earn on the third test. 

    [`
\begin{aligned}
   \frac{[$a]+[$b]+x}{3} & \geq  [$m] \\
   3 \cdot \frac{[$a]+[$b]+x}{3} & \geq  3 \cdot [$m] \\
   [$a]+[$b]+x & \geq  [$s1] \\
   [$s2] +x & \geq  [$s1] \\
   [$s2] +x \mathbf{{}-[$s2]}& \geq  [$s1]\mathbf{{}-[$s2]} \\
          x  & \geq  [$ans2]
\end{aligned}
`]

You have to score at least [`[$ans2]`] on the third test in order to earn [`[$m]`] in this course.

To earn at least [`[$m]`] in this course, the interval for your score on the third test is [`[$interval]`].

END_PGML_SOLUTION

ENDDOCUMENT();

