# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Compositions and combinations of functions)
## Institution(Loyola University Chicago)
## Author(Adam Spiegler)
## MLT(ccf-harder-powers)
## MLTleader(1)
## MO(1)
## TitleText1('Functions Modeling Change')
## TitleText2('Functions Modeling Change')
## TitleText3('Functions Modeling Change')
## AuthorText1('Connally')
## AuthorText2('Connally')
## AuthorText3('Connally')
## EditionText1('3')
## EditionText2('4')
## EditionText3('5')
## Section1('8.3')
## Section2('10.3')
## Section3('10.3')
## Problem1('15')
## Problem2('15')
## Problem3('15')
## KEYWORDS('composition','function')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "extraAnswerEvaluators.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context("Numeric");

$showPartialCorrectAnswers = 1;

$m = random(2,9,1);
$b = random(2,9,1);
$k = random(2,9,1);

$p1 = 2*$k;
$ans1 = $m**2;
$ans2 = 2*$b*$m;
$ans3 = $b**2;

$fchk = MultiAnswer( Compute("$ans1*e^($p1*x)"), Compute("$ans2*e^($k*x)"),
		     Formula("$ans3") )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ans ) = @_;
	if ( $c->[0] == $s->[0] && $c->[1] == $s->[1] && $c->[2] == $s->[2] ) {
	    return [1, 1, 1];
	} else {
	    my %c_used = ( 0=>0, 1=>0, 2=>0 );
	    my %s_used = ( 0=>0, 1=>0, 2=>0 );
	    my @res = (0, 0, 0);
	    for ( my $i=0; $i<3; $i++ ) {
		for ( my $j=0; $j<3; $j++ ) {
		    if ( $c->[$i] == $s->[$j] && ! $c_used{$i} &&
			 ! $s_used{$j} ) {
			$c_used{$i} = 1;
			$s_used{$j} = 1;
			$res[$j] = 1;
		    }
		}
	    }
	    return [ @res ];
	}
} );

Context()->texStrings;
BEGIN_TEXT

Let \( u(x) = e^{$k x} \) and \( v(x) = $m x + $b \).  Find a
simplified formula for the  function below.
$PAR
\( \bigg( v \big( u(x) \big) \bigg)^2 = \)
\{ $fchk->ans_rule(20) \} \( + \)
\{ $fchk->ans_rule(20) \} \( + \)
\{ $fchk->ans_rule(20) \}


END_TEXT
Context()->normalStrings;

ANS( $fchk->cmp() );

# ANS(num_cmp( $ans1 ));
# ANS(num_cmp( $ans2 ));
# ANS(num_cmp( $ans3 ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
We start by computing
\[ v \big( u(x) \big) = $m (u(x)) + $b = $m e^{$k x} + $b. \]
Next we compute  \( \bigg( v \big( u(x) \big) \bigg)^2 \) by squaring
the expression above:
\[\begin{aligned}
\bigg( v \big( u(x) \big) \bigg)^2
&= \left( $m e^{$k x} + $b \right)^2 \\
&= \left( $m e^{$k x} \right)^2 + 2 \left( $m e^{$k x} \right) ($b) + $b^2 \\
&= $ans1 e^{$p1 x} + $ans2 e^{$k x} + $ans3.
\end{aligned}\]

END_SOLUTION
Context()->normalStrings;


ENDDOCUMENT();
