# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>
# ENDDESCRIPTION

## DBsubject('Precalculus')
## DBchapter('Compositions, Inverses, And Combinations Of Functions')
## DBsection('Combinations Of Functions')
## KEYWORDS('composition','function')
## TitleText1('Functions Modeling Change')
## EditionText1('3')
## AuthorText1('Connally')
## Section1('8.3')
## Problem1('16')
## TitleText2('Functions Modeling Change');
## EditionText2('4')
## AuthorText2('Connally')
## Section2('10.3')
## Problem2('16')
## Author('Adam Spiegler')
## Institution('Loyola University Chicago')

DOCUMENT();

loadMacros("PG.pl",
           "PGbasicmacros.pl",
           "PGchoicemacros.pl",
           "PGanswermacros.pl",
           "PGgraphmacros.pl",
           "PGauxiliaryFunctions.pl",
           "extraAnswerEvaluators.pl",
           "MathObjects.pl",
           "PGcourse.pl",
	   "parserMultiAnswer.pl",
);

TEXT(beginproblem());

Context("Numeric");


$showPartialCorrectAnswers = 1;

$k = random(2,9,1);
$m = random(2,9,1);
$b = random(2,9,1);

$p1 = 2*$k;

$ans1 = $m;
$ans2 = 0;
$ans3 = $b;

$fchk = MultiAnswer( Compute("$ans1*e^($p1*x)"), Formula("$ans3") )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ans ) = @_;
	if ( ( $c->[0] == $s->[0] && $c->[1] == $s->[1] ) ||
	     ( $c->[1] == $s->[0] && $c->[0] == $s->[1] ) ) {
	    return [1, 1];
	} elsif ( $c->[0] == $s->[0] || $c->[1] == $s->[0] ) {
	    return [1, 0];
	} elsif ( $c->[0] == $s->[1] || $c->[1] == $s->[1] ) {
	    return [0, 1];
	} else {
	    return [0,0];
	}
} );

Context()->texStrings;
BEGIN_TEXT

Let \( u(x) = e^{$k x} \) and \( v(x) = $m x + $b \).  Find a
simplified formula for the  function below.
$PAR
\( v \big( u(x)^2 \big) = \)
\{ $fchk->ans_rule(20) \} \( + \)
\{ $fchk->ans_rule(20) \}

END_TEXT
Context()->normalStrings;

ANS( $fchk->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
We start by computing
\[ \big( u(x) \big)^2  = \big( e^{$k x} \big)^2 = e^{2($k) x} = e^{$p1 x}. \]
Next we compose this with \( v(x) \):
\[ v \big( u(x)^2 \big) = v \left( e^{$p1 x} \right) = $m e^{$p1 x} + $b. \]

END_SOLUTION
Context()->normalStrings;


ENDDOCUMENT();
