# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 11/04/08

## Author('Gavin LaRose');
## DBsubject(Calculus - single variable)
## DBchapter(Limits and continuity)
## DBsection(Continuity - intermediate value theorem)
## Institution(University of Michigan)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('1.7')
## Section2('1.7')
## Section3('1.7')
## Section4('1.7')
## Problem1('13')
## Problem2('13')
## Problem3('13')
## Problem4('21')
## KEYWORDS('calculus','derivative','continuity','functions')

## Textbook tags
## HHChapter1('A Library of Functions')
## HHChapter2('A Library of Functions')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = random(0.5, 9, 0.5);
$c = random(1, 5, 1);

$zero = Compute(0);
$one = Compute(1);
$fzero = Compute( "$c - 1" );
$fone = Compute( "$a - cos(1) + $c" );

$func = Compute("$a*x - cos(x) + $c");

$mp = MultiAnswer( $fzero, $fone, $zero, $one )->with(
    singleResult => 0,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	@stuk = ( $student->[0], $student->[1] );
	@stuc = ( $student->[2], $student->[3] );

## need to check that values of k are between 0 and 1 and in order,
##   and that the corresponding values of c are f(kmin) and f(kmax)
	if ( $stuc[0] == $zero && $stuc[1] == $one &&
	     $stuk[0] == $fzero && $stuk[1] == $fone ) {
	    return [1,1,1,1];
	} else {
	    if ( $stuc[0] > $stuc[1] ) {
		$self->setMessage(3, "Your values of c don't satisfy the " .
				  "inequality");
		return [0,0,0,0]
	    } elsif ( $stuc[0] == $stuc[1] ) {
		$self->setMessage(3, "Your values of c give a zero-length " .
				  "range of c values");
		return [0,0,0,0]
	    } else {
		$sk0 = $func->substitute( x=>$stuc[0] );
		$sk1 = $func->substitute( x=>$stuc[1] );

		@ret = (0, 0, 1, 1);

		if ( $sk0 == $stuk[0] ) {
		    $ret[0] = 1;
		}
		if ( $sk1 == $stuk[1] ) {
		    $ret[1] = 1;
		}

		return [ @ret ];
	    }
	}
    }
);

Context()->texStrings;
BEGIN_TEXT

Consider the function \(f(x) = $a x - \cos(x) + $c\) on the interval
\(0\le x\le 1\).  The Intermediate Value Theorem guarantees that
there is a value \(c\) such that \(f(c) = k\) for which values of
\(c\) and \(k\)?  Fill in the following mathematical statements, giving
an interval with non-zero length in each case.
$PAR
For every \(k\) in
\{ $mp->ans_rule(20) \} \( \le k \le \) \{ $mp->ans_rule(20) \},
$BR
there is a \(c\) in
\{ $mp->ans_rule(20) \} \( \le c \le \) \{ $mp->ans_rule(20) \}
$BR
such that \( f(c) = k \).

END_TEXT
Context()->normalStrings;

ANS($mp->cmp() );
# ANS($fzero->cmp() );
# ANS($fone->cmp() );
# ANS($zero->cmp() );
# ANS($one->cmp() );

## might try $fone_disp = $func->substitute(x=>1) if $fone
##    turns into a decimal
$apc = $a + $c;
$fone_disp = "$apc - \cos(1)";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

We know \(f\) is continuous, and that \(f(0) = $fzero\) and
\(f(1) = $fone_disp\).  Thus, we know by the Intermediate Value
Theorem that for any \(k\) between these values there must be
a \(c\) between 0 and 1 for which \(f(c) = k\).  (Of course, we
could also pick any other range of values \(c_0 \le x \le c_1\),
noting that the Intermediate Value Theorem guarantees that there
is a \(c\) in that range such that \(f(c) = k\) for any \(k\) in
\(f(c_0) \le k \le f(c_1)\).)

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
