## DESCRIPTION
## Diagonalize a 3x3 matrix.
##    The matrix is guaranteed to have integer eigenvalues between -5 and 5, at least 2 of which are different.
##    The matrix will have small integer entries and will have rank 2 or 3.
##    There are infinitely many different possible diagonalizations. All of them are accepted.
##    There is a solution with small integer entries.
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(March 4, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('3.4')
## Problem1('')
## KEYWORDS('linear algebra','matrix','eigenvalue','diagonalization')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

$flg = random(1,5);
if($flg == 1) {
    $P = Matrix([[2,5,-2],[0,3,-1],[-1,-2,1]]);
    $PI = Matrix([[1,-1,1],[1,0,2],[3,-1,6]]);
} elsif($flg == 2) {
    $P = Matrix([[3,0,1],[1,1,1],[2,0,1]]);
    $PI = Matrix([[1,0,-1],[1,1,-2],[-2,0,3]]);
} elsif($flg == 3) {
    $P = Matrix([[1,-2,2],[-2,5,-4],[-1,2,-1]]);
    $PI = Matrix([[3,2,-2],[2,1,0],[1,0,1]]);
} elsif($flg == 4) {
    $P = Matrix([[1,1,-1],[0,1,-1],[2,0,1]]);
    $PI = Matrix([[1,-1,0],[-2,3,1],[-2,2,1]]);
} else {
    $P = Matrix([[2,1,1],[3,2,1],[1,1,1]]);
    $PI = Matrix([[1,0,-1],[-2,1,1],[1,-1,1]]);
}

do {
    $eig1 = random(-5,5);
    $eig2 = random(-5,5);
    $eig3 = random(-5,5);
    
    $D = Matrix([[$eig1,0,0],[0,$eig2,0],[0,0,$eig3]]);
    
    $A = $P * $D * $PI;
} until ($A->norm_one < 20 && (($eig1 != $eig2) || ($eig1 != $eig3) || ($eig2 != $eig3)) && ((abs($eig1) + abs($eig2) != 0) || (abs($eig1) + abs($eig3) != 0) || (abs($eig2) + abs($eig3) != 0)));

$multians = MultiAnswer($P,$D)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if(abs($s0->det) > 0.0001 && $s1->element(1,2) == 0 && $s1->element(1,3) == 0 && $s1->element(2,1) == 0 && $s1->element(2,3) == 0 && $s1->element(3,1) == 0 && $s1->element(3,2) == 0 && $s0*$s1 == $A*$s0) {
            return 1;
        } else {
            return 0;
        }
    }
);


##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Diagonalize the matrix \( \displaystyle A = $A \). That is, find an invertible matrix \( P \) and a diagonal matrix \( D \) such that \( A = PDP^{-1} \).
$BR
$BR
\( P \) = \{ $multians->ans_array \}
$BR
\( D \) = \{ $multians->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
