## DESCRIPTION
## Compute a square root of a 2x2 matrix (via diagonalization).
##    The matrix is guaranteed to have distinct integer eigenvalues equal to either 1, 4, or 9.
##    The matrix will have small-ish integer entries.
##    There will be 4 possible different solutions. They are all accepted.
##    At least one of the solutions will have small integer entries.
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(March 4, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('3.4')
## Problem1('')
## KEYWORDS('linear algebra','matrix','eigenvalue','diagonalization','matrix square root')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

$flg = random(1,3);
if($flg == 1) {
    $P = Matrix([[2,1],[1,1]]);
    $PI = Matrix([[1,-1],[-1,2]]);
} elsif($flg == 2) {
    $P = Matrix([[1,2],[1,3]]);
    $PI = Matrix([[3,-2],[-1,1]]);
} else {
    $P = Matrix([[1,-3],[2,-5]]);
    $PI = Matrix([[-5,3],[-2,1]]);
}
    
do {
    $eig1 = random(1,3);
    $eig2 = random(1,3);
    $D = Matrix([[$eig1,0],[0,$eig2]]);
    
    $Sq = $P * $D * $PI;
    $A = $P * $D * $D * $PI;
} until ($eig1 != $eig2);


##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a square root of the matrix \( \displaystyle $A \).
$BR
$BR
Square root: \{ $Sq->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS($Sq->cmp(
    checker => sub {
        my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
        if($student*$student == $A) {
            return 1;
        } else {
            return 0;
        }
    }
));

ENDDOCUMENT();
