## DESCRIPTION
## Compute a square root of a 3x3 matrix (via diagonalization).
##    The matrix is guaranteed to have integer eigenvalues equal to either 1, 4, or 9, at least 2 of which are different.
##    The matrix will have small-ish integer entries.
##    There will be multiple possible different solutions. They are all accepted.
##    At least one of the solutions will have small integer entries.
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(March 4, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('3.4')
## Problem1('')
## KEYWORDS('linear algebra','matrix','eigenvalue','diagonalization','matrix square root')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

$flg = random(1,5);
if($flg == 1) {
    $P = Matrix([[2,5,-2],[0,3,-1],[-1,-2,1]]);
    $PI = Matrix([[1,-1,1],[1,0,2],[3,-1,6]]);
} elsif($flg == 2) {
    $P = Matrix([[3,0,1],[1,1,1],[2,0,1]]);
    $PI = Matrix([[1,0,-1],[1,1,-2],[-2,0,3]]);
} elsif($flg == 3) {
    $P = Matrix([[1,-2,2],[-2,5,-4],[-1,2,-1]]);
    $PI = Matrix([[3,2,-2],[2,1,0],[1,0,1]]);
} elsif($flg == 4) {
    $P = Matrix([[1,1,-1],[0,1,-1],[2,0,1]]);
    $PI = Matrix([[1,-1,0],[-2,3,1],[-2,2,1]]);
} else {
    $P = Matrix([[2,1,1],[3,2,1],[1,1,1]]);
    $PI = Matrix([[1,0,-1],[-2,1,1],[1,-1,1]]);
}
    
do {
    $eig1 = random(1,3);
    $eig2 = random(1,3);
    $eig3 = random(1,3);
    $D = Matrix([[$eig1,0,0],[0,$eig2,0],[0,0,$eig3]]);
    
    $Sq = $P * $D * $PI;
    $A = $P * $D * $D * $PI;
} until ($eig1 != $eig2 || $eig1 != $eig3 || $eig2 != $eig3);


##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a square root of the matrix \( \displaystyle $A \).
$BR
$BR
Square root: \{ $Sq->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS($Sq->cmp(
    checker => sub {
        my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
        if($student*$student == $A) {
            return 1;
        } else {
            return 0;
        }
    }
));

ENDDOCUMENT();
