##DESCRIPTION
## sbrummel tagged and PAID on 2-20-2004
## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Mixed techniques)
## Date(6/3/2002)
## Institution(Indiana)
## Author(Arnie Pizer)
## MO(1)
## Level(3)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('6')
## Section1('7.2')
## Problem1('26 28')
## TitleText2('Calculus: Early Transcendentals')
## AuthorText2('Rogawski')
## EditionText2('1')
## Section2('7.4')
## Problem2('17')
## KEYWORDS('integrals', 'trig', 'substitution','Substitution' 'Trig Substitution' 'Trigonometric Substitution')
##ENDDESCRIPTION

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric");

$a = random(2,9,1);
$k = random(4,9,1);

$aa = $a*$a;
$n = 2*$k + 3;
$m = 2*$k + 1;
$p = 2*$k;

$integrand = Compute("x^($p)/($aa-x^2)^($n/2)");

$ans = Compute("(1/($aa*$m)) * (x^($m)/($aa - x^2)**($m/2))");

Context()->texStrings;
BEGIN_TEXT
Evaluate the indefinite integral
$BR \[\int \frac{ x^{$p} }{($aa - x^2)^{$n/2}} \, dx \]
$PAR \{ans_rule( 50) \}+C 
END_TEXT

BEGIN_HINT
Make a trigonometric substitution.  To attack the resulting
integral, remember that \( (\sin x)/(\cos x) = \tan x \)
and that \( 1/(\cos x) = \sec x \).
END_HINT

Context()->normalStrings;

#check that the derivative of the student's answer equals $integrand
ANS( $integrand->cmp( checker=>sub{
  my ($correct, $student, $ansHash ) = @_;
  return $correct == $student->D(x);
}));

$soln_den = $aa*$m;

Context()->texStrings;
BEGIN_SOLUTION
Use the right triangle 
$BCENTER
\{ image("triangle.png", width=>137, height=>79, tex_size=>500, 
extra_html_tags=>'alt="right triangle: theta is angle, a is hypotenuse, x is side opposite theta, square root of (a squared minus x squared) is side adjacent to theta."') \}
$ECENTER
with \(a = $a\) to create a trig substitution 

\[x=$a \sin(\theta), \text{ so } dx = $a \cos(\theta)\,d\theta.\] 

\[
\begin{aligned}
\int \frac{x^{$p}}{\left($aa-x^2\right)^{$n/2}}\,dx
&=        \int \frac{x^{$p}}{\left(\sqrt{$aa-x^2}\right)^{$n}}\,  dx      \\\\
&=        \int \frac{\left($a \sin(\theta)\right)^{$p} \cdot $a \cos(\theta) }
                {\left(\sqrt{$a^2-$a^2\sin^2\theta}\right)^{$n}}  \, d\theta      \\\\
&=        \int \frac{$a^{$m}\sin^{$p}(\theta)\cos(\theta) }
                {\left(\sqrt{$a^2\cos^2(\theta)}\right)^{$n}}  \, d\theta      \\\\
&=        \int \frac{$a^{$m}\sin^{$p}(\theta)\cos(\theta) }
                {\left($a \cos(\theta)\right)^{$n}} \, d\theta       \\\\
&=        \frac{1}{$aa} \int \tan^{$p}(\theta) \sec^2(\theta)  \, d\theta       \\\\
&=        \frac{\tan^{$m}(\theta)}{$aa \cdot $m} +C = \frac{\tan^{$m}(\theta)}{$soln_den} +C
\end{aligned}
\]

Now use the right triangle to find \(\tan(\theta)\) in terms of \(x\).
\[\tan(\theta)=\frac{\text{opp}}{\text{adj}}=\frac{x}{\sqrt{$aa - x^2}},\]
so back substitution gives
\[  \frac{\tan^{$m}(\theta)}{$soln_den} +C  
=        \frac{1}{$soln_den}\frac{x^{$m}}{($aa-x^2)^{$m/2}}+C\]
That is, 
\[\int \frac{x^{$p}}{\left($aa-x^2\right)^{$n/2}}\,dx =  \frac{1}{$soln_den}\frac{x^{$m}}{($aa-x^2)^{$m/2}}\]
leaving off the arbitrary constant of integration.
$BR

END_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.
