## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS ('linear algebra','matrix','eigenvalue','complex')
## Tagged by cmd6a 4/29/06

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Eigenvalues')
## Date('')
## Author('')
## Institution('Rochester')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros("PGmatrixmacros.pl",
           "PGstandard.pl",
           "MathObjects.pl");

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;
Context("Numeric");

$a = non_zero_random(-7,7,2);
$b1 = random(2,9,1);
$r = random(-1,1,2);
$b = $b1 * $r;
$c = non_zero_random(-9,9,1);
if ($b * $c > 0) { $c = -1 * $c; }
$d = random($a-2, $a+2,1);

# the real and imaginary parts of an eigenvalue are 

$re = Compute("($a+$d)/2");
$im = Compute("sqrt(-($a - $d)^2 -4*$b*$c)/2");

Context()->texStrings;
BEGIN_TEXT

The matrix 
$BR 
\{ mbox( '\( A = \)', display_matrix([[$a, $b], [$c, $d]]) ) \} 
$BR
has complex eigenvalues, \( \lambda_{1,2}= a \pm b i\), where \(a=\) \{ans_rule(10)\} and \(b=\) \{ans_rule(10)\}.

END_TEXT
Context()->normalStrings;

ANS($re->cmp() );
ANS($im->cmp( checker=>sub { my ( $correct, $student, $ansHash ) = @_;
    return $correct == abs($student); } ) );

ENDDOCUMENT();