## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Quotient rule (without trigonometric functions))
## Institution(UCSB)
## MLT(quotRuleNoTrigLinearoverLinearTangent)
## Level(4)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('5')
## Section1('3.2')
## Problem1('42')
## KEYWORDS('Product','Quotient','Differentiation')
## MO(1)
## updated 12/2014

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(&beginproblem);
$showPartialCorrectAnswers = 1;
Context('Numeric');

$a=random(1,5,1);
$b=2*($a);


$line1 = Formula("(1/$b)(x-$a)");
$line2 = Formula("(1/$b)(x+3*($a))+2");

$multians = MultiAnswer($line1, $line2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
BEGIN_TEXT


Find the equations of the tangent lines to the curve \(\displaystyle y=\frac{x-$a}{x+$a}\) that are parallel to the line \(x-$b y=2\). 

$PAR

\(y=\)\{$multians->ans_rule(20)\} $BR$BR
\(y=\)\{$multians->ans_rule(20)\}


END_TEXT



ANS( $multians->cmp() );

ENDDOCUMENT();

