#DESCRIPTION
#  Series: Ratio test
#ENDDESCRIPTION


## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Ratio test)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## MLT(displayed_ratio_limit_conclusion)
## TitleText1('Calculus')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('series','ratio test')


###############################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "unionTables.pl",
  "parserPopUp.pl",
  "answerHints.pl",
  "PGgraders.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem);


###############################
#  Setup

Context("Numeric")->variables->are(n=>"Real");

$a = random(2,9,1);
$b = random(1,2,1);
$c = random(2,30,1);

$bp1 = $b + 1;

# TeX
$series = "\sum_{n=$c}^{\infty} \frac{n $a^n}{(n+$b)!}";
$fraction = "\lim_{n \to \infty} \left| \frac{ a_{n+1} }{ a_{n} } \right| = \lim_{n \to \infty}";


$num1 = Formula("$a (n+1)");
$den1 = Formula("n (n+$bp1)"); 

$num2 = Formula("$a + $a/n");
$den2 = Formula("n+$bp1"); 

$numbogus[0] = Formula("(n+1)*($a)^(n+1) * (n+$b)!");
$numbogus[1] = Formula("(n+1)*($a)^(n+1)");
$numbogus[2] = Formula("(n+1) * (n+$b)!");

foreach my $i (0..2) {
$numbogus[$i]->{test_points} = [[1],[2],[3],[4]];
}

$denbogus[0] = Formula("((n+$bp1)!) * n*$a^n");
$denbogus[1] = Formula("((n+$bp1)!) * n");
$denbogus[2] = Formula("(n+$bp1) * n");

foreach my $i (0..2) {
$denbogus[$i]->{test_points} = [[1],[2],[3],[4]];
}


$multians = MultiAnswer($num1, $den1)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $stu1, $stu2 ) = @{$student};

      if ( ($numbogus[0] == $stu1 && $denbogus[0] == $stu2) || 
           ($numbogus[1] == $stu1 && $denbogus[1] == $stu2) ||
           ($numbogus[2] == $stu1 && $denbogus[2] == $stu2)
         )
      {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];
      } elsif (($num1 == $stu1 && $den2 == $stu2) || 
               ($num2 == $stu1 && $den1 == $stu2)) {
          $ansHash->setMessage(1,"Check your algebra");
          $ansHash->setMessage(2,"Check your algebra");
          return [0,0];
      }
    
      
      if (($num1 == $stu1 && $den1 == $stu2) ||
          ($num2 == $stu1 && $den2 == $stu2) ) {
          return [1,1];
      } elsif ($num1 == $stu1 || $num2 == $stu1) {
          return [1,0];
      } elsif ($den1 == $stu2 || $den2 == $stu2) {
          return [0,1];
      } else {
          return [0,0];
      }
  }
);



$limit = Formula("0");
$popup = PopUp(["Choose","Converges","Diverges","Inconclusive"],"Converges");



# 
#  Display the fraction and answer blanks nicely
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
  $showfraction =
  "\[ $fraction ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $showfraction =
  ColumnTable(
  "\( \displaystyle $fraction \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}
Context()->normalStrings;



#################################
#  Main text

BEGIN_TEXT
Use the ratio test to determine whether 
\( \displaystyle $series \) converges or diverges.
$BR
$BR
(a) Find the ratio of successive terms.  Write your answer as
a fully simplified fraction.
For \( n \geq $c \), 
$showfraction
$BR
(b) Evaluate the limit in the previous part.  Enter \( \infty \) 
as ${BITALIC}infinity${EITALIC} and \( -\infty \) 
as ${BITALIC}-infinity.${EITALIC}  If the limit does not exist, 
enter ${BITALIC}DNE.${EITALIC}
$BR
\( \displaystyle \lim_{n\to\infty} \left| \frac{a_{n+1}}{a_{n}} \right| \, \) = 
\{ ans_rule(20) \}
$BR
$BR
(c) By the ratio test, does the series
converge, diverge, or is the test inconclusive?
\{ $popup->menu() \}
END_TEXT


##########################################
#  Answer evaluation

$showPartialCorrectAnswers=1;
install_problem_grader(~~&custom_problem_grader_fluid);
$ENV{'grader_numright'} = [2,4];
$ENV{'grader_scores'} = [0.4,1];
$ENV{'grader_message'} = "You can earn " .
"40% partial credit for 2 - 3 correct answers.";

ANS( $multians->cmp() );

ANS( $limit->cmp() );
ANS( $popup->cmp() );

BEGIN_SOLUTION
$PAR SOLUTION $PAR

Using the ratio test to test convergence of the series, \( {\displaystyle\sum_{n=$c}^{\infty} \frac{n $a^n}{(n+$b)!}}\), we have  
 $BR
 \(a_n =  \frac{n $a^n}{(n+$b)!}\), and   \(a_{n+1} = \frac{(n+1) $a^{n+1}}{(n+$bp1)!}\), so the ratio of consecutive terms is 
\[\left|\frac{a_{n+1}}{a_n}\right| = \left|\frac{\left((n+1) $a^{n+1}\right) / \left((n+$bp1)!\right) }{\left(n $a^n\right) / \left((n+$b)!\right) }\right|\]
\[=   \frac{(n+1) $a^{n+1}}{(n+$bp1)!}  \frac{(n+$b)!} {n $a^n}= \frac{(n+1) $a^{n+1}}{n $a^n}  \frac{(n+$b)!}{(n+$bp1)!}  \]
\[=  \left(\frac{n+1}{n}\right)\left(\frac{\left($a^n\right)($a)}{$a^n}\right)  \left(\frac{(n+$b)!}{(n+$bp1)(n+$b)!}\right) \]
\[ = \frac{$a(n+1)}{n(n+$bp1)}    = \frac{$a n +$a}{n^2+$bp1 n}  \]

So the limit \(L = \displaystyle{ \lim_{n\to \infty}} \left|\frac{a_{n+1}}{a_n}\right|\) is 
\[L = \lim_{n\to \infty}\left|\frac{a_{n+1}}{a_n}\right| = \lim_{n\to \infty}   \frac{$a n +$a}{n^2+$bp1 n} =0  \]
And since  \(L= 0 <1\), the series converges.
$BR 

END_SOLUTION



ENDDOCUMENT();


$numbogus[0] = Formula("(n+1)*($a)^(n+1) * (n+$b)!");
$numbogus[1] = Formula("(n+1)*($a)^(n+1)");
$numbogus[2] = Formula("(n+1) * (n+$b)!");

foreach my $i (0..2) {
$numbogus[$i]->{test_points} = [[1],[2],[3],[4]];
}

ANS( $num->cmp() 
->withPostFilter(AnswerHints( 
  [$numbogus[0],$numbogus[1],$numbogus[2]]
  => "Simplify your answer further (factor and cancel)",
))
);


$denbogus[0] = Formula("((n+$bp1)!) * n*$a^n");
$denbogus[1] = Formula("((n+$bp1)!) * n");
$denbogus[2] = Formula("(n+$bp1) * n");

foreach my $i (0..2) {
$denbogus[$i]->{test_points} = [[1],[2],[3],[4]];
}

ANS( $den->cmp() 
->withPostFilter(AnswerHints( 
  [$denbogus[0],$denbogus[1],$denbogus[2]]
  => "Simplify your answer further (factor and cancel)",
))
);
