#DESCRIPTION
#  Series: Ratio test
#ENDDESCRIPTION


## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Ratio test)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## MLT(displayed_ratio_limit_conclusion)
## TitleText1('Calculus')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('series','ratio test')


#######################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "unionTables.pl",
  "parserPopUp.pl",
  "answerHints.pl",
  "PGgraders.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem);


################################
#  Setup

Context("Numeric")->variables->are(n=>"Real");

$a = random(2,9,1);
$c = random(2,30,1);

$ap1 = $a + 1;

$num1 = Formula("n+$ap1");
$den1 = Formula("(n+1)(n+$a)"); 

$num2 = Formula("1+$ap1/n");
$den2 = Formula("n+$ap1+$a/n");

$numbogus = Formula("(n+$ap1)*n!");
$numbogus->{test_points} = [[1],[2],[3],[4]];

$denbogus = Formula("(n+$a)*((n+1)!)");
$denbogus->{test_points} = [[1],[2],[3],[4]];


$multians = MultiAnswer($num1, $den1)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $stu1, $stu2 ) = @{$student};

      if ( $numbogus == $stu1 && $denbogus == $stu2 )
      {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];
      } elsif (($num1 == $stu1 && $den2 == $stu2) || 
               ($num2 == $stu1 && $den1 == $stu2)) {
          $ansHash->setMessage(1,"Check your algebra");
          $ansHash->setMessage(2,"Check your algebra");
          return [0,0];
      }
    
      
      if (($num1 == $stu1 && $den1 == $stu2) ||
          ($num2 == $stu1 && $den2 == $stu2) ) {
          return [1,1];
      } elsif ($num1 == $stu1 || $num2 == $stu1) {
          return [1,0];
      } elsif ($den1 == $stu2 || $den2 == $stu2) {
          return [0,1];
      } else {
          return [0,0];
      }
  }
);



$limit = Formula("0");
$popup = PopUp(["Choose","Converges","Diverges","Inconclusive"],"Converges");

# TeX
$series = "\sum_{n=$c}^{\infty} \frac{n+$a}{n!}";
$fraction = "\lim_{n \to \infty} \left| \frac{ a_{n+1} }{ a_{n} } \right| = \lim_{n \to \infty}";


# 
#  Display the fraction and answer blanks nicely
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
  $showfraction =
  "\[ $fraction ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $showfraction =
  ColumnTable(
  "\( \displaystyle $fraction \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}
Context()->normalStrings;


######################################
#  Main text

BEGIN_TEXT
Use the ratio test to determine whether 
\( \displaystyle $series \) converges or diverges.
$BR
$BR
(a) Find the ratio of successive terms.  Write your answer as
a fully simplified fraction.
For \( n \geq $c \), 
$showfraction
$BR
(b) Evaluate the limit in the previous part.  Enter \( \infty \) 
as ${BITALIC}infinity${EITALIC} and \( -\infty \) 
as ${BITALIC}-infinity.${EITALIC}  If the limit does not exist, 
enter ${BITALIC}DNE.${EITALIC}
$BR
\( \displaystyle \lim_{n\to\infty} \left| \frac{a_{n+1}}{a_{n}} \right| \, \) = 
\{ ans_rule(20) \}
$BR
$BR
(c) By the ratio test, does the series
converge, diverge, or is the test inconclusive?
\{ $popup->menu() \}
END_TEXT


###########################################
#  Answer evaluation

$showPartialCorrectAnswers=1;
install_problem_grader(~~&custom_problem_grader_fluid);
$ENV{'grader_numright'} = [2,4];
$ENV{'grader_scores'} = [0.4,1];
$ENV{'grader_message'} = "You can earn " .
"40% partial credit for 2 - 3 correct answers.";

ANS( $multians->cmp() );


ANS( $limit->cmp() );
ANS( $popup->cmp() );

BEGIN_SOLUTION
$PAR SOLUTION $PAR

Using the ratio test to test convergence of the series, \( {\displaystyle\sum_{n=$c}^{\infty}\frac{n+$a}{n!}}\), we have  
 $BR
 \(a_n =  \frac{n+$a}{n!}\), and   \(a_{n+1} = \frac{n+$ap1}{(n+1)!}\), so the ratio of consecutive terms is 
\[\left|\frac{a_{n+1}}{a_n}\right| = \left|\frac{\left(n+$ap1\right) / \left((n+1)!\right) }{\left(n+$a\right) / \left((n)!\right) }\right|\]
\[=  \frac{(n+$ap1)}{(n+1)!}  \frac{n!}{(n+$a)}  =  \frac{(n+$ap1)}{(n+$a)}  \frac{n!}{(n+1)!}  \]
\[=  \left(\frac{n+$ap1}{n+$a}\right)  \left(\frac{n!}{(n+1)(n!)}\right)  =  \left(\frac{n+$ap1}{(n+$a)(n+1)}\right)  \left(\frac{n!}{n!}\right)  \]
\[= \left(\frac{n+$ap1}{(n+$a)(n+1)}\right) = \left(\frac{n+$ap1}{n^2+$ap1 n+$a}\right)  \]

So the limit \(L = \displaystyle{ \lim_{n\to \infty}} \left|\frac{a_{n+1}}{a_n}\right|\) is 
\[L = \lim_{n\to \infty}\left|\frac{a_{n+1}}{a_n}\right| = \lim_{n\to \infty}   \frac{n+$ap1}{n^2+$ap1 n+$a} =0  \]
And since  \(L= 0 <1\), the series converges.
$BR 

END_SOLUTION


ENDDOCUMENT();

$numbogus = Formula("(n+$ap1)*n!");
$numbogus->{test_points} = [[1],[2],[3],[4]];


ANS( $num->cmp() 
->withPostFilter(AnswerHints( 
  $numbogus
  => "Simplify your answer further (factor and cancel)",
))
);


$denbogus = Formula("(n+$a)*((n+1)!)");
$denbogus->{test_points} = [[1],[2],[3],[4]];

ANS( $den->cmp() 
->withPostFilter(AnswerHints( 
  $denbogus
  => "Simplify your answer further (factor and cancel)",
))
);


