# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Basic vocabulary questions
#
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Function')
## DBsection()
## KEYWORDS('function',)
## DBCCSS('')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "contextInequalities.pl",
  "parserMultiAnswer.pl",
  "parserOneOf.pl",
  "PGcourse.pl",
);


##############################################
TEXT(beginproblem());

Context("Inequalities");

Context()->variables->are(a=>'Real');
Context()->constants->add(b=>exp(sqrt(2)));
Context()->constants->set(b => {TeX=>'b'});

Context()->strings->add(
  "increasing"=>{},
  "decreasing"=>{},
);

$ans1 = OneOf("increasing","decreasing");
$ans2 = OneOf("a<b","b<a");

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2 ) = @{$student};
      if ( (String("increasing") == $stu1 && Inequality("a<b") == $stu2) ||
           (String("decreasing") == $stu1 && Inequality("b<a") == $stu2) ) {
          return 1;
      } else {
              return 0;
      }
  }
);

##############################################

BEGIN_TEXT

A function \(f\) is said to be \{$multians->ans_rule(10)\} on an interval \(I\) if for any two numbers \(a\), \(b\) in \(I\) where \{$multians->ans_rule(10)\}, it's true that \(f(a)<f(b)\).


END_TEXT


##############################################

ANS($multians->cmp(correct_ans_latex_string=>'\text{increasing (or decreasing)};\,a<b\text{ (or }b<a\text{)}'));

BEGIN_PGML_SOLUTION

A function [`f`] is said to be _increasing_ on an interval [`I`] if for any two numbers [`a`], [`b`] in [`I`] where [`a<b`], it's true that [`f(a)<f(b)`].

A function [`f`] is said to be _decreasing_ on an interval [`I`] if for any two numbers [`a`], [`b`] in [`I`] where [`a>b`], it's true that [`f(a)<f(b)`].



END_PGML_SOLUTION


ENDDOCUMENT();