# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Applications - other)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Applied Calculus')
## TitleText5('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## AuthorText5('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('4')
## EditionText5('7')
## Section1('11.6')
## Section2('11.6')
## Section3('11.6')
## Section4('10.5')
## Section5('11.6')
## Problem1('19')
## Problem2('23')
## Problem3('22')
## Problem4('17')
## Problem5('24')
## KEYWORDS('calculus', 'integral', 'differential equations', 'modeling')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are( k => 'Real', y => 'Real',
			   a => 'Real', t => 'Real' );

$dy1 = random( 0.10, 0.20, 0.02 );
## ensure that we have a concave up exponential with a final
##    value > 0 (the latter requirement ensures the former
##    on the domain we're picking for dy1)
$dy2bound = $dy1*(1 - $dy1);
$dy2 = random( 0.05, ($dy2bound - 0.01), 0.01 );

$y1  = 1 - $dy1;
$y2  = $y1 - $dy2;
$y1p = round( 100*$y1 );
$y2p = round( 100*$y2 );

$k = ln((1 - $y1)/($y1 - $y2));
# $a = ($y2 - $y1*$y1)/(1 + $y2 - 2*$y1);
$a = 1 - ($dy1*$dy1)/($dy1 - $dy2);

$deqn = Compute( "-k*(y - a)" );
$icon = Compute( 1 );
$ysol = Compute( "(1 - a)*e^(-k*t) + a" );
$yfin = Compute( "(1 - $a)*e^(-$k*14) + $a" );

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context()->texStrings;
BEGIN_TEXT

As you know, when a course ends, students start to forget the material
they have learned.  One model (called the Ebbinghaus model) assumes
that the rate at which a student forgets material is proportional to
the difference between the material currently remembered and some
positive constant, \( a \).
$PAR
${BBOLD}A.$EBOLD
Let \( y=f(t) \) be the fraction of the original material remembered
\( t \) weeks after the course has ended.  Set up a differential
equation for \( y\), using \(k\) as any constant of proportionality
you may need (let \(k>0\)).  Your equation will contain two constants;
the constant \( a \) (also positive) is less than \( y \) for all 
\( t \). 
$BR
\(\frac{dy}{dt} = \) \{ ans_rule(35) \}
$BR
What is the initial condition for your equation?
$BR
\( y(0) = \) \{ ans_rule(15) \}

$PAR
${BBOLD}B.$EBOLD
Solve the differential equation.
$BR
\( y = \) \{ ans_rule(35) \}

$PAR
${BBOLD}C.$EBOLD
What are the practical meaning (in terms of the amount remembered)
of the constants in the solution \( y=f(t)\)?  If after one week the
student remembers $y1p percent of the material learned in the semester,
and after two weeks remembers $y2p percent, how much will she or he
remember after summer vacation (about 14 weeks)?
$BR
percent = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( $deqn->cmp() );
ANS( $icon->cmp() );
ANS( $ysol->cmp() );
ANS( $yfin->cmp( checker=> sub {
     my ( $correct, $student, $ansHash ) = @_;
     return $correct == $student || 100*$correct == $student ||
       round(100*$correct) == $student; } ) );

$aval = round(1000*$a)/1000;
$kval = round(1000*$k)/1000;
$as = ( abs($aval - $a) < 0.0000001 ) ? "=" : "\approx";
$ks = ( abs($kval - $k) < 0.0000001 ) ? "=" : "\approx";

$twoy1 = 2*$y1;
$y2p1 = 1 + $y2;
$y1sq = $y1*$y1;
$ysum = round(1000*($a + (1 - $a)*exp(-$k*14)))/1000;
$yperc = round(1000*$ysum)/10;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Our differential equation is 
\[ 
  \frac{dy}{dt} = -k(y-a),
\]
where \(k >0\) and \(a\) are constants.
We can solve this using separation of variables, getting 
\( \int \frac{dy}{y-a} = \int -k\,dt\), so 
\(\ln |y-a| = \ln(y-a) = -kt+C\). Thus,
\(y-a=Ae^{-kt}\) where \(A=e^C\).  

Initially nothing has been forgotten, so \(y(0)=1\). 
Therefore,  \( 1-a = Ae^0 = A\), so 
\(y - a = (1 - a) e^{-kt}\), or
\(y = ( 1 - a ) e^{-kt} + a\).

As \(t \to \infty\),  \(e^{-kt} \to 0\), so \(y\to a\).
Thus, \(a\) represents the fraction of material which is remembered in
the long run.  The constant \(k\) tells us about the rate at which
material is forgotten.

$PAR
Then, after one week we have 
\[ $y1 = (1 - a) e^{-k} + a, \]
and after two,
\[ $y2 = (1 - a) e^{-2k} + a. \]
Solving the first of these for \(e^{-k}\), we get
\[ e^{-k} = \frac{$y1 - a}{1 - a}, \]
so that, plugging this in to the second,
\[ $y2 = (1 - a)\left(\frac{$y1 - a}{1 - a}\right)^2 + a, \]
or
\[ $y2 - a = \frac{($y1 - a)^2}{1-a}. \]
Cross-multiplying and expanding, we get
\[ a^2 - $y2p1 a + $y2 = a^2 - $twoy1 a + $y1sq, \]
so that 
\[ a $as $aval. \]

$PAR
Plugging this back into the expression for \(e^{-k}\) and taking the
natural log of both sides, we get \(k $ks $kval\).  Thus after the
summer we have
\[ y(14) = (1 - $aval) e^{-14 ($kval)} + $aval \approx $ysum, \]
or about $yperc percent of the material remembered.


END_SOLUTION
Context()->normalStrings;


ENDDOCUMENT();
