# DESCRIPTION
# Factor an expression of the form a^2 x^2 + 2ac x + c^2, with gcd(a,c) = 1.
# ENDDESCRIPTION

## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Institution(University of Missouri-Columbia)
## Level(3)
## KEYWORDS('factoring', 'factor', 'factorization', 'factoring, diference of squares')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
$var = list_random('x','y','z','w','p','q','r','s','t','u','v');
Context()->variables->are($var=>"Real", uc($var)=>"Real");

do {
  $a = random(2,12);
  $c = random(2,12);
} until (gcd($a,$c) == 1);

$a2 = ($a)**2;
$c2 = ($c)**2;
$e2 = 2*$a*$c;
$fac1 = "$a*$var + $c";    $afac1 = Formula($fac1)->reduce;

$expn = "$a2 $var^2 +$e2 $var + $c2";

$answer = MultiAnswer($afac1, $afac1)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    if (($c1 == $s1 && $c2 == $s2)||($c1 == -$s1 && $c2 == -$s2)) {
      return [1,1];
    } elsif (($c1==$s1)||($c1==-$s1)){
      return [1,0];
    } elsif (($c2==$s2)||($c2==-$s2)){
      return [0,1];
    }
    return [0,0];
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());



COMMENT('Desc: Factor an expression of the form a^2 x^2 + 2ac x + c^2, with gcd(a,c) = 1.');
ENDDOCUMENT();
