# DESCRIPTION
# Factor an expression of the form x^2 + 2c x + c^2.
# ENDDESCRIPTION

## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Institution(University of Missouri-Columbia)
## Level(2)
## KEYWORDS('factoring', 'factor', 'factorization', 'quadratic')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
#  "contextLimitedPowers.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
$var = list_random('x','y','z','w','p','q','r','s','t','u','v');
Context()->variables->are($var=>"Real", uc($var)=>"Real");

$c = random(2,12);

$c2 = ($c)**2;
$e2 = 2*$c;
$fac1 = "$var + $c";    $afac1 = Formula($fac1)->reduce;

$expn = "$var^2 +$e2 $var + $c2";

$answer = MultiAnswer($afac1, $afac1)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    if (($c1 == $s1 && $c2 == $s2) || ($c1 == -$s1 && $c2 == -$s1)) {
      return [1,1];
    } elsif (($c1 == $s1) || ($c1 == -$s1 )) {
      return [1,0];
    } elsif (($c2 == $s2) || ($c2 == -$s2)) {
      return [0,1];
    } 
    return [0,0];
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());

COMMENT('Desc: Factor an expression of the form x^2 + 2c x + c^2.');
ENDDOCUMENT();
