# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 03/25/10


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Triple integrals)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(Triple_Integral_Write_a_Triple_Integral)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('16.3')
## Section2('16.3')
## Section3('16.3')
## Problem1('16')
## Problem2('35')
## Problem3('42')
## KEYWORDS('triple integral', 'definite integrals', 'functions', 'multivariable')

## Textbook tags
## HHChapter1('Integrating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( y=>'Real', z=>'Real' );
$showPartialCorrectAnswers = 1;

$c1 = random(1,5,1);
$c2 = random(1,5,1);
$d1 = $c1 + random(1,3,1);
$d2 = $c2 + random(1,3,1);
## make sure that at least one of the planes is not parallel
##    to the xy-plane, so that we can't write the integral as
##    a single integral with z as one of the outer limits
$d2++ if ( $c1 == $c2 && $d1 == $d2 );

$z0 = random(0,8,1);

$p1 = Compute( "$c1*x + $c2*y + z" )->reduce();
$p2 = Compute( "$d1*x + $d2*y + z" )->reduce();

$l0 = random(1,3,1);

$mp = MultiAnswer( Formula("1"),
		   Compute("z"), Compute("y"), Compute("x"),
		   Formula("0"), Formula("$l0"),
		   Formula("0"), Compute("$l0-x"),
		   Compute("$z0-$d1*x-$d2*y"),
		   Compute("$z0-$c1*x-$c2*y") )->with(
    singleResult=>0,
    checker=>sub {
	my( $c, $s, $ansH ) = @_;
	my @ret = (0,0,0, 0, 0,0, 0,0, 0,0);
	$ret[0] = ( $c->[0] == $s->[0] ) ? 1 : 0;

	if ( $c->[1] == $s->[1] && $c->[2] == $s->[2] && $c->[3] == $s->[3] ) {
	    ## case of dz dy dx
	    $ret[1] = 1;  $ret[2] = 1;  $ret[3] = 1;
	    $ret[4] = ( $c->[4] == $s->[4] ) ? 1 : 0;
	    $ret[5] = ( $c->[5] == $s->[5] ) ? 1 : 0;
	    $ret[6] = ( $c->[6] == $s->[6] ) ? 1 : 0;
	    $ret[7] = ( $c->[7] == $s->[7] ) ? 1 : 0;
	    $ret[8] = ( $c->[8] == $s->[8] ) ? 1 : 0;
	    $ret[9] = ( $c->[9] == $s->[9] ) ? 1 : 0;
	} elsif ( $c->[1] == $s->[1] && $c->[2] == $s->[3] &&
		  $c->[3] == $s->[2] ) {
	    ## case of dz dx dy
	    $ret[1] = 1;  $ret[2] = 1;  $ret[3] = 1;
	    $ret[4] = ( $c->[4] == $s->[4] ) ? 1 : 0;
	    $ret[5] = ( $c->[5] == $s->[5] ) ? 1 : 0;
	    $ret[6] = ( $c->[6] == $s->[6] ) ? 1 : 0;
	    $ret[7] = ( Compute("$l0-y") == $s->[7] ) ? 1 : 0;
	    $ret[8] = ( $c->[8] == $s->[8] ) ? 1 : 0;
	    $ret[9] = ( $c->[9] == $s->[9] ) ? 1 : 0;
	}
	return [ @ret ];
    }
);

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Write a triple integral, including limits of
integration, that gives the volume between
\( $p1 = $z0 \) and \( $p2 = $z0 \) and above
\( x+y\le $l0, x\ge 0, y\ge 0 \).

$PAR
volume = \( \int_a^b\int_c^d\int_e^f \)\{ $mp->ans_rule(15) \}
\(d\)\{ $mp->ans_rule(3) \} \(d\)\{ $mp->ans_rule(3) \}
\(d\)\{ $mp->ans_rule(3) \}
$BR
where
\(a = \) \{$mp->ans_rule(15) \},
\(b = \) \{$mp->ans_rule(15) \},$BR
\(c = \) \{$mp->ans_rule(15) \},
\(d = \) \{$mp->ans_rule(15) \},$BR
\(e = \) \{$mp->ans_rule(15) \}, and
\(f = \) \{$mp->ans_rule(15) \}.
$BR
${BEMPH}(Note: values for all answer blanks must be supplied for
this problem to be able to check the answers provided.)$EEMPH

END_TEXT
Context()->normalStrings;

ANS( $mp->cmp() );

$p1d = Compute( "$z0 - $c1*x - $c2*y" )->reduce();
$p2d = Compute( "$z0 - $d1*x - $d2*y" )->reduce();

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The two surfaces are planes given by
\[ z = $p1d \]
and
\[ z = $p2d. \]

For \(x,y\ge0\), the plane \(z=$p1d\) is above \(z=$p2d\).  The region
in the \(xy\)-plane is the triangular region under \(y = $l0-x\),
above the \(x\)-axis and to the right of the \(y\)-axis.  Thus the
volume is
\[
V = \int_0^{$l0}\int_0^{$l0-x}\int_{$p2d}^{$p1d}1\,dz\,dy\,dx.
\]
The order of integration of \(x\) and \(y\) can be reversed, giving
\[
V = \int_0^{$l0}\int_0^{$l0-y}\int_{$p2d}^{$p1d}1\,dz\,dx\,dy.
\]

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
