# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 03/25/10


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Triple integrals)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('16.3')
## Section2('16.3')
## Section3('16.3')
## Problem1('23')
## Problem2('40')
## Problem3('47')
## KEYWORDS('triple integral', 'definite integrals', 'functions', 'multivariable')

## Textbook tags
## HHChapter1('Integrating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( y=>'Real', z=>'Real' );
$showPartialCorrectAnswers = 1;

## pick a figure
$whichF = list_random(1,2,3,4);
$fname = "q23f$whichF.png";

## the variable from y or z that has one sign
@pvar = ( '', 'y', 'y', 'z', 'z' );
$sgn  = ( $whichF == 2 || $whichF == 3 ) ? '\le' : '\ge';

$len = random(1,3,1);
$r0  = random(1,3,1);
Context()->flags->set( limits=>[-$r0,$r0] );
$r0s = $r0*$r0;

## the sizes of these are
@h = ( '', 338, 387, 244, 249 );
@w = ( '', 288, 270, 334, 324 );
@alt = ( '',
	 'a quarter cylinder extending from the yz-plane along the ' .
	     'positive x-axis, including only that portion of the ' .
	     'cylinder having x and y coordinates both positive.',
	 'a quarter cylinder extending from the yz-plane along the ' .
	     'positive x-axis, including only that portion of the ' .
	     'cylinder having positive x and negative y coordinates.',
	 'a quarter cylinder extending from the yz-plane along the ' .
	     'positive x-axis, consisting of that portion of the ' .
	     'cylinder with z negative.',
	 'a quarter cylinder extending from the yz-plane along the ' .
	     'positive x-axis, consisting of that portion of the ' .
	     'cylinder with z positive.', );

## then the default limits are
if ( $whichF == 1 ) {
    ## the limits for dy dz dx
    %limyzx = ( y => [ Formula("0"), Compute("sqrt($r0s - z^2)") ],
	    	z => [ Formula("-$r0"), Formula("$r0") ],
	    	x => [ Formula("0"), Formula("$len") ] );
    ## and the limits for dz dy dx
    %limzyx = ( z => [ Compute("-sqrt($r0s - y^2)"),
		       Compute("sqrt($r0s - y^2)") ],
		y => [ Formula("0"), Formula("$r0") ],
		x => [ Formula("0"), Formula("$len") ] );
} elsif ( $whichF == 2 ) {
    ## for figure 2, we have the same limits with y values multiplied by -1
    ## the limits for dy dz dx
    %limyzx = ( y => [ Compute("-sqrt($r0s - z^2)"), Formula("0") ],
	    	z => [ Formula("-$r0"), Formula("$r0") ],
	    	x => [ Formula("0"), Formula("$len") ] );
    ## and the limits for dz dy dx
    %limzyx = ( z => [ Compute("-sqrt($r0s - y^2)"),
		       Compute("sqrt($r0s - y^2)") ],
		y => [ Formula("-$r0"), Formula("0") ],
		x => [ Formula("0"), Formula("$len") ] );
} elsif ( $whichF == 3 ) {
    ## for figure 3, those for figure 2 with y and z variables exchanged
    ## the limits for dy dz dx
    %limyzx = ( y => [ Compute("-sqrt($r0s - z^2)"),
		       Compute("sqrt($r0s - z^2)") ],
	    	z => [ Formula("-$r0"), Formula("0") ],
	    	x => [ Formula("0"), Formula("$len") ] );
    ## and the limits for dz dy dx
    %limzyx = ( z => [ Compute("-sqrt($r0s - y^2)"), Formula("0") ],
		y => [ Formula("-$r0"), Formula("$r0") ],
		x => [ Formula("0"), Formula("$len") ] );
} else {
    ## for figure 4, those for figure 1 with y and z variables exchanged
    ## the limits for dy dz dx
    %limyzx = ( y => [ Compute("-sqrt($r0s - z^2)"),
		       Compute("sqrt($r0s - z^2)") ],
	    	z => [ Formula("0"), Formula("$r0") ],
	    	x => [ Formula("0"), Formula("$len") ] );
    ## and the limits for dz dy dx
    %limzyx = ( z => [ Formula("0"), Compute("sqrt($r0s - y^2)") ],
		y => [ Formula("-$r0"), Formula("$r0") ],
		x => [ Formula("0"), Formula("$len") ] );
}

$mp = MultiAnswer( Compute("z"), Compute("y"), Compute("x"),
		   $limzyx{"x"}->[0], $limzyx{"x"}->[1],
		   $limzyx{"y"}->[0], $limzyx{"y"}->[1],
		   $limzyx{"z"}->[0], $limzyx{"z"}->[1] )->with(
    singleResult=>0,
    checker=>sub {
	my( $c, $s, $ansH ) = @_;

	my @res = ( 0,0,0, 0,0, 0,0, 0,0 );

	## variables we've found in the integral
	my %donevar  = ( x => 0, y => 0, z => 0 );
	my $yz = 0;
	my @varorder = ( );

	if ( $c->[0] == $s->[0] || $c->[1] == $s->[0] || $c->[2] == $s->[0] ) {
	    $res[0] = 1;
	    $donevar{$s->[0]->string()} = 1;
	    push( @varorder, $s->[0]->string() );
	    $yz = 1 if ( $s->[0]->string() eq 'y' );
	} else {
	    push( @varorder, 0 );
	}
	if ( ( $c->[0] == $s->[1] && ! $donevar{$c->[0]->string()} ) ||
	     ( $c->[1] == $s->[1] && ! $donevar{$c->[1]->string()} ) ||
	     ( $c->[2] == $s->[1] && ! $donevar{$c->[2]->string()} ) ) {
	    $res[1] = 1;
	    $donevar{$s->[1]->string()} = 1;
	    push( @varorder, $s->[1]->string() );
	    $yz = 1 if ( $s->[1]->string() eq 'y' && ! $donevar{'z'} );
	} else {
	    push( @varorder, 0 );
	}
	if ( ( $c->[0] == $s->[2] && ! $donevar{$c->[0]->string()} ) ||
	     ( $c->[1] == $s->[2] && ! $donevar{$c->[1]->string()} ) ||
	     ( $c->[2] == $s->[2] && ! $donevar{$c->[2]->string()} ) ) {
	    $res[2] = 1;
	    $donevar{$s->[2]->string()} = 1;
	    push( @varorder, $s->[2]->string() );
	} else {
	    push( @varorder, 0 );
	}
	## now @varorder is a list of the correct variables that we have
        ##    in the differential. to check the limits for each we need
	##    to know if the order is yz or zy (the location of the x 
        ##    differential doesn't matter)
        $lim = ( $yz ) ? ~~%limyzx : ~~%limzyx;
	my $i = 7;
	foreach my $v ( @varorder ) {
	    if ( $v ) {
		$res[$i] = ( $lim->{$v}->[0] == $s->[$i] ) ? 1 : 0;
		$res[$i+1] = ( $lim->{$v}->[1] == $s->[$i+1] ) ? 1 : 0;
	    }
	    $i -= 2;
	}

	return [ @res ];
} );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Write limits of integration for the
integral \( \int_W g(x,y,z)\,dV \), where \( W \) is the
quarter cylinder shown, if the length of the cylinder is
$len and its radius is $r0.
$PAR
$BCENTER
\{ image( $fname, tex_size=>200, height=>$h[$whichf], width=>$w[$whichf],
	  extra_html_tags=>'alt="' . $alt[$whichf] . '"' ) \}
$ECENTER

\( \int_W g(x,y,z)\,dV = \int_a^b\int_c^d\int_e^f g(x,y,z)\, d\)\{ $mp->ans_rule(3) \}
\(d\)\{ $mp->ans_rule(3) \}
\(d\)\{ $mp->ans_rule(3) \}
$BR
where
\(a = \) \{$mp->ans_rule(15) \},
\(b = \) \{$mp->ans_rule(15) \},$BR
\(c = \) \{$mp->ans_rule(15) \},
\(d = \) \{$mp->ans_rule(15) \},$BR
\(e = \) \{$mp->ans_rule(15) \}, and
\(f = \) \{$mp->ans_rule(15) \}.
$BR
${BEMPH}(Note: values for all answer blanks must be supplied for
this problem to be able to check the answers provided.)$EEMPH

END_TEXT
Context()->normalStrings;

ANS( $mp->cmp() );

if ( $whichF == 1 || $whichF == 2 ) {
    %lim1 = %limyzx;
    @var1 = ( 'y', 'z', 'x' );
    %lim2 = %limzyx;
    @var2 = ( 'z', 'y', 'x' );
} else {
    %lim1 = %limzyx;
    @var1 = ( 'z', 'y', 'x' );
    %lim2 = %limyzx;
    @var2 = ( 'y', 'z', 'x' );
}

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

A slice through \(W\) for a fixed value of \(x\) is a
semi-circle the boundary of which is \(y^2 + z^2 = $r0s\), for
\($pvar[$whichF]$sgn 0\), so we may take the inner integral
to be
\[
    \int_{$lim1{$var1[0]}->[0]}^{$lim1{$var1[0]}->[1]} g(x,y,z)\,d$var1[0].
\]
Lining up these stacks parallel to \($var1[1]\)-axis gives a slice from
\($var1[1]=$lim1{$var1[1]}->[0]\) to
\($var1[1]=$lim1{$var1[1]}->[1]\) giving
\[
  \int_{$lim1{$var1[1]}->[0]}^{$lim1{$var1[1]}->[1]}
    \int_{$lim1{$var1[0]}->[0]}^{$lim1{$var1[0]}->[1]} g(x,y,z)\,
	d$var1[0]\,d$var1[1].
\]
Finally, there is a slice for each \(x\) between 0 and $len, so the integral
we want is
\[
\int_{0}^{$len}
  \int_{$lim1{$var1[1]}->[0]}^{$lim1{$var1[1]}->[1]}
    \int_{$lim1{$var1[0]}->[0]}^{$lim1{$var1[0]}->[1]} g(x,y,z)\,
	d$var1[0]\,d$var1[1]\,dx.
\]
$PAR
We could also write this as
\[
\int_{0}^{$len}
  \int_{$lim2{$var2[1]}->[0]}^{$lim2{$var2[1]}->[1]}
    \int_{$lim2{$var2[0]}->[0]}^{$lim2{$var2[0]}->[1]} g(x,y,z)\,
	d$var2[0]\,d$var2[1]\,dx.
\]

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
