# DESCRIPTION
# Skills building problem
# ENDDESCRIPTION

## DBsubject(Precalculus)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Level(2)
## KEYWORDS('algebra rational functions, operations with rational expressions, rational expression, fraction, simplify')

DOCUMENT();
loadMacros("PGstandard.pl",
           "MathObjects.pl",
           "PGML.pl",
	   "parserPopUp.pl",
	   "parserMultiAnswer.pl",
# "PGchoicemacros.pl",
# "PGgraphmacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->operators->undefine("/");
Context()->constants->add( h => sqrt(3)/pi );
Context()->flags->set( formatStudentAnswer=>'parsed' );

$a = random(2,9,1);

$ma = MultiAnswer( Compute("-1"), Compute("$a*($a + h)") )->with(
    singleResult => 1,
    separator => ' / ',
    tex_separator => ' / ',
    checker => sub {
        my ( $cor, $stu, $self ) = @_;
	return ($cor->[0] == $stu->[0] && $cor->[1] == $stu->[1]) || \
	  ($cor->[0] == -1*$stu->[0] && $cor->[1] == -1*$stu->[1]);
    }
);

TEXT(beginproblem());
BEGIN_PGML

Simplify the expression 
>> [`` E = \frac{\frac{1}{[$a]+h} - \frac{1}{[$a]}}{h} ``] <<
as much as possible.  Assume that [`h \ne 0`], and write your
answer as a simple proper fraction.

    [` E = (`] [___]{$ma}{15} [` ) / ( `] [___]{$ma}{15} [` ) `]

END_PGML

BEGIN_PGML_SOLUTION

Getting a common denominator in the numerator, dividing, and simplifying,
we find
>> [`` \begin{aligned}
       \frac{\frac{1}{[$a]+h} - \frac{1}{[$a]}}{h}
       &= \frac{\frac{[$a]}{[$a]([$a]+h)} - \frac{[$a]+h}{[$a]([$a]+h)}}{h} \\
       &= \frac{[$a] - ([$a]+h)}{[$a]([$a]+h)h} \\
       &= \frac{-h}{[$a]([$a]+h)h} \\
       &= -\frac{1}{[$a]([$a]+h)}.
    \end{aligned} ``] <<

END_PGML_SOLUTION

ENDDOCUMENT();
