# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions and their Graphs')
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1()
## Problem1()
## Author('Alex Jordan')
## Institution('PCC')



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "PGcourse.pl",

);


TEXT(beginproblem());

##############################################


Context("Numeric");

$r[0] = random(-4,-1,1);
for my $i (1..3)
{
  $r[$i] = $r[$i-1] + random(1,5,1);
}

$p = Compute("(x-$r[0])*(x-$r[1])/((x-$r[2])*(x-$r[3]))")->reduce;

$f = RandomVariableName(type=>'function');

##############################################


BEGIN_PGML

Find a formula for a rational function [`[$f]`] with degree-one zeros at [`[@$r[0]@]`] and [`[@$r[1]@]`], and degree-one vertical asymptotes at [`[@$r[2]@]`] and [`[@$r[3]@]`].

    [`[$f](x)={}`][______________________________________]

END_PGML


##############################################

$showPartialCorrectAnswers = 1;

ANS( $p->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $p");
      return $correct == $student;
    }
) );

##############################################

BEGIN_PGML_SOLUTION

Since this is a rational function, having zeros at [`[@$r[0]@]`] and [`[@$r[1]@]`] means that [`(x-[@$r[0]@])`] and [`(x-[@$r[1]@])`] are factors of the numerator. Since they are degree-1 zeros, there is no exponent on these factors (other than [`1`]).

Having vertical asymptotes at [`[@$r[2]@]`] and [`[@$r[3]@]`] means that [`(x-[@$r[2]@])`] and [`(x-[@$r[3]@])`] are factors of the denominator. Since they are degree-1 asymptotes, there is no exponent on these factors (other than [`1`]).

So far, this tells us that [`[$f](x)=c[$p]`] for some constant [$c]. You can enter any formula that has this pattern to get credit for this problem, including just [`[$p]`] (where [`c=1`]).

END_PGML_SOLUTION

##############################################

ENDDOCUMENT();
