# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Given a quadratic function's formula, build a table and graph it.
# y=ax^2, where |a|>1, and a could be positive or negative.
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Coordinates and Graphs')
## DBsection('Graphs of Equations')
## KEYWORDS('quadratic','function','parabola','table','graph')
## DBCCSS('F.IF.7.a')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGgraphmacros.pl",
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "pccTables.pl",
  "parserMultiAnswer.pl",
  "parserRadioButtons.pl",
  "PCCgraphMacros.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
);

##############################################
TEXT(beginproblem());
Context("Numeric");
Context()->variables->add(y=>'Real');
Context()->flags->set(showExtraParens=>0, reduceConstants=>0);
Context()->noreduce("(-x)-y","(-x)+y");
$refreshCachedImages=1;

#This array stores value of a in y=ax^2.
#$a[0] is the correct graph. The others are used in wrong graphs.
#$a[4] is used in the correct graph.
@a=();
$a[0] = list_random(-5..-2,2..5);
$a[1] = $a[0]-1;
$a[2] = $a[0]+1;
$a[3] = $a[0]+2;
$a[4] = $a[0];

$left = Compute("y");
$right = Compute("$a[0]*x**2")->reduce;

@defaultAnswers = ();
for my $i (0..4) {
  $x[$i] = $i-2;
  $y[$i] = $right->eval(x=>$x[$i]);
  push(@defaultAnswers,$x[$i]);
  push(@defaultAnswers,$y[$i]);
}

$multians = MultiAnswer(@defaultAnswers)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @sxy = @{$student};
      my @xs, @ys;
     for my $k (0..4) 
       {$xs[$k] = $sxy[2*$k];
        $ys[$k] = $sxy[2*$k+1];}
     #
     #  Loop though the student answers
     ##
     for ($i = 0; $i <= 4; $i++) {
       my $ith = Value::List->NameForNumber($i+1);
       #
       #  Check that the student's answers are Reals
       #
       Value->Error("Enter a number for your $ith x-value") if ($xs[$i] eq "");
       Value->Error("Enter a number for your $ith y-value") if ($ys[$i] eq "");
       Value->Error("Your $ith x-value is not a number") if ($xs[$i]->type ne "Number");
       Value->Error("Your $ith y-value is not a number") if ($ys[$i]->type ne "Number");
       #
       #  Check that the point hasn't been given before
       #
       for ($j = 0, $used = 0; $j < $i; $j++) {
         my $jth = Value::List->NameForNumber($j+1);
         if ($xs[$j] == $xs[$i] and $ys[$j] == $ys[$i]) {
           Value->Error("Your $ith row in the table is the same as your $jth row") unless $ansHash->{isPreview};
           $used = 1; last;
         }
       }
       #
       #  If not already used, check that it satisfies the equation
       #    and increase the score if so.
       #
       if (!$used) {
         Value->Error("Your $ith row is not correct") unless $ansHash->{isPreview} or $left->eval(x=>$xs[$i],y=>$ys[$i]) == $right->eval(x=>$xs[$i],y=>$ys[$i]);
         }
       }
     return 1;
});

#Create a hash and its inverse that randomizes the order of the four graphs.
%scramble = ();
%inverse = ();
for my $i (0..3) {
  do {$temp = list_random(0..3)} until !(grep(/^$temp$/,values %scramble) );
  $scramble{$i} = $temp;
  $inverse{$temp} = $i;
};

$scramble{4} = 4;
$inverse{4} = 4;

$xmin = -5;
$xmax = -$xmin;
if ($a[0]>0) {
   $ymin = -5;
   $ymax = 50;
} else {
   $ymin = -50;
   $ymax = 5;
}

#These two variables decide how many labels are on each axis.
$xmark = 1;
$ymark = 5;

@gr=();

#$gr[0] through $gr[3] are options; $gr[4] is used in solution.
for ($i=0;$i<=4;$i++) {
   $gr[$i] = init_graph($xmin,$ymin,$xmax,$ymax,
   axes=>[0,0],
   grid=>[$xmax-$xmin,($ymax-$ymin)/5],
   size=>[240,240]
   );
   $gr[$i]->lb('reset');
   for ($j=-$xmax+$xmark;$j<$xmax;$j=$j+$xmark) {
      $gr[$i]->lb( new Label($j,  -0.1, $j,'black','center','top'));
   }
   for ($j=$ymin+$ymark;$j<$ymax;$j=$j+$ymark) {
      if (Real($j) == 0) {next;}
      $gr[$i]->lb( new Label(0,  $j,' '.$j,'black','left','middle'));
   }

   $gr[$i]->lb( new Label($xmin, $ymax, ' '.$ALPHABET[$inverse{$i}], 'black', 'left', 'top')) unless ($i==4);
   add_functions($gr[$i], "$a[$i]*x**2 for x in <$xmin,$xmax> using color:red and weight:1");
   $xpoint1 = $xmin;
   $xpoint2 = 0;
   $xpoint3 = $xmax;

   $ypoint1 = $a[$i]*($xpoint1**2);
   $ypoint2 = $a[$i]*($xpoint2**2);
   $ypoint3 = $a[$i]*($xpoint3**2);
   if($i==4)
   {
   push(@alt,"Graph $ALPHABET[$inverse{0}]: graph of a parabola passing through the points ($xpoint1,$ypoint1), ($xpoint2,$ypoint2) and ($xpoint3,$ypoint3)");
   }
   else
   {
    push(@alt,"Graph $ALPHABET[$inverse{$i}]: graph of a parabola passing through the points ($xpoint1,$ypoint1), ($xpoint2,$ypoint2) and ($xpoint3,$ypoint3)");
   }


}

$radio = RadioButtons(
  ["A","B","C","D"],
  $ALPHABET[$inverse{0}], # correct answer
  order => ["A","B","C","D"],
);

$f =RandomVariableName(type=>'function');
##############################################

BEGIN_PGML

Make a table of values for the function [`[$f]`], defined by [`[$f](x)=[$right]`]. Based on values in the table, sketch a graph of [`[$f]`].

END_PGML

Context()->texStrings;
BEGIN_TEXT

\{ 
DataTable([[{data=>"\(x\)",header=>"CH"},{data=>"\($f(x)\)",header=>"CH"},],
 [$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
], allcellcss=>"border:1px solid;padding:4pt;text-align:center;", caption => "\($f(x)=$right\)", captioncss =>"padding:4pt; text-decoration:none;");
\}  


$PAR

END_TEXT

Context()->normalStrings;

#Context()->texStrings;
#BEGIN_TEXT
#$PAR
#$BCENTER
#\{ 
#LayoutTable([[image( insertGraph($gr[$scramble{0}]), width=>240, height=>240, tex_size=>TeXscalar(), extra_html_tags=>'alt = "$alt[$scramble{0}]" title = "$alt[$scramble{0}]"'  ),
#image( insertGraph($gr[$scramble{1}]), width=>240, height=>240, tex_size=>TeXscalar(), extra_html_tags=>'alt = "$alt[$scramble{1}]" title = "$alt[$scramble{1}]"'  ),
#],
#[image( insertGraph($gr[$scramble{2}]), width=>240, height=>240, tex_size=>TeXscalar(), extra_html_tags=>'alt = "$alt[$scramble{2}]" title = "$alt[$scramble{2}]"'  ),
#image( insertGraph($gr[$scramble{3}]), width=>240, height=>240, tex_size=>TeXscalar(), extra_html_tags=>'alt = "$alt[$scramble{3}]" title = "$alt[$scramble{3}]"'  )
#]], allcellcss => "padding:20pt;")

#\}
#$ECENTER
#$PAR

#END_TEXT
#Context()->normalStrings;

#BEGIN_PGML
#    The correct graph is graph    
#    [@$radio->buttons()@]*
#END_PGML

##############################################

ANS( $multians->cmp() ); 
#ANS( $radio->cmp() );

##############################################

#Plot some points in the correct solution graph.
for my $i (0..4) {
   $gr[4]->stamps( closed_circle($defaultAnswers[2*$i],   $defaultAnswers[2*$i+1],'black') );
   if ($defaultAnswers[2*$i]>=0) {
      $gr[4]->lb( new Label($defaultAnswers[2*$i],$defaultAnswers[2*$i+1], " ($defaultAnswers[2*$i],$defaultAnswers[2*$i+1]) ",'black','left','top'));
   } else {
      $gr[4]->lb( new Label($defaultAnswers[2*$i],$defaultAnswers[2*$i+1], " ($defaultAnswers[2*$i],$defaultAnswers[2*$i+1]) ",'black','right','top'));
   }
}

sub format_list {
   return "" if !@_;
   my $last = pop(@_);
   return '\('.$last.'\)' if !@_;
   return '\('.join('\), \(', @_) .'\) and \('. $last.'\)';
}

$xvals = format_list(@x);

for my $i (0..4) {
$ls[$i] = $left->substitute(x=>$x[$i]);
$rs[$i] = $right->substitute(x=>$x[$i]);
}

Context()->texStrings;
BEGIN_SOLUTION

This solution will use the \(x\)-values $xvals. The choice to use these \(x\)-values is arbitrary. Since they are small numbers, they might make calculations easier. It's important to include negative numbers if they are in the domain of the function.

One at a time, we substitute these \(x\)-values into the formula for \($f(x)\), and simplify:

\[\begin{aligned}
$f($x[0]) = $rs[0] &\implies $f($x[0])=$y[0]\\
$f($x[1]) = $rs[1] &\implies $f($x[1])=$y[1]\\
$f($x[2]) = $rs[2] &\implies $f($x[2])=$y[2]\\
$f($x[3]) = $rs[3] &\implies $f($x[3])=$y[3]\\
$f($x[4]) = $rs[4] &\implies $f($x[4])=$y[4]\\
\end{aligned}
\]

So the table may be completed as:

\{ 
DataTable([[{data=>"\(x\)",header=>"CH"},{data=>"\($f(x)\)",header=>"CH"},],
 ["\($x[0]\)", "\($y[0]\)"],
["\($x[1]\)", "\($y[1]\)"],
["\($x[2]\)", "\($y[2]\)"],
["\($x[3]\)", "\($y[3]\)"],
["\($x[4]\)", "\($y[4]\)"],
], allcellcss=>"border:1px solid;padding:4pt;text-align:center;width:60pt;", caption => "\($f(x)=$right\)", captioncss =>"padding:4pt; text-decoration:none;");
\}  

#$PAR
#By values in the table, graph \{ $radio->correct_ans() \} is the correct graph.
#$PAR
Here is a graph of [`[$f]`].
$BCENTER
\{ image( insertGraph($gr[4]), width=>xScreen(), height=>yScreen(), tex_size=>TeXscalar(), extra_html_tags=>'alt = "$alt[$scramble{4}]" title = "$alt[$scramble{4}]"'  )  \}
$ECENTER
$PAR
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();


