##DESCRIPTION
##ENDDESCRIPTION

##KEYWORDS('differential equations')

## DBsubject('Differential Equations')
## DBchapter('Introduction')
## DBsection('Classification of Differential Equations')
## Date('05/22/2012')
## Author('Anna Ghazaryan, Fuli Zhang, and Paul Pearson')
## Institution('Miami Univ. Ohio and Hope College')
## TitleText1('Differential equations An introduction to modern methods and applications')
## AuthorText1('James R. Brannan, William E. Boyce')
## Section1('1.4')
## Problem1('6')

########################################################################

DOCUMENT();      

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
#"unionLists.pl",
#"parserPrime.pl",
"parserPopUp.pl",
#"answerHints.pl",
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");


##############################################################
#
#  Text
#
#

$a = random(2,9,1);
do { $b = random(2,9,1); } until ($b != $a);

$answer[0] = Compute("$a");
$answer[1] = PopUp(["Choose","Linear","Nonlinear"],"Linear");

Context()->texStrings;
BEGIN_TEXT


Determine the order of the given differential equation and state whether the equation is linear or nonlinear.

\[\displaystyle\  t\frac{dy}{dt}+ \frac{d^{$a}y}{dt^{$a}}+ (\cos^{2}t)y = t^{$b}\]$BR${BBOLD}You have only one attempt for credit. ${EBOLD}
$BR
$BR
(a) The order of this differential equation is \{ ans_rule(10) \}. 
$BR
$BR
(b) The equation is \{ $answer[1]->menu() \}.



END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
$showPartialCorrectAnswers = 0;


foreach my $i (0..1) {
ANS( $answer[$i]->cmp(
  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( ($numOfAttempts < 1) && ($correct == $student) ) { 
       return 1; 
    } elsif ( ($numOfAttempts > 0) && ($correct == $student) ) {
       Value->Error("Your answer is correct, but your score did not change since you used up the attempts allowed."); 
    } else {
       return 0;
    };
  }
));

}

COMMENT('MathObject version.  Uses a custom answer checker to 
limit the number of attempts that count for credit to one.');

ENDDOCUMENT();    

