##DESCRIPTION
##ENDDESCRIPTION

##KEYWORDS('differential equations')

## DBsubject('Differential Equations')
## DBchapter('Introduction')
## DBsection('Verifying Solutions')
## Date('05/22/2012')
## Author('Anna Ghazaryan, Fuli Zhang, and Paul Pearson')
## Institution('Miami Univ. Ohio and Hope College')
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## Section1('1.1')
## Edition('7')
## Problem1('25')

########################################################################

DOCUMENT();      

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");


##############################################################
#
#  Text
#
#

$answer[1] = PopUp(["Choose","Yes","No"],"Yes");

Context()->texStrings;
BEGIN_TEXT

Is the piecewise-defined function 
\[ y = \left\lbrace \begin{array}{ll} -x^{2}, & \mbox{ if }  x < 0 \\  x^{2}, & \mbox{ if } x \geq 0 \end{array} \right. \]
a solution of the differential equation \(xy'-2y=0\) on \((-\infty,\infty)\)?
 $BR$BR${BBOLD}You have only one attempt for credit. ${EBOLD}
$BR
$BR
\{ $answer[1]->menu() \}



END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
$showPartialCorrectAnswers = 0;

ANS( $answer[1]->cmp(  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( ($numOfAttempts < 1) && ($correct == $student) ) { 
       return 1; 
    } elsif ( ($numOfAttempts > 0) && ($correct == $student) ) {
       Value->Error("Your answer is correct, but your score did not change since you used up the attempts allowed."); 
    } else {
       return 0;
    };
  }
) );

COMMENT('MathObject version.  Uses a custom answer checker to 
limit the number of attempts that count for credit to one.');

ENDDOCUMENT();    

