##DESCRIPTION
##ENDDESCRIPTION


## Edition('7')
## DBsubject(Differential equations)
## DBchapter(Introductory concepts)
## DBsection(Verification of solutions)
## Date(05/22/2012)
## Institution(Miami Univ. Ohio and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, and Paul Pearson)
## Level(2)
## Static(1)
## MO(1)
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## EditionText1('')
## Section1('1.1')
## Problem1('21')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");


##############################################################
#
#  Text
#
#

$answer[1] = PopUp(["Choose","Yes","No"],"Yes");

Context()->texStrings;
BEGIN_TEXT


Is the family of functions 
\( \displaystyle P=\frac{c_{1}e^{t}}{1+c_{1}e^{t}} \)
a solution of the differential equation \(\displaystyle  \frac{dP}{dt} = P(1-P) \)? Assume an appropriate interval \(I\) of definition for the solution.
 $BR$BR${BBOLD}You have only one attempt for credit. ${EBOLD}
$BR
$BR
\{ $answer[1]->menu() \}



END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
$showPartialCorrectAnswers = 0;

ANS( $answer[1]->cmp(  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( ($numOfAttempts < 2) && ($correct == $student) ) { 
       return 1; 
    } elsif ( ($numOfAttempts > 1) && ($correct == $student) ) {
       Value->Error("Your answer is correct, but your score did not change since you used up the attempts allowed."); 
    } else {
       return 0;
    };
  }
) );

COMMENT('Uses a custom answer checker to 
limit the number of attempts that count for credit to one.');

ENDDOCUMENT();    

