##DESCRIPTION
##  Combinatorial game theory: Red-Blue hackenbush
##ENDDESCRIPTION

##KEYWORDS('combinatorial games', 'hackenbush', 'red/blue', 'value')

## DBsubject('Combinatorial game theory')
## DBchapter('Red/Blue Hackenbush')
## DBsection('Game values')
## Date('7/22/2018')
## Author('John Perry')
## Institution('University of Southern Mississippi')


########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$ans = -4;

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

What is the value of the Hackenbush game depicted in this image?
$PAR
\{ image( "blue_flower.svg", width => 200, height => 200 ) \}
$PAR
\{ ans_rule(5) \}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS( Compute($ans) -> cmp(
  checker => sub {
    my ( $correct, $student, $hash ) = @_;
    if ( $student == $ans ) { return 1; }
    else {
      if ( $student > 0 ) { Value->Error( "Try a negative value." ); }
      return 0;
    }
  }
) );

ENDDOCUMENT();        
