##DESCRIPTION
##  Nim: find the best move for a configuration
##ENDDESCRIPTION

##KEYWORDS('game theory', 'nim', 'move')

## DBsubject('Combinatorial Game Theory')
## DBchapter('Nim')
## DBsection('Winning position')
## Date('7/11/2018')
## Author('John Perry')
## Institution('University of Southern Mississippi')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "parserMultiAnswer.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 0;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$value = 0;
@rows = ( 0, 0, 0 );

while ($value == 0) {
  @rows[0] = random(1,4,1);
  @rows[1] = random($rows[0],6,1);
  @rows[2] = random($rows[1],8,1);
  $value = (@rows[0] ^ @rows[1]) ^ @rows[2];
}

$ma = MultiAnswer( 0, 0 )->with(
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $num, $where ) = @{ $student };
    $where -= 1;
    @new_row = ( int(@rows[2]), int(@rows[1]), int(@rows[0]) );
    @new_row[$where] -= $num;
    my ( $a, $b, $c ) = @new_row;
    if ( ( ( int($a) ^ int($b) ) ^ int($c) ) == 0 ) {
      return [1,1];
    } else {
      return [0,0];
    }
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

Suppose a configuration of Nim has three rows of @rows[2], @rows[1], and @rows[0] sticks. What is the best move to make?
$PAR
Remove \{ $ma->ans_rule(5)\} sticks from row \{ $ma->ans_rule(5) \}&nbsp;.

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS( $ma->cmp() );


ENDDOCUMENT();        

