## DESCRIPTION
## Higher order ODEs: constant coefficient second order nonhomogeneous linear ODEs
## ENDDESCRIPTION

## KEYWORDS('differential equations','second order linear ODE','constant coefficients','nonhomogeneous')

## DBsubject('Differential Equations')
## DBchapter('Higher Order Differential Equations')
## DBsection('Second Order Linear Nonhomogeneous')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Notes on Diffy Qs')
## EditionText1('December 9, 2010')
## AuthorText1('Jiri Lebl')
## Section1('2.5')
## Problem1('4')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"parserAssignment.pl",
"unionLists.pl",
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->add(
A=>"Real",B=>"Real",y=>"Real",
'a'=>'Parameter','b'=>'Parameter'
);

$k = random(2,5,1);
$k2 = 2 * $k;
$kk = $k * $k;

$y0 = random(2,9,1);
do { $yprime0 = random(1,5,1); } until ($yprime0 != $y0);

$num = 2*$k*$yprime0+1;
$den = 2*$k**2;

$diffeq = "y'' + $kk y = \cos($k x) + \sin($k x)"; # tex

$yparticular = Compute("-1/$k2 x cos($k x) + 1/$k2 x sin($k x) + a cos($k x) + b sin($k x)");
$yhomogeneous = Compute("A cos($k x) + B sin($k x)");
$yinitvalprob = Compute("-1/$k2 x cos($k x) + 1/$k2 x sin($k x) + $y0 cos($k x) + $num/$den sin($k x)");


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

$ITEM Find a particular solution to the 
nonhomogeneous differential equation \( $diffeq \).
$BR
$BR
\( y_p = \) 
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}


$ITEMSEP
$ITEM Find the most general solution to the 
associated homogeneous differential equation.  
Use \( A \) and \( B \) in your answer to denote 
arbitrary constants.
$BR
$BR
\( y_h = \)
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}


$ITEMSEP
$ITEM Find the solution to the original 
nonhomogeneous differential equation satisfying
the initial conditions \( y(0) = $y0 \) and 
\( y'(0) = $yprime0 \).
$BR
$BR
\( y = \)
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $yparticular->cmp() );

ANS( $yhomogeneous->cmp( 
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my $stu   = Formula($student);

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('A'))==Formula(0) || 
      Formula($stu->D('B'))==Formula(0) 
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(1.24);

    my $a11 = $stu->eval(A=>1,B=>0,x=>$x,y=>0);
    my $a12 = $stu->eval(A=>0,B=>1,x=>$x,y=>0);
 
    my $a21 = $stu->D('x')->eval(A=>1,B=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval(A=>0,B=>1,x=>$x,y=>0);

    my $wronskian = $a11*$a22 - $a21*$a12;

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($wronskian==Real(0));

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stux  = Formula($stu->D('x'));
    my $stuxx = Formula($stu->D('x','x'));
    my $stuDE = Formula($stuxx + $kk*$stu)
    ->with(test_points=>[[1,1,1,0],[1,1,0,0],[2,1,1,0],[1,2,0,0]]);

    return ($stuDE==Formula(0));

}));



ANS( $yinitvalprob->cmp() );


COMMENT("MathObject version.");

ENDDOCUMENT();