## DESCRIPTION
## Differential Equations: Exact Equations
## ENDDESCRIPTION

## Tagged by cmd6a 3/10/06

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Institution(ASU)
## Level(3)
## MO(1)

DOCUMENT() ;        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add(y=>["Real",limits=>[.5,2]]);

TEXT(beginproblem()) ;
$showPartialCorrectAnswers = 1;

do {$b= non_zero_random(-4,4,1);} until ($b!=-1);

$n = random(1,6,1);
$n1=$n+1;
do {$m = random(2,6,1);} until ($m!=$n);
$m1= $m+1;
$p = $m1-$n1;
$p1 = $p + 1;
if  ($n1-$m1==1) {
  $Fxy = Compute("-ln(y)+x/y^$n+$b x");
} else {
  $Fxy = Compute("-y^$p1/$p1+x/y^$n+$b x");
}

$M=Compute("y + $b y^$n1")->reduce;
$minusN = Compute("y^{$m1}+$n x");
$N=-$minusN;

$multians = MultiAnswer($M, $N, $Fxy)->with(
  singleResult => 0,
  allowBlankAnswers=>1,
  checkTypes=>0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $Mstu, $Nstu, $Fxystu ) = @{$student};
      my ( $M, $N, $Fxycorrect) = @{$correct};
      my @score = (0,0,0);
      my $mult = Compute("y^$n1");
      # check types and values of first two inputs
      if ( $M->typeMatch($Mstu) && $N->typeMatch($Nstu) ){
         # check values of first two inputs
         if ( ($Mstu != Formula(0) || $Nstu != Formula(0)) && $Mstu*$N - $Nstu*$M == Formula(0) ) { 
            $score[0] = $score[1] = 1;
         }
      }
      else { return [0,0,0]; }
      # check type and value of third input
      if ( $score[0]==1 && $score[1]==1 && $Fxycorrect->typeMatch($Fxystu) ) {
         my $DFx = $Fxystu->D('x');
         my $DFy = $Fxystu->D('y');
         if ( ($DFx != Formula(0) || $DFy != Formula(0)) 
                && $DFx*$N == $DFy*$M )
         { $score[2]=1; }
      }
      return [@score];
  }
);

BEGIN_PGML

The differential equation 
[``` [$M] = \left([$minusN] \right)y' ```]
can be written in differential form:
[``` M(x,y) \, dx + N(x,y) \, dy  = 0 ```]
where 
    [`M(x,y)= `][__________]{$multians} , and  
    [`N(x,y)= `][__________]{$multians}.  (To receive credit both answers must be correct.)

The left-hand side [` M(x,y) \, dx + N(x,y) \, dy `] becomes an exact differential if it is divided by [` y^[$n1] `]. Integrate the resulting equation to obtain a solution of the differential equation: 
    [_______________]{$multians} [`= C `].

END_PGML

ENDDOCUMENT() ;        # This should be the last executable line in the problem.

