##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Algebra')
## DBchapter('Fundamentals')
## DBsection('Real Numbers')
## Date('6/3/2002')
## Author('')
## Institution('')
## TitleText1('Precalculus')
## EditionText1('3')
## AuthorText1('Stewart, Redlin, Watson')
## Section1('1.1')
## Problem1('22')

########################################################################
#Source Code for calc_compound_problem2.pg:

#This problem is a compound problem with three parts. The student is given an #integral that must be done using integration by parts. In the first part of #the problem the student must give the u and dv functions. In the second part #the student must give the du and v functions. In the third part the student #must #produce the antiderivative.

    DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "compoundProblem.pl",
  "Parser.pl",
  "PGunion.pl",
  "PGgraders.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);









#install_problem_grader(~~&full_partial_grader);
$showPartialCorrectAnswers = 1;

Context()->variables->are(x=>'Real',a=>'Real',b=>'Real',c=>'Real',m=>'Real',u=>'Real',w=>'Real');

$auxeq = Compute("a * m^2 + b * m + c");
$D = Compute("b^2 - 4 * a * c");

$r = random(2,9,1);
$s = random(2,9,1);
$p = random(2,9,1);
$q = random(2,9,1);

$a1 = $r * $s;
$b1 = $r * $q - $p *$s;
$c1 = -1 * $p * $q;

$auxeq1 = Compute("$a1 * m^2 + $b1 * m + $c1");

$a2 = $p * $p;
$b2 = 2 * $p * $s;
$c2 = $s * $s;

$auxeq2 = Compute("$a2 * m^2 + $b2 * m + $c2");

$a3 = $q * $q;
$b3 = 2 * $q * $r;
$c3 = $r * $r + $s * $s;

$auxeq3 = Compute("$a3 * m^2 + $b3 * m + $c3");

$m1 = Compute("(-$b1-sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)");
$m2 = Compute("(-$b1+sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)");


$multipartm = MultiAnswer($m1, $m2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );



$y11 = Formula("e^((-$b1-sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)*x)")->reduce;
$y12 = Formula("e^((-$b1+sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)*x)")->reduce;


$multipart2 = MultiAnswer($y11, $y12)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct2, $student2, $self2 ) = @_;
        my ( $f1stu2, $f2stu2 ) = @{$student2};
        my ( $f12, $f22 ) = @{$correct2};
        if ( ($f12 == $f1stu2 && $f22 == $f2stu2) ||
             ($f12 == $f2stu2 && $f22 == $f1stu2) ) {
            return [1,1];
        } else {
            if ($f12 == $f1stu2 || $f22 == $f1stu2) {
                return [1,0];
            } elsif ($f12 == $f1stu2 || $f22 == $f2stu2) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


$multipart2b = MultiAnswer($y11, $y12)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct2b, $student2b, $self2b ) = @_;
        my ( $f1stu2b, $f2stu2b ) = @{$student2b};
        my ( $f12b, $f22b ) = @{$correct2b};
        if ( ($f12b == $f1stu2b && $f22b == $f2stu2b) ||
             ($f12b == $f2stu2b && $f22b == $f1stu2b) ) {
            return [1,1];
        } else {
            if ($f12b == $f1stu2b || $f22b == $f1stu2b) {
                return [1,0];
            } elsif ($f12b == $f1stu2b || $f22b == $f2stu2b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );





$P = Compute("b/a");
$y41 = Compute("e^(-x*b/(2*a))");
$g = Compute("e^(-x*b/a)");
$xy41 = Compute("x*e^(-x*b/(2*a))");

$multipart3 = MultiAnswer($y41, $xy41)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct3, $student3, $self3 ) = @_;
        my ( $f1stu3, $f2stu3 ) = @{$student3};
        my ( $f13, $f23 ) = @{$correct3};
        if ( ($f13 == $f1stu3 && $f23 == $f2stu3) ||
             ($f13 == $f2stu3 && $f23 == $f1stu3) ) {
            return [1,1];
        } else {
            if ($f13 == $f1stu3 || $f23 == $f1stu3) {
                return [1,0];
            } elsif ($f13 == $f1stu3 || $f23 == $f2stu3) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );







$m5 = Compute("-$b2/(2*$a2)");
$y51 = Compute("e^(-$b2/(2*$a2)*x)");
$xy51 = Compute("x*e^(-$b2/(2*$a2)*x)");


$multipart4 = MultiAnswer($y51, $xy51)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct4, $student4, $self4 ) = @_;
        my ( $f1stu4, $f2stu4 ) = @{$student4};
        my ( $f14, $f24 ) = @{$correct4};
        if ( ($f14 == $f1stu4 && $f24 == $f2stu4) ||
             ($f14 == $f2stu4 && $f24 == $f1stu4) ) {
            return [1,1];
        } else {
            if ($f14 == $f1stu4 || $f24 == $f1stu4) {
                return [1,0];
            } elsif ($f14 == $f1stu4 || $f24 == $f2stu4) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );



$multipart4b = MultiAnswer($y51, $xy51)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct4b, $student4b, $self4b ) = @_;
        my ( $f1stu4b, $f2stu4b ) = @{$student4b};
        my ( $f14b, $f24b ) = @{$correct4b};
        if ( ($f14b == $f1stu4b && $f24b == $f2stu4b) ||
             ($f14b == $f2stu4b && $f24b == $f1stu4b) ) {
            return [1,1];
        } else {
            if ($f14b == $f1stu4b || $f24b == $f1stu4b) {
                return [1,0];
            } elsif ($f14b == $f1stu4b || $f24b == $f2stu4b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );






$y6g1 = Compute("e^{-b*x/(2*a)}*sin(x*sqrt{4*a*c-(b)^2}/(2*a))");
$y6g2 = Compute("e^{-b*x/(2*a)}*cos(x*sqrt{4*a*c-(b)^2}/(2*a))");

$y6e1 = Compute("e^{-$b3*x/(2*$a3)}*sin(x*sqrt{4*$a3*$c3-($b3)^2}/(2*$a3))");
$y6e2 = Compute("e^{-$b3*x/(2*$a3)}*cos(x*sqrt{4*$a3*$c3-($b3)^2}/(2*$a3))");

$multipart5 = MultiAnswer($y6e1, $y6e2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct5, $student5, $self5 ) = @_;
        my ( $f1stu5, $f2stu5 ) = @{$student5};
        my ( $f15, $f25 ) = @{$correct5};
        if ( ($f15 == $f1stu5 && $f25 == $f2stu5) ||
             ($f15 == $f2stu5 && $f25 == $f1stu5) ) {
            return [1,1];
        } else {
            if ($f15 == $f1stu5 || $f25 == $f1stu5) {
                return [1,0];
            } elsif ($f15 == $f1stu5 || $f25 == $f2stu5) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );



$multipart5b = MultiAnswer($y6e1, $y6e2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct5b, $student5b, $self5b ) = @_;
        my ( $f1stu5b, $f2stu5b ) = @{$student5b};
        my ( $f15b, $f25b ) = @{$correct5b};
        if ( ($f15b == $f1stu5b && $f25b == $f2stu5b) ||
             ($f15b == $f2stu5b && $f25b == $f1stu5b) ) {
            return [1,1];
        } else {
            if ($f15b == $f1stu5b || $f25b == $f1stu5b) {
                return [1,0];
            } elsif ($f15b == $f1stu5b || $f25b == $f2stu5b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );



$u61 =Compute("sqrt{4*a*c-(b)^2}/(2*a)");
$u62 =Compute("sin(x*sqrt{4*a*c-(b)^2}/(2*a))");
$u63 =Compute("cos(x*sqrt{4*a*c-(b)^2}/(2*a))");

    TEXT(beginproblem);

    ##############################################

    #This gives the professor (in this case jmarin) the ability to "go backwards" in 
    #a compound problem - useful for debugging.
    $isProfessor = ($studentLogin eq 'admin' || $studentLogin eq 'professor');

    #
    #  Start a compound problem.  See the compoundProblem.pl
    #  file for more details about the parameters you
    #  can supply.
    #
    $cp = new compoundProblem(
      parts => 6,                  # the total number of parts in this problem
      totalAnswers => 48,           # total answers in all parts combined
      parserValues => 1,           # make MathObjects from student answers
      allowReset => $isProfessor,  # professors get Reset button for testing
    );
    $part = $cp->part;             # look up the current part


    ##############################################
    #
    #  Part 1
    #

    if ($part == 1) {

    BEGIN_TEXT
$PAR
This problem has multiple parts.  After you submit your answers, you will be given the option to move on to the next part.  I suggest that you print each page so that you can refer to them if necessary.

$PAR 
Throughout this WeBWork problem, let \(a,b,\) and \(c\) be real numbers with \(a \neq 0\).  In this exercise we will discover a method that will allow us to solve linear homogeneous differential equations of the form
\[ a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\]
$PAR 

To solve this equation, we will assume that the solution \(y\) of the differential equation is of the form \(y=e^{mx}\) where \(m\) is some number that we need to find. If we substitute \(y=e^{mx}\) into 
\( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\), we find that 
$BR
$BCENTER
\{ BeginTable(center=>0).
      Row(['\(0=a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy\)', ],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 48pt =e^{mx}\Big(\)',ans_rule(15),'\(\Big)\)', ],separation=>1).
   EndTable();
\}
$ECENTER
$BR

    END_TEXT

    ANS($auxeq->cmp());


    #
    #  the useGrader() method sets the grader for this part. The student is only 
    #given credit for each part of the problem if all answer blanks in that part are 
    #correct (so no partial credit within parts of the problem).
    #
    $cp->useGrader(~~&std_problem_grader);
    $showPartialCorrectAnswers = 0;
    #Students are not shown which answer is incorrect.

    }

    ##############################################
    #
    #  Part 2
    #

    if ($part == 2) {

    BEGIN_TEXT


Good! If we assume that \(y=e^{mx}\) is a solution to the differential equation  
\( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\), we find that 
\[
\begin{aligned}
0&=a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy\\ 
 &\\
 &=e^{mx}\left(am^2+bm+c\right).
\end{aligned}
\]

Since \(e^{mx}\) never equals zero; it follows that \(am^2+bm+c=0\).  The equation \[am^2+bm+c=0\] is called the $BBOLD auxiliary equation $EBOLD for the differential equation \[ a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\]
$PAR

$BBOLD Let's look at some examples. $EBOLD
$PAR

$BBOLD Example 1: $EBOLD The auxiliary equation for the differential equation
\[ $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\] is 
\{ ans_rule(15) \}=\{ ans_rule(1) \}.
$PAR

$BBOLD Example 2: $EBOLD The auxiliary equation for the differential equation
\[ $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\] is  
\{ ans_rule(15) \}=\{ ans_rule(1) \}.
$PAR

$BBOLD Example 3: $EBOLD The auxiliary equation for the differential equation
\[ $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\] is  
\{ ans_rule(15) \}=\{ ans_rule(1) \}.






    END_TEXT

ANS($auxeq1->cmp());
ANS(Compute("0")->cmp());
ANS($auxeq2->cmp());
ANS(Compute("0")->cmp());
ANS($auxeq3->cmp());
ANS(Compute("0")->cmp());


    }

    ##############################################
    #
    #  Part 3
    #

    if ($part == 3) {

    BEGIN_TEXT

So, if we assume that \(y=e^{mx}\) is a solution to the differential equation  
\( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\), we arrive at the auxiliary equation \(am^2+bm+c=0.\)

$PAR

From the quadratic formula, we know that \(am^2+bm+c=0\) when 
\{ BeginTable(center=>0).
      Row(['\(m = \)', ans_rule(3).'\(\pm\sqrt{}\)'.ans_rule(10).$BR."-----------------------------------------".$BR.'\(\hskip 30 pt\)'.ans_rule(3)],separation=>1).
   EndTable();
\}

$PAR

We will need to consider three different cases; that is, when  
\(b^2-4ac > 0, b^2-4ac = 0\), and \(b^2-4ac<0.\)
$PAR
$BBOLD Case 1:$EBOLD If \( b^2-4ac>0\), then we have found two linearly independent solutions 
\(y_1(x)= e^{\displaystyle{\frac{-b-\sqrt{b^2-4ac}}{2a}}x}\) 
and
\(y_2(x)= e^{\displaystyle{\frac{-b+\sqrt{b^2-4ac}}{2a}}x}\) 
to the differential equation \( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\)  In this case, the general solution to \[ a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\] is the function \[y= c_1 y_1(x) + c_2 y_2(x)\] where \(c_1\) and \(c_2\) are arbitrary constants.
$PAR
$BBOLD Example 4 :$EBOLD
In Example 1, we found that the auxiliary equation of the differential equation 
\[ $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\] is 
\( $a1 m^2+ $b1 m+ $c1 =0.\) (Notice that \( b^2-4ac>0\).)
$BR
Since \( $a1 m^2+ $b1 m+ $c1 =0\), we have that
\(m=\)\{ ans_rule(25) \} 
and 
\(m=\)\{ ans_rule(25) \}.
$BR
In this case, 
\(y_1(x)=\)\{ ans_rule(15) \}
and 

\(y_2(x)=\)\{ ans_rule(15) \}
$BR
are linearly independent solutions to 
\( $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\) and
$BR 
\(y = c_1 \)\{ ans_rule(15) \}\( + c_2 \)\{ ans_rule(15) \} where \(c_1\) and \(c_2\) are arbitrary constants 
is the general solution to \( $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\).

    END_TEXT

ANS(Compute("-b")->cmp());
ANS($D->cmp());
ANS(Compute("2 * a")->cmp());
ANS($multipartm->cmp());
ANS($multipart2->cmp());
ANS($multipart2b->cmp());
#ANS($m1->cmp());
#ANS($m2->cmp());
#ANS($y11->cmp());
#ANS($y12->cmp());
#ANS($y11->cmp());
#ANS($y12->cmp());

    }

   ##############################################
    #
    #  Part 4
    #

    if ($part == 4) {

    BEGIN_TEXT
So, if we assume that \(y=e^{mx}\) is a solution to the differential equation  
\( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\), we arrive at the auxiliary equation \(am^2+bm+c=0.\)

$PAR

From the quadratic formula, we know that \(am^2+bm+c=0\) when 
\[m=\frac{-b\pm\sqrt{b^2-4ac}}{2a}\]
$PAR

We now consider the case when  
\(b^2-4ac = 0.\)
$PAR
$BBOLD Case 2:$EBOLD If \( b^2-4ac=0\), then the only solution to the auxiliary equation \(am^2+bm+c=0\) is 
\(m=\)\{ ans_rule(5) \} and in this case, we get only one solution 
\(y_1(x)=\) \{ ans_rule(15) \}
to the differential equation \( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\)
$PAR

Since a second order linear homogeneous differential equation must have two linearly independent solutions, we need a second solution that is linearly independent to \(y_1\).
$PAR

Recall that if \( y_1(x)\) is a solution to the differential equation 
\[ a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\]
then if \( P(x)=a_1(x)/a_2(x)\), a second solution to  
\( a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\) is given by 
\[
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx.
\]
Furthermore, \(y_1\) and \(y_2\) are linearly independent.

$PAR

For the differential equation 
\[ a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0,\]
\(P=\)\{ ans_rule(15) \}, 
\(e^{-\int P(x)dx}=\)\{ ans_rule(15) \}, and 
\(
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx
=\)\{ ans_rule(15) \}.


$PAR
So, if \( b^2-4ac=0\), then we have found two linearly independent solutions
$BR
\(y_1(x)=\)\{ ans_rule(15) \}
and
\(y_2(x)=\)\{ ans_rule(15) \}
to the differential equation \( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\)  In this case, the general solution to the differential equation 
\[ a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\]
is 
$BR
$BCENTER
\(y(x) = c_1\)\{ ans_rule(15) \}\(+c_2\)\{ ans_rule(15) \}
$ECENTER
$BR
$PAR
$BBOLD Example 5 :$EBOLD
In Example 2, we found that the auxiliary equation of the differential equation 
\[ $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\] is 
\( $a2 m^2+ $b2 m+ $c2 =0.\) (Notice that \( b^2-4ac=0\).)
$BR
Since \( $a2 m^2+ $b2 m+ $c2 =0\), we have that
\(m=\)\{ ans_rule(5) \}.
$BR
In this case, 
\(y_1(x)=\)\{ ans_rule(15) \}
and 
\(y_2(x)=\)\{ ans_rule(15) \}
$BR
are linearly independent solutions to 
\( $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\) and 
$BR
$BCENTER
\(y = c_1\) \{ ans_rule(15) \}\( + c_2 \)\{ ans_rule(15) \}
$ECENTER
$BR
where \(c_1\) and \(c_2\) are arbitrary constants 
is the general solution to \( $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\).





    END_TEXT
ANS(Compute("-b/(2a)")->cmp());
ANS($y41->cmp());
ANS($P->cmp());
ANS($g->cmp());
ANS($xy41->cmp());
ANS($y41->cmp());
ANS($xy41->cmp());
ANS($multipart3->cmp());
#ANS($y41->cmp());
#ANS($xy41->cmp());
ANS($m5->cmp());
ANS($multipart4->cmp());
#ANS($y51->cmp());
#ANS($xy51->cmp());
ANS($multipart4b->cmp());
#ANS($y51->cmp());
#ANS($xy51->cmp());
    }


   ##############################################
    #
    #  Part 5
    #

    if ($part == 5) {

    BEGIN_TEXT
Before we tackle the last case, I want to remind you of a very useful identity from complex analysis.  Don't forget that in complex analysis 
\(i = \sqrt{-1}\).
$PAR

The formula I want to remind you of is Euler's formula, which states that 
\[e^{ix}=\cos(x)+i\sin(x).\]
$PAR

So, if \(u\) and \(w\) are real numbers, 
\{ BeginTable(center=>0).
      Row(['\(e^{u+iw}\)','\(=\)','\(e^{u}e^{iw}\)'],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 22pt\)','\(=\)','\(e^{u}\Big(\)',ans_rule(15),'\(+i\)',ans_rule(15),'\(\Big)\)'],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 22pt\)','\(=\)',ans_rule(15),'\(+i\)',ans_rule(15),'\(\Big)\)'],separation=>1).
   EndTable();
\}
$BBOLD
As this last equation will be used on the next page, please make sure that you write it down.
$EBOLD
END_TEXT
ANS(Compute("cos(w)")->cmp());
ANS(Compute("sin(w)")->cmp());
ANS(Compute("e^(u)*cos(w)")->cmp());
ANS(Compute("e^(u)*sin(w)")->cmp());

}


   ##############################################
    #
    #  Part 6
    #

    if ($part == 6) {

    BEGIN_TEXT
So, if we assume that \(y=e^{mx}\) is a solution to the differential equation  
\( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\), we arrive at the auxiliary equation \(am^2+bm+c=0.\)

$PAR

From the quadratic formula, we know that \(am^2+bm+c=0\) when 
\[m=\frac{-b\pm\sqrt{b^2-4ac}}{2a}\]
$PAR

We now consider the case when  
\(b^2-4ac < 0.\)
$PAR
$BBOLD Case 3:$EBOLD If \( b^2-4ac<0\), then the solutions to the auxiliary equation \(am^2+bm+c=0\) are 
\(m=\displaystyle{\frac{-b-\sqrt{b^2-4ac}}{2a}}\) and
\(m=\displaystyle{\frac{-b+\sqrt{b^2-4ac}}{2a}}\).  For now let's focus on what happens if \(m=\displaystyle{\frac{-b+\sqrt{b^2-4ac}}{2a}}\).

Notice that
\[
\begin{aligned}
   m &=\frac{-b+\sqrt{b^2-4ac}}{2a}
     &\\
     &=\frac{-b}{2a}+\frac{\sqrt{-1(4ac-b^2)}}{2a}\\ 
     &\\
\end{aligned}
\] 
\(\hskip 30pt=\)\{ ans_rule(15) \} \(+i\) \{ ans_rule(15) \}.
So,
$BR
$BCENTER
  \(y = e^{mx}\)
$BR
   \(\hskip 20pt =e^{\frac{-b}{2a}x}\Big(\) \{ ans_rule(25) \}\(\Big)
     +ie^{\frac{-b}{2a}x}\Big(\) \{ ans_rule(25) \} \(\Big)\)
$BR
\(\hskip 20pt =y_1(x)+iy_2(x).\)
$ECENTER
$BR


$PAR
If \( b^2-4ac<0\), it turns out that  
$BR
\(y_1(x)=\)\{ ans_rule(30) \}
$BR
and
$BR
\(y_2(x)=\)\{ ans_rule(30) \}
$BR
are linearly independent solutions 
to the differential equation \( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\) 
(You should check this.) 
Since a fundamental set for a second order linear differential equation has two linearly independent solutions, \(\lbrace y_1, y_2  \rbrace\) forms a fundamental set for 
\( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\) when \( b^2-4ac<0\).

If \( b^2-4ac<0\), the general solution to the differential equation 
\[ a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0\]
is 
$BR
$BCENTER
\(y(x) = c_1\)\{ ans_rule(30) \}\(+c_2\\)\{ ans_rule(30) \}
$ECENTER
$BR
$PAR
$BBOLD Example 5 :$EBOLD
In Example 3, we found that the auxiliary equation of the differential equation 
\[ $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\] is 
\( $a3 m^2+ $b3 m+ $c3 =0.\) (Notice that \( b^2-4ac<0\).)
$BR
Since \( $a3 m^2+ $b3 m+ $c3 =0\), we have that
\(m=\)\{ ans_rule(5) \}\(\pm i\) \{ ans_rule(5) \}.
$BR
In this case, 
$BR
\(y_1(x)=\)\{ ans_rule(15) \}
and 
\(y_2(x)=\)\{ ans_rule(15) \}
$BR
are linearly independent solutions to 
\( $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\) and 
$BR
$BCENTER
\(y = c_1 \)\{ ans_rule(15) \}\(+ c_2 \)\{ ans_rule(15) \}
$ECENTER
$BR
where \(c_1\) and \(c_2\) are arbitrary constants 
is the general solution to \( $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\).
$PAR

$BBOLD
The material in this webwork assignment is covered in Section 4.3 of our text.
$EBOLD





    END_TEXT
ANS(Compute("-b/(2*a)")->cmp());
ANS(fun_cmp( $u61,var=>['a','b','c'],limits=>[[10,20],[0, 1],[10,20]]));
ANS(fun_cmp( $u63,var=>['a','b','c','x'],limits=>[[10,20],[0, 1],[10,20],[-10,10]]));
ANS(fun_cmp( $u62,var=>['a','b','c','x'],limits=>[[10,20],[0, 1],[10,20],[-10,10]]));
ANS(fun_cmp( $y6g2,var=>['a','b','c','x'],limits=>[[10,20],[0, 1],[10,20],[-10,10]]));
ANS(fun_cmp( $y6g1,var=>['a','b','c','x'],limits=>[[10,20],[0, 1],[10,20],[-10,10]]));
ANS(fun_cmp( $y6g2,var=>['a','b','c','x'],limits=>[[10,20],[0, 1],[10,20],[-10,10]]));
ANS(fun_cmp( $y6g1,var=>['a','b','c','x'],limits=>[[10,20],[0, 1],[10,20],[-10,10]]));
ANS(Compute("-$b3/(2*$a3)")->cmp());
ANS(Compute("sqrt{4*$a3*$c3-($b3)^2}/(2*$a3)")->cmp());
ANS($multipart5->cmp());
ANS($multipart5b->cmp());

    }











    ##############################################

    ENDDOCUMENT();        # This should be the last executable line in the problem.


