# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 12/19/08

## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Conceptual understanding of derivatives)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(AppsInterpretDeriv)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Applied Calculus')
## TitleText5('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## AuthorText5('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('4')
## EditionText5('7')
## Section1('2.4')
## Section2('2.4')
## Section3('2.4')
## Section4('2.Review')
## Section5('2.4')
## Problem1('1')
## Problem2('2')
## Problem3('4')
## Problem4('31')
## Problem5('4')
## KEYWORDS('calculus','derivative')

## Textbook tags
## HHChapter1('Key Concept: The Derivative')
## HHChapter2('Key Concept: The Derivative')



DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserNumberWithUnits.pl",
  "parserPopUp.pl",
  "parserAutoStrings.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

Context("Numeric");
# AutoStrings();
DefineStrings("degC","min","degC/min","min/degC","hr","degC/hr","hr/degC");
$showPartialCorrectAnswers = 1;

# include javascript for pop-up explaining units
# HEADER_TEXT(<<EOF);
# <script language="javascript" type="text/javascript">
# <!-- //
# function unitspopup() {
#     var url = "/webwork2_files/units.html";
#     var opt = "height=625,width=600,location=no,menubar=no,status=no" +
#               "resizable=yes,scrollbars=yes,toolbar=no,";
#     window.open(url,'examdata_info',opt).focus();
# }
# // -->
# </script>
# EOF

$t = random(5,35,5);
$f = random(50,80,1);
$fp = random(0.1,2,0.1);
$posneg = PopUp( [ "?", "positive", "negative" ], "negative" );
# $units = NumberWithUnits( "1", "degC/min" );
$units = String("degC/min");
$temp = PopUp( ["?", "derivative", "temperature",
		"change in temperature" ], "temperature" );
$inc = PopUp( ["?", "increase", "decrease"], "decrease" );
$dt = random(0.25,1.5,0.25);
$dtsec = $dt*60;

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
The temperature, \( H \), in degrees Celsius, of a cup of coffee
placed on the kitchen counter is given by \( H = f(t) \), where
\( t \) is in minutes since the coffee was put on the counter.
$PAR
${BBOLD}(a)$EBOLD
Is \( f'(t) \) positive or negative?
\{ $posneg->menu() \}
$BR
${BITALIC}(Be sure that you are able to give a reason for your
answer.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
What are the units of \( f'($t) \)?
\{ ans_rule(15) \} \{ AnswerFormatHelp("units") \}
$BR$BR

Suppose that \( |f'($t)| = $fp \) and \( f($t) = $f \).  Fill in the
blanks (including units where needed) and select the appropriate terms
to complete the following statement about the temperature of the coffee
in this case.
$BR
At \{ ans_rule(5) \} minutes after the coffee was put on the counter,
its \{ $temp->menu() \} is \{ ans_rule(10) \} and will \{ $inc->menu() \}
by about \{ ans_rule(10) \} in the next $dtsec seconds.
$PAR
${BBOLD}Note: If you are using MathQuill click the textbox (Tt) button before entering an answer that contains units.${EBOLD} 


END_TEXT
Context()->normalStrings;

ANS( $posneg->cmp() );
# ANS( $units->cmp( checker=>sub {
#     my ( $correct, $student, $ansHash ) = @_;
#     return ( NumberWithUnits( "1", "degC/min" ) == 
# 	     NumberWithUnits( "1", $student ) ) ? 1 : 0;
# } ) );
ANS( $units->cmp() );
ANS( Compute( $t )->cmp() );
ANS( $temp->cmp() );
ANS( NumberWithUnits( $f, "degC" )->cmp() );
ANS( $inc->cmp() );
ANS( NumberWithUnits( "$fp*$dt", "degC" )->cmp() );

$delf = $fp*$dt;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
As the cup of coffee cools, the temperature decreases, so
\(f'(t)\) is negative.

$PAR
${BBOLD}(b)$EBOLD
Since \(f'(t) = dH/dt\), the units are degrees Celsius per
minute (degC/min).  The quantity \(f'($t)\) represents the rate
at which the coffee is cooling, in degrees per minute, $t minutes
after the cup is put on the counter.  Thus we can complete the statement:
$BR
At $t minutes after the coffee was put on the counter, its temperature
is $f degC and will decrease by about $delf degC in the next $dtsec
seconds.

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
