## DBsubject(Abstract algebra)
## DBchapter(Rings)
## DBsection(Ring axioms)
## Institution(University of Massachusetts Amherst)
## Author(Daniel Nichols and Siman Wong)
## Level(3)
## MO(1)
## KEYWORDS('ring axioms, inverse')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "algebraMacros.pl",
  "PGcourse.pl"
);

Context()->variables->add( n => 'Real' );

TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

# pick 2 random moduli
@chooseModuli = ( 6 .. 50 );
$mod1 = splice( @chooseModuli, random( 0, @chooseModuli - 1, 1 ), 1 );
$mod2 = splice( @chooseModuli, random( 0, @chooseModuli - 1, 1 ), 1 );

$ans[0] = Compute( $mod1 - 1 );
$ans[1] = Compute( $mod2 - 1 );
$ans[2] = Formula( "n - 1" );


TEXT(EV3(<<'EOT'));

(a) Find the multiplicative inverse of \( \{ $mod1 - 1 \} \) in \( \{ cyclic( $mod1 ) \} \). $BR
\{ ans_rule( 10 ) \} $BR $BR

(b) Find the multiplicative inverse of \( \{ $mod2 - 1 \} \) in \( \{ cyclic( $mod2 ) \} \). $BR
\{ ans_rule( 10 ) \} $BR $BR

(c) In general, what is the multiplicative inverse of \( ( n - 1 ) \) in \( \{ cyclic( n ) \} \)? $BR
\{ ans_rule( 10 ) \} $BR $BR

(d) (Optional: see your instructor) Prove that your answer for (c) is correct. $BR

EOT

# we can't use modChecker from algebraMacros.pl because each answer uses a different modulus

# Check answers using modular arithmetic
ANS( $ans[0]->cmp(
	checker => sub{
		my ( $correct, $student, $ansHash ) = @_;  # get correct and student MathObjects
		return ( $student % $mod1 == $correct % $mod1 ? 1 : 0 );
	}
) );

# Check answers using modular arithmetic
ANS( $ans[1]->cmp(
	checker => sub{
		my ( $correct, $student, $ansHash ) = @_;  # get correct and student MathObjects
		return ( $student % $mod2 == $correct % $mod2 ? 1 : 0 );
	}
) );

ANS( $ans[2]->cmp );


SOLUTION(EV3(<<'END_SOLUTION'));
$SOLUTION

For any integer \( n \), we have $BR
\( (n - 1 ) \times ( n - 1 ) = ( n - 1 )^2 \) $BR
\( \phantom{ (n - 1 ) \times ( n - 1 ) } = n^2 + 2n + 1 \) $BR
\( \phantom{ (n - 1 ) \times ( n - 1 ) } \equiv 1 \pmod{n} \) $BR
So \( n - 1 \) is its own inverse in \( \{ cyclic( n ) \} \). $BR

END_SOLUTION

COMMENT( "Complete solution included, possible accompanying written part" );


ENDDOCUMENT(); 
