## DESCRIPTION
## Algebra: finding an equation of a line in slope-intercept form.
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: slope-intercept form)
## Date(12/22/2016)
## Institution(Red Rocks Community College, Colorado Community College System)
## Author(Brenda Forland; Edit: James Morski 7/17/2017, and Steven Wray, 8/2/17)
## Static(1)
## MO(1)
## KEYWORDS('algebra', 'slope-intercept')


###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserAssignment.pl",
"contextFraction.pl",
"AnswerFormatHelp.pl",
"PGML.pl",
"PGcourse.pl",
"contextArbitraryString.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


##########################
#  Setup

Context("Fraction")->variables->are(x=>"Real",y=>"Real");
parser::Assignment->Allow;
parser::Assignment->Function("f");

$a = random(-10,10,1);
$b = random(-9,9,1);
$c = random(3,7,1);
$d = non_zero_random(-4,4,1);

$slope = - $c;
$intercept = $a*$c+$b;

Context("ArbitraryString");

if ($intercept > 0) {
  $ans1 = "y=" . $slope . "x+" . $intercept;
  $ans2 = "y=" . $slope . "*x+" . $intercept;
  $ans3 = "y=(" . $slope . ")x+" . $intercept;
  $ans4 = "y=(" . $slope . ")*x+" . $intercept;
} 
else {
  $ans1 = "y=" . $slope . "x-" . (-1*$intercept);
  $ans2 = "y=" . $slope . "*x-" . (-1*$intercept);
  $ans3 = "y=(" . $slope . ")x-" . (-1*$intercept);
  $ans4 = "y=(" . $slope . ")*x-" . (-1*$intercept);
  $ans5 = "y=" . $slope . "x+-" . (-1*$intercept);
  $ans6 = "y=" . $slope . "*x+-" . (-1*$intercept);
  $ans7 = "y=(" . $slope . ")x+(-" . (-1*$intercept) . ")";
  $ans8 = "y=(" . $slope . ")*x+(-" . (-1*$intercept) . ")";
  $ans9 = "y=(" . $slope . ")x+-" . (-1*$intercept);
  $ans10 = "y=(" . $slope . ")*x+-" . (-1*$intercept);
}

$video = MODES(
HTML=> 
'<iframe width="420" height="315" src="//www.youtube.com/embed/H98iBqt1fFQ" frameborder="0" allowfullscreen></iframe>', 
TeX =>
"An embedded YouTube video."
);



#################################
#  Main text


BEGIN_PGML

Find an equation of the line perpendicular to [` [$c]*y=x-[$d] `] that passes through the point [`([$a],[$b])`]. 

Write the answer in slope-intercept form, [`y = mx + b `], with all fractions written in lowest terms.

[____________________]{$ans} [@ AnswerFormatHelp("equation") @]*

[$video]*
END_PGML
BEGIN_TEXT
$BR

If the video does not work, \{ htmlLink ("https://www.youtube.com/watch?v=H98iBqt1fFQ", "click here to watch on YouTube.", "TARGET='_blank'")\}
$BR
END_TEXT


#################################
#  Solution

#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION

ANS( String($ans1) -> cmp( checker => sub{
    my ($correct, $student, $ansHash) = @_;
    $s = $student;
    $s =~ s/ //g;      # compress all blank space
    if ($intercept > 0) {
      if (String($s) == String($ans1) ||
        String($s) == String($ans2)  ||
        String($s) == String($ans3) ||
        String($s) == String($ans4) 
      ) {
        return(1);
      }
      else {
        return(0);
      }
    }
   else {   # intercept is negative
     if (String($s) == String($ans1) ||
        String($s) == String($ans2) ||
        String($s) == String($ans3) ||
        String($s) == String($ans4) ||
        String($s) == String($ans5) ||
        String($s) == String($ans6) || 
        String($s) == String($ans7) ||
        String($s) == String($ans8) 
      ) {
        return(1);
      }
      else {
        return(0);
      }
    }
} ) );

COMMENT('Uses PGML.');

ENDDOCUMENT();