##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(LU factorization)
## Date(6/3/2013)
## Author(Darwyn Cook)
## MLT(LDU2x2bracket)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.4')
## Problem1('28')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Matrix");
$sample = Matrix([[2,3,-1],[-2,-5,6]]);

$L = Matrix([[1,0],
[Compute(random(-9,9)),1]]);

$D = Matrix([[Compute(non_zero_random(-9,9)),0],[0,Compute(non_zero_random(-9,9))]]);
$U = Matrix([[1,Compute(random(-9,9))],[0,1]]);
 
$A = $L * $D*$U;
$U2 = $D*$U;

$multians = MultiAnswer($L,$D,$U)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $Lstu, $Dstu, $Ustu ) = @{$student};
      my ( $Lcorrect, $Dcorrect,$Ucorrect ) = @{$correct};    
      #Check to see if the student L matrix is unit lower triangular
      unless (($Lstu->element(1,2)==0)&&($Lstu->element(1,1)==1)&&($Lstu->element(2,2)==1)){
          $self->setMessage(1,'Your L matrix is not unit lower triangular');
          $self->setMessage(2,'The D matrix cannot be checked for accuracy '.
                              'since L does not have the correct format.');
          $self->setMessage(3,'The U matrix cannot be checked for accuracy '.
                              'since L does not have the correct format.');
          return[0,0,0];
          };
      #Check to see if the student U matrix is unit upper triangular
      unless(($Ustu->element(2,1)==0)&&($Ustu->element(1,1)==1)&&($Ustu->element(2,2)==1)){
          $self->setMessage(1,'The L matrix cannot be checked for accuracy '.
                              'since U does not have the correct format.');
          $self->setMessage(2,'The D matrix cannot be checked for accuracy '.
                              'since U does not have the correct format.');
          $self->setMessage(3,'Your U matrix is not unit upper triangular.');
          return[0,0,0];
          };
      #Check to see if the student D matrix is diagonal
      unless (($Dstu->element(2,1)==0)&&($Dstu->element(1,2)==0)){
           
           $self->setMessage(1,'The L matrix cannot be checked for accuracy '.
                              'since D does not have the correct format.');
           $self->setMessage(2,'Your D matrix should be diagonal.');
           $self->setMessage(3,'The U matrix cannot be checked for accuracy '.
                              'since D does not have the correct format.');
          return[0,0,0];
          };
      #Check to see if they have the right product
      unless($Lstu*$Dstu*$Ustu == $Lcorrect*$Dcorrect*$Ucorrect){
          $self->setMessage(1,'Your L matrix is in the right format, '.
                              'but you have the wrong product.');
          $self->setMessage(2,'Your D matrix is in the right format, '.
                              'but you have the wrong product.');
          $self->setMessage(3,'Your U matrix is in the right format, '.
                              'but you have the wrong product.');
          return[0,0,0]
          }
      return [1,1,1];
  }
);


##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
\{ textbook_ref_exact("Holt Linear Algebra", "3.4","28") \}
$BR
For this problem enter matrices using bracket notation. So the matrix 
$BR \($sample\)
$BR
would be entered as [[2,3,-1],[-2,-5,6]].
$BR
$BR
Find an LDU factorization for 
$BR
\($A = $L $U2\)
$BR
\(L = \)\{$multians->ans_rule(40)\}
$BR
\(D = \)\{$multians->ans_rule(40)\}
$BR
 \(U = \)\{$multians->ans_rule(40)\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS( $multians->cmp() );

Context()->texStrings;
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD} 
We are given that \(L = $L\). Next we divide the rows of \(U = $U2\) by the diagonal entries to obtain \(D = $D\) and \(U = $U\).
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
