##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(LU factorization)
## Date(6/3/2013)
## Author(Darwyn Cook)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.4')
## Problem1('55')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Matrix");

$L = Matrix([[1,0,0,0],[0,1,0,0],[0,0,1,0],[0,0,0,1]]);
$U = Matrix([[1,0,0],[0,1,0],[0,0,1],[0,0,0]]);
##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
\{ textbook_ref_exact("Holt Linear Algebra", "3.4","55") \}
$BR
Find a matrix \(A\) that has an LU factorization where \(L\) is \(4 \times 4\) and
\(U\) is \(4 \times 3\).
$BR
$BR
\(L = \)\{$L->ans_array\}
$BR
$BR
\(U = \)\{$U->ans_array\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($L->cmp(checker=>sub {
      my ( $correct, $Lstu, $ansHash ) = @_;
      #Check to see if the student L matrix is lower triangular
      Value->Error("Your L matrix is not lower triangular") unless 
             (($Lstu->element(1,2)==0)&&($Lstu->element(1,3)==0)&&
              ($Lstu->element(1,4)==0)&&($Lstu->element(2,3)==0)&&
              ($Lstu->element(2,4)==0)&&($Lstu->element(3,4)==0));
      #Check to see if the student L has 1's on the diagonal
      Value->Error("Your L matrix should have ones on the diagonal.") unless    
             (($Lstu->element(1,1)==1)&&($Lstu->element(2,2)==1)&&
             ($Lstu->element(3,3)==1)&&($Lstu->element(4,4)==1));
      return 1;
} ) );
ANS($U->cmp(checker=>sub {
      my ( $correct, $Ustu, $ansHash ) = @_;
      #Check to see if the student L matrix is lower triangular
      Value->Error("Your U matrix is not upper triangular") unless 
             (($Ustu->element(2,1)==0)&&($Ustu->element(3,1)==0)&&
              ($Ustu->element(3,2)==0)&&($Ustu->element(4,1)==0)&&
              ($Ustu->element(4,2)==0));
      return 1;
} ) );

Context()->texStrings;
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD} 
Such an example would be
\[
A = \left[\begin{array}{rrr} 1 & 0 & 0 \\ 0 & 1 & 0 \\ 0 & 0 & 1 \\ 0 & 0 & 0 \end{array}\right]
\]
which has LU factorization
\[
A = LU = \left[\begin{array}{rrrr} 1 & 0 & 0 & 0 \\ 0 & 1 & 0 & 0 \\ 0 & 0 & 1 & 0 \\ 0 & 0 & 0 & 1 \end{array}\right] \left[\begin{array}{rrr} 1 & 0 & 0 \\ 0 & 1 & 0 \\ 0 & 0 & 1 \\ 0 & 0 & 0 \end{array}\right]
\]
That is, \(L\) is the \(4\times 4\) identity matrix and \(U = A\).
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();        
