##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(LU factorization)
## Date(6/3/2013)
## Author(Darwyn Cook)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.4')
## Problem1('58')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Matrix");

$A = Matrix([[1,0],[0,1],[0,0]]);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
\{ textbook_ref_exact("Holt Linear Algebra", "3.4","58") \}
$BR
Find a \(3 \times 2 \) matrix \(A\) and has LU factorization where \(L\) is diagonal.
$BR
$BR
\(A = \)\{$A->ans_array\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#


ANS($A->cmp(checker=>sub {
      my ( $correct, $Astu, $ansHash ) = @_;
      #Check to see if the student L matrix is lower triangular
      if (($Astu->element(2,1)==0)&&($Astu->element(3,1)==0)&&
              ($Astu->element(3,2)==0)){
         return 1;
      }
      else {return 0};
} ) );

Context()->texStrings;
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD} 
An example is
\[
A = \left[\begin{array}{rr} 1 & 0 \\ 0 & 1 \\ 0 & 0 \end{array}\right]
\]
which has LU factorization
\[
A = LU = \left[\begin{array}{rrr} 1 & 0 & 0 \\ 0 & 1 & 0 \\ 0 & 0 & 1 \end{array}\right] \left[\begin{array}{rr} 1 & 0 \\ 0 & 1 \\ 0 & 0 \end{array}\right]
\]
That is, \(L\) is the \(3 \times 3 \) identity matrix and \(U = A\).
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();        
