##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Markov chains)
## Date(6/3/2013)
## Author(Darwyn Cook)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.5')
## Problem1('33')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Matrix");

Context("Numeric");

##
## If the student has the right matrix, the eigenvalues will all have 
## |lambda| = 1
##
## The student matrix is [[a,c,e][b,d,f],[1-(a+b),1-(c+d),1-(e+f)]]

$multians = MultiAnswer(1,0,0,1,1,1,0,0)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $a11,$a12,$a21,$a22,$b11,$b12,$b21,$b22 ) = @{$student};
      my ( $ansa11, $ansa12,$ansa21,$ansa22,$ansb11, $ansb12,$ansb21,$ansb22 ) 
          = @{$correct};
      #Check to see if the student thinks the answer is true
       if (($a11==1)&&($a12==1)&&($a21==1)&&($a22==1)&&($a11==1)&&
               ($a12==1)&&($a21==1)&&($a22==1)){
          $self->setMessage(1,'The answer is false, you need a counterexample.');
          return[0,0,0,0,0,0,0,0];
          };
       #Check that the matrices are stochastic
       unless (($a11+$a21==1)){
          $self->setMessage(1,'Your matrix A is not stochastic.');
          return[0,0,0,0,0,0,0,0];
          };
       unless (($a12+$a22==1)){
          $self->setMessage(3,'Your matrix A is not stochastic.');
          return[0,0,0,0,0,0,0,0];
          };
       unless (($b11+$b21==1)){
          $self->setMessage(5,'Your matrix A is not stochastic.');
          return[0,0,0,0,0,0,0,0];
          };
       unless (($b12+$b22==1)){
          $self->setMessage(8,'Your matrix A is not stochastic.');
          return[0,0,0,0,0,0,0,0];
          };
      #Now we check to see if they have a counterexample.
      $m11 = $a11*$b11+$a12*$b12;
      $m12 = $a11*$b21+$a12*$b22;
      $m21 = $a21*$b11+$a22*$b12;
      $m22 = $a21*$b21+$a22*$b22;
      if (($m11+$m21 == 1)&&($m12+$m22 == 1)){
          $self->setMessage(1,'The product AB^T is stochastic, you need one that is not stochastic');
          return[0,0,0,0,0,0,0,0];
          }  
      else {
          return[1,1,1,1,1,1,1,1];
          };
  }
);


##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
\{ textbook_ref_exact("Holt Linear Algebra", "3.5","33") \}
$BR
True or False: If \(A\) and \(B\) are stochastic n × n matrices, then \(AB^{T}\) is also stochastic. If the answer is true put 1's in all of the answer blanks, if the answer is false provide a counterexample.
$BR

$BR
$BR
\{mbox("\(A = \)", display_matrix([ [$multians->ans_rule(5),$multians->ans_rule(5)] ,[$multians->ans_rule(5),$multians->ans_rule(5)] ]),"\(B = \)",display_matrix([ [$multians->ans_rule(5),$multians->ans_rule(5)] ,[$multians->ans_rule(5),$multians->ans_rule(5)] ]))\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#


ANS( $multians->cmp() );

Context()->texStrings;
$A = Matrix([1,0],[0,1]);  $B = Matrix([1,1],[0,0]);
$ABt = $A*($B->transpose);
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD} 
False, since if \(A = $A\) and \(B = $B\) then \(A\) and \(B\) are stochastic, but \(AB^T = $ABt\) is not.
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();        
     
