##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Markov chains)
## Date(6/3/2013)
## Author(Darwyn Cook)
## Level(2)
## MO(1)
## Static(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.5')
## Problem1('36')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserRadioButtons.pl",
  "parserMultiAnswer.pl",
  "PGmatrixmacros.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Matrix");
Context()->variables->are(a=>"Real",b=>"Real");

$tf = RadioButtons(['True','False'],'True', separator => $SPACE x 4 );

$A = Matrix([[Formula("a"),Formula("1-b")],[Formula("1-a"),Formula("b")]]);

$multians = MultiAnswer(Formula("(1-b)/(2-a-b)"),Formula("(1-a)/(2-a-b)"))->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1,$stu2 ) = @{$student};
      my ( $correct1, $correct2 ) = @{$correct};
      if ($stu1 + $stu2 != 1){
         Value::Error('Your vector is not a state vector.');
      }
      if ( (Formula("a*$stu1+(1-b)*$stu2") != $stu1) or 
           (Formula("(1-a)*$stu1+b*$stu2") != $stu2) )
         {Value->Error("Your vector is not a steady-state vector for \(A\)");}
      else {return 1};
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
Every 2x2 stochastic matrix \(A\) has the form 
\[A = $A\] $BR
where \(0\leq a\leq 1\) and \(0\leq b \leq 1\).  
$PAR
1) True or False? Every 2×2 stochastic matrix has at least one steady-state vector.$BR
$SPACE$SPACE$SPACE$SPACE \{$tf->buttons\}$PAR

2) Assume \(A\neq \begin{bmatrix}1 & 0 \\ 0 & 1\end{bmatrix} \).  If \(A\) has a steady-state vector enter it here.  If it has no steady-state matrix enter the zero vector.
$BR
\{display_matrix([[$multians->ans_rule],[$multians->ans_rule]])\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#


ANS($tf->cmp);
ANS($multians->cmp());

ENDDOCUMENT();        
