##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Markov chains)
## Date(6/3/2013)
## Author(Darwyn Cook)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.5')
## Problem1('36')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->variables->are(a=>"Real",b=>"Real");

$multians = MultiAnswer(Formula("(1-b)/(2-a-b)"),Formula("(1-a)/(2-a-b)"))->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1,$stu2 ) = @{$student};
      my ( $correct1, $correct2 ) = @{$correct};
      if (($stu1 == 1) && ($stu2 == 1)){
         $self->setMessage(1,'The answer is true, you need to find a vector.');
         return [0,0];
      }
      elsif ($stu1 + $stu2 != 1){
         $self->setMessage(1,'Your vector is not a state vector.');
         return [0,0];
      }
      elsif (Formula("a*$stu1+b*$stu2") != $stu1){
         return [0,0];
      }
      elsif (Formula("(1-a)*$stu1+(1-b)*$stu2") != $stu2){
         return [1,0];
      }
      else {return [1,1]};
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
\{ textbook_ref_exact("Holt Linear Algebra", "3.5","36") \}
$BR
True or False: Every 2 × 2 stochastic matrix has at least one steady-state
vector.
$BR
 \{mbox("Note: Every 2 x 2 stochastic matrix has the form A = ", display_matrix([["\({\rm a}\)","\({\rm 1-b}\)"],["\({\rm 1-a}\)","\({\rm b}\)"]]))\} where \(0 \leq {\rm a} \leq 1\) and \(0 \leq {\rm b} \leq 1\). If there is a steady-state vector for this matrix enter it below, otherwise enter all 1's for the vector.
\{display_matrix([[$multians->ans_rule],[$multians->ans_rule]])\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($multians->cmp());

ENDDOCUMENT();        
