##DESCRIPTION
##
##ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Markov chains)
## Date(6/3/2013)
## Author(Darwyn Cook)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Holt')
## EditionText1('1')
## Section1('3.5')
## Problem1('30')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

##
## If the student has the right matrix, the eigenvalues will all have 
## |lambda| = 1
##
## The student matrix is [[a,c,e][b,d,f],[1-(a+b),1-(c+d),1-(e+f)]]

$multians = MultiAnswer(0,1,0,0,0,1,1,0,0)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $a,$c,$e,$b,$d,$f,$x,$y,$z ) = @{$student};
      my ( $ans11, $ans12,$ans21,$ans22 ) = @{$correct};
      #Check that the student answer is stochastic
       unless (($a>=0)&&($b>=0)&&($c>=0)&&($d>=0)&&($e>=0)&&
              ($f>=0)&&($x>=0)&&($y>=0)&&($z>=0)){
          $self->setMessage(1,'Your matrix is not stochastic.');
          return[0,0,0,0,0,0,0,0,0];
          };
       unless (($a+$b+$x==1)){
          $self->setMessage(1,'Your matrix is not stochastic.');
          return[0,0,0,0,0,0,0,0,0];
          };
       unless (($c+$d+$y==1)){
          $self->setMessage(4,'Your matrix is not stochastic.');
          return[0,0,0,0,0,0,0,0,0];
          };
       unless (($e+$f+$z==1)){
          $self->setMessage(7,'Your matrix is not stochastic.');
          return[0,0,0,0,0,0,0,0,0];
          };
      #Now that we know we have a stochastic matrix, we look at the eigenvalues,
      #other than the eigenvalue 1 we know must exist.
      $discriminant = 
      Real("($a+$d-$e-$f)**2+4*($b*$c-$a*$d-$b*$e+$d*$e+$a*$f-$c*$f)");
      if ($discriminant < 0){
         $magnitudelambda = 
         Compute("(1/2($a+$d-$e-$f))**2+$discriminant/4");
        if ($magnitudelambda == 1){
          $self->setMessage(1,'Your matrix has more than one solution');
          return[0,0,0,0,0,0,0,0,0];
          }
        else{
          return[1,1,1,1,1,1,1,1,1];
          }
        }
      else {
         $lambda1 = Compute("1/2($a+$d-$e-$f-sqrt($discriminant))");
         $lambda2 = Compute("1/2($a+$d-$e-$f+sqrt($discriminant))");
         if (($lambda1 == 1)||($lambda2 == 1)){
          $self->setMessage(1,'Your matrix has more than one solution');
          return[0,0,0,0,0,0,0,0,0];
          }
         else {
          return[1,1,1,1,1,1,1,1,1];
          }
       };
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
\{ textbook_ref_exact("Holt Linear Algebra", "3.5","30") \}
$BR
Find a \(3 \times 3 \) stochastic matrix \(A\) that has exactly one initial state
vector \(x_0\) that will generate a Markov chain with a steady-state
vector.
$BR
$BR
\(A = \) \{display_matrix([ [$multians->ans_rule(5),$multians->ans_rule(5),$multians->ans_rule(5)] ,[$multians->ans_rule(5),$multians->ans_rule(5),$multians->ans_rule(5)],[$multians->ans_rule(5),$multians->ans_rule(5),$multians->ans_rule(5)] ])\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#


ANS( $multians->cmp() );

Context()->texStrings;
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD} 
The matrix
\[
A = \left[\begin{array}{rrr}1 & 0 & 0 \\ 0 & 0 & 1 \\ 0 & 1 & 0 \end{array}\right]
\]
has only 
\[
\mathbf{x}_0 = \left[\begin{array}{r} 1/3 \\ 1/3 \\ 1/3 \end{array}\right]
\]
that generates a Markov chain with steady-state vector. For all other vectors 
\[
\mathbf{x}_0 = \left[\begin{array}{r} a_1 \\ a_2 \\ a_3 \end{array}\right]
\]
we obtain the cycle
\[
\mathbf{x}_1 = \left[\begin{array}{r} a_2 \\ a_3 \\ a_1 \end{array}\right],\, \mathbf{x}_2 = \left[\begin{array}{r} a_3 \\ a_1 \\ a_2 \end{array}\right],\,
\mathbf{x}_3 = \left[\begin{array}{r} a_1 \\ a_2 \\ a_3 \end{array}\right] = \mathbf{x}_0.
\]
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();        
