# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 06/2/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Institution(Howard University)
## Author(Richard Bayne)
## MO(1)
## MLT(basis_by_deleting_vectors)
## MLTleader(1)
## Level(3)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('')
## TitleText2('')
## AuthorText2('')
## EditionText2('')
## Section2('.')
## Problem2('21')
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","21") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector");

$n11 = 0;
$n12 = random(2,5,1);
$n13 = random(3,7,2);
$n14 = 1;
$n21 = 0;
$n22 = 0;
$n23 = random(-1,-3,1);
$n24 = $n13;
$n31 = 0;
$n32 = 0;
$n33 = 0;
$n34 = $n13;

$N1 = ColumnVector([$n11,$n21,$n31]);
$N2 = ColumnVector([$n12,$n22,$n32]);
$N3 = ColumnVector([$n13,$n23,$n33]);
$N4 = ColumnVector([$n14,$n24,$n34]);

$multians = MultiAnswer($N2,$N3,$N4) -> with(
  singleResult => 1,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($corvec1, $corvec2, $corvec3) = @{$correct};
    my ($stuvec1, $stuvec2, $stuvec3) = @{$student};
    if ( List($corvec1,$corvec2,$corvec3) == List($stuvec1,$stuvec2,$stuvec3) ||
         List($corvec1,$corvec2,$corvec3) == List($stuvec1,$stuvec3,$stuvec2) ||
         List($corvec1,$corvec2,$corvec3) == List($stuvec2,$stuvec1,$stuvec3) ||
         List($corvec1,$corvec2,$corvec3) == List($stuvec2,$stuvec3,$stuvec1) ||
         List($corvec1,$corvec2,$corvec3) == List($stuvec3,$stuvec1,$stuvec2) ||
         List($corvec1,$corvec2,$corvec3) == List($stuvec3,$stuvec2,$stuvec1)  
       ){
      return 1;
    }
    else {
      # extract the coordinates from student vectors
      my ( $s11,$s21, $s31 ) = ($stuvec1.i, $stuvec1.j, $stuvec1.k);
      my ( $s12,$s22, $s32 ) = ($stuvec2.i, $stuvec2.j, $stuvec2.k);
      my ( $s13,$s23, $s33 ) = ($stuvec3.i, $stuvec3.j, $stuvec3.k);
      my $stumatrix = Matrix([
        [ $s11,$s21, $s31 ],
        [ $s12,$s22, $s32 ],
        [ $s13,$s23, $s33 ]
      ]);
      # check that student's vectors are linearly independent
      if ($stumatrix->det == 0){
        Value::Error("Basis vectors must be linearly independent");
      }
      else {
        Value::Error("The basis vectors must be taken from the given set.");
      }
      return 0;
    }
  }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Find a basis of the given subspace by deleting linearly dependent vectors.
$BR
span of \(\Bigg $LBRACE $N1, $N2, $N3, $N4 \Bigg $RBRACE\)
$PAR
A basis is \(\Bigg $LBRACE\)\{$multians->ans_array\}, \{$multians->ans_array\}, \{$multians->ans_array\}\(\Bigg $RBRACE\)
END_TEXT
Context()->normalStrings;

ANS($multians->cmp());

Context()->texStrings;
BEGIN_SOLUTION
$PAR$PAR{BBOLD}SOLUTION:${EBOLD} $PAR
The first vector is eliminated since the zero vector is always linearly dependent. The remaining vectors are linearly dependent since the correspondng matrix
$PAR
 \[ \left[ \begin{array}{ccc}  $n12 & $n13 & $n14 \\ $n22 & $n23 & $n24 \\ $n32 & $n33 & $n34 \end{array} \right] \] is row-reduced with 3 pivots. Hence, a basis is
$PAR
 \[\left$LBRACE \left [ \begin{array}{c} $n12 \\ $n22 \\ $n32 \end{array} \right], \left [ \begin{array}{c} $n13 \\ $n23 \\ $n33 \end{array} \right], \left [ \begin{array}{c} $n14 \\ $n24 \\ $n34 \end{array} \right] \right $RBRACE \]
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();

