# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 05/29/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Institution(Howard University)
## Author(Richard Bayne)
## MLT(find_basis_with_instructions)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('8')
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","8") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixReduce.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Matrix");

$m = random(-3,-2,1);
$k = random(2,4,1);
$n = random(-2,-1,1);
$a11 =random(2,5,1);
$a12 = -1;
$a13 = random(1,3,1);
$b22 = 1;
$b23 = random(-5,-3,1);


# row-reduced matrix 
$R=Matrix([
  [$a11,$a12,$a13],
  [0   ,$b22,$b23],
  [0   ,0   ,   0]
]);
$E1=Matrix([
  [  1, 0, 0],
  [-$m, 1, 0],
  [  0, 0, 1]
]);
$E2=Matrix([
  [1       ,0 ,0],
  [0       ,1, 0],
  [$k*$m-$n,$k,1]
]);
# un-reduced matrix
$A = $E2*$E1*$R;

#spanning set
$V1 = ColumnVector($A->row(1));
$V2 = ColumnVector($A->row(2));
$V3 = ColumnVector($A->row(3));

#expected answer in solution
$B1 = ColumnVector($R->row(1));
$B2 = ColumnVector($R->row(2));

# problem asked only for "echelon form" which isn't unique so accept all bases

$multians = MultiAnswer($B1,$B2)->with(
  singleResult=>1,
  checker => sub {
    my ($correct,$student,$self)=@_;
    my ($cor1,$cor2)=@{$correct};
    my ($stu1,$stu2)=@{$student};
    # convert to row vectors
    $cor1=Vector($cor1);
    $cor2=Vector($cor2);
    $stu1=Vector($stu1);
    $stu2=Vector($stu2);
    my $cor = Matrix([
     [@{$cor1->{data}}],
     [@{$cor2->{data}}]
    ]);
    my $stu = Matrix([
     [@{$stu1->{data}}],
     [@{$stu2->{data}}]
    ]);
    #reduce
    my $diff = rref($cor)-rref($stu);
    if ( $diff->isZero ){
      return 1;
    }
    else {
      if ($stu->row(2)->isZero){
        Value::Error("Basis vectors must be linearly independent.");
      }
      else {
        Value::Error("Basis vectors must span the given subspace.");
      }
      return 0;
    }
  }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Suppose a subspace is spanned by the set of vectors shown. Find a basis for the subspace, using the method of transforming a matrix to echelon form, where the rows of the matrix represent vectors spanning the subspace. $BR
\(\Bigg $LBRACE\)\($V1\), \($V2\), \($V3\)\(\Bigg $RBRACE\)

$PAR
A basis = \(\Bigg $LBRACE\)\{$multians->ans_array\}, \{$multians->ans_array\} \(\Bigg $RBRACE\)


$PAR
What is the dimension of the subspace? \{ans_rule(5) \}
END_TEXT
Context()->texStrings;

ANS($multians->cmp());
ANS(Real(2)->cmp());

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR$PAR{BBOLD}SOLUTION:${EBOLD} $PAR
Row-reduce the matrix which has the given vectors as rows.
$BR$BR
\($A\) \(\sim\) \(\left[ \begin{array}{ccc} $a11 & $a12 & $a13 \\ 0 & $b22 & $b23\\ 0 & 0 & 0 \end{array} \right]  \)$PAR
so that a basis for the subspace is given by the non-zero row vectors
$PAR
 \(\left$LBRACE \left [ \begin{array}{c} $a11 & $a12 & $a13 \end{array} \right] , \left [ \begin{array}{c} 0 & $b22 & $b23 \end{array} \right]\right$RBRACE \)

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
