## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Level(4)
## Static(1)
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Date('')
## Author('Richard Bayne')
## Institution('Howard University')
## TitleText1('Linear Algebra with Applications')
## EditionText1('1')
## AuthorText1('Jeffrey Holt')
## Section('4.2')
## Problem1('25')


## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","25") \}

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PG.pl",
"PGbasicmacros.pl",
"PGanswermacros.pl",
"PGmatrixmacros.pl", 
"PGmorematrixmacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"PGauxiliaryFunctions.pl"

);

TEXT(beginproblem());

$a11 = random(-5,-1,1);;
$a21 = random(2,4,1);
$a31 = random(1,3,2);
Context('Matrix');

$ma= MultiAnswer(Compute('[[1],[0],[0]]'), Compute("[[0],[1],[0]]") )->with(
  singleResult=>1,
  checker => sub {
	  my ($correct,$student,$self) = @_;
	  my ($b1,$b2) = @{$student};
	  my $m = Matrix([[$b1->element(1,1), $b1->element(2,1), $b1->element(3,1)],
	  [$b2->element(1,1), $b2->element(2,1), $b2->element(3,1)], 
	  [$a11, $a21, $a31]]);
	  return 1 if $m->det != Real(0);
	  return 0;
	  },
);

BEGIN_TEXT
Expand the set below to form a basis for \( \mathbb{R}^3\) by adding vectors to the set.

\[\left\lbrace \begin{pmatrix} $a11 \\ $a21 \\ $a31 \end{pmatrix}
\right\rbrace \]

The vectors to add are  \{$ma->ans_rule(20) \} and  \{$ma->ans_rule(20) \} 
$BR$BR
$BBOLD Note:$EBOLD
you can represent a column vector using square brackets. 
 For instance, the vector \(\begin{pmatrix} 
 $a11 \\ $a21 \\ $a31 
 \end{pmatrix} \) can be written as [ [$a11], [$a21], [$a31] ].

END_TEXT

#  **The following needs to be checked.**
# ANS(basis_cmp( [[1,0,0]]) );
# ANS(basis_cmp( [[0,1,0]]) );
ANS($ma->cmp());

ENDDOCUMENT()
