# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 05/27/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Institution(Howard University)
## Author(Richard Bayne)
## MLT(find_basis_with_instructions)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('11')
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","11") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector");

$a11 =random(1,5,2);
$a21 = random(4,7,3);
$mult = random(2,4,1);
$a12 = $mult * $a11;
$a22 = -$mult * $a21;

$C1 = ColumnVector($a11,$a21);
$C2 = ColumnVector($a12,$a22);

$b22 = $a22 - $a12*$a21;

$multians = MultiAnswer($C1,$C2)->with(
  singleResult => 1,
  checker => sub {
    my ($correct,$student,$self) = @_;
    my ($cor1, $cor2) = @{$correct};
    my ($stu1, $stu2) = @{$student};
    # convert to rows
    $cor1 = Vector($cor1);
    $cor2 = Vector($cor2);
    $stu1 = Vector($stu1);
    $stu2 = Vector($stu2);
    my $M = Matrix([
     [@{$stu1->{data}}],
     [@{$stu2->{data}}]
    ]);
    if ($M->det->isZero){
      Value::Error("Basis vectors must be linearly independent.");
      return 0;
    }
    elsif (($cor1==$stu1 && $cor2==$stu2) || ($cor1==$stu2 && $cor2==$stu1)){
      return 1;
    }
    else {
      Value::Error("Must be a subset of the given set.");
      return 0;
    }
  }
);


TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Suppose a subspace is spanned by the set \(S\) of vectors shown below.
$PAR 
\(S = \Bigg $LBRACE\)\($C1\), \($C2\)\(\Bigg $RBRACE\)
$PAR
Find a subset of \(S\) that forms a basis for the subspace, using the method of transforming a matrix to echelon form, where the columns of the matrix represent vectors spanning the subspace.

$PAR

A basis for the subspace is \(\Bigg $LBRACE\)\{ $multians->ans_array\}, \{ $multians->ans_array\} \(\Bigg $RBRACE\)

$PAR
What is the dimension of the subspace? \{ans_rule(5) \}
END_TEXT
Context()->normalStrings;

ANS($multians->cmp());
ANS(Real(2)->cmp());

Context()->texStrings;
BEGIN_SOLUTION
$SOL $PAR
Row-reduce the matrix which has the given vectors as columns:
$PAR
\(\left[ \begin{array}{cc} $a11 & $a12 \\ 
$a21 & $a22 \end{array}\right] \sim \left[ \begin{array}{cc} $a11 & $a12 \\ 
0 & $b22 \end{array} \right]  \)$PAR
So a basis for the subspace is given by columns 1 and 2 of the original matrix corresponding to the pivot columns of the row-reduced matrix.
In other words, a basis for the subspace is
$PAR
\(\Bigg $LBRACE\)\($C1\), \($C2\)\(\Bigg $RBRACE\).
$PAR
The dimension of the subspace is therefore 2.
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();

