# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 06/2/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Institution(Howard University)
## Author(Richard Bayne)
## MLT(basis_by_deleting_vectors)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('')
## TitleText2('')
## AuthorText2('')
## EditionText2('')
## Section2('.')
## Problem2('20')
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","20") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixReduce.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Matrix");

$h11 = random(1,3,2);
$h21 = -$h11;
$h31 = $h11;
$mult5 = random(-5,-3,1);
$j12 = $mult5 * $h11;
$j22 = $mult5 *$h21;
$j32 = $mult5 * $h31;
$k13 = 2*$h11;
$k23 = $h21 +8;
$k33 = $h31 +7;

$V1 = ColumnVector($h11, $h21, $h31);
$V2 = ColumnVector($j12, $j22, $j32);
$V3 = ColumnVector($k13, $k23, $k33);

$multians = MultiAnswer($V1,$V3)->with(
  singleResult => 1,
  checker => sub { 
    my ($correct,$student,$self) = @_;
    my ($cor1,$cor2)=@{$correct};
    my ($stu1,$stu2)=@{$student};
    $cor1alt = $V2;
    if ( ($cor1 == $stu1 && $cor2 == $stu2) || 
         ($cor1 == $stu2 && $cor1 == $stu2) ||
         ($cor1alt == $stu1 && $cor2 == $stu2) ||
         ($cor1alt == $stu2 && $cor2 == $stu1) ){
         return 1;
    } 
    else {
      $stu1 = Vector($stu1);
      $stu2 = Vector($stu2);
      $stu = Matrix([
	[@{$stu1->{data}}],
	[@{$stu2->{data}}]
      ]);
      $stu = rref($stu);
      if ($stu->row(2)->isZero){
	Value::Error("The vectors in a basis must be linearly independent.");
      }
      else {
        Value::Error("Answer should be a subset of S.");
      }
      return 0;
    }
  }
);


TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Find a basis of the subspace \(V = \text{ span of } S\) by deleting linearly dependent vectors from the set.  Your answer should be a subset of \(S\).
$PAR
S = span of \(\Bigg $LBRACE\)\($V1\), \($V2\), \($V3\)\(\Bigg $RBRACE\)

$PAR
A basis is \(\Bigg $LBRACE\)\{$multians->ans_array\}, \{$multians->ans_array\}\(\Bigg $RBRACE\)

END_TEXT
Context()->normalStrings;

ANS($multians->cmp());

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR$PAR{BBOLD}SOLUTION:${EBOLD} $PAR
The second vector is $mult5 times the first, and is eliminated as a dependent vector. Since the remaining vectors are linearly dependent, a basis is
$PAR
 \(\left$LBRACE \left [ \begin{array}{c} $h11 \\ $h21 \\ $h31 \end{array} \right], \left [ \begin{array}{c} $k13 \\ $k23 \\ $k33 \end{array} \right] \right$RBRACE \)

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();

