## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Level(4)
## MO(1)
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Date('')
## Author('Richard Bayne')
## Institution('Howard University')
## TitleText1('Linear Algebra with Applications')
## EditionText1('1')
## AuthorText1('Jeffrey Holt')
## Section('4.2')
## Problem1('25')


## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","25") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "PGmorematrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$a11 = random(-5,-1,1);;
$a21 = random(2,4,1);
$a31 = random(1,3,2);
Context('Matrix');

$ma= MultiAnswer(Vector(1,0,0), Vector(0,1,0) )->with(
  singleResult=>1,
  checker => sub {
	  my ($correct,$student,$self) = @_;
	  my ($b1,$b2) = @{$student};
	  @bv1=$b1->value;
	  @bv2=$b2->value;
	  my $m = Matrix([[$bv1[0], $bv1[1], $bv1[2]],
	  [$bv2[0], $bv2[1], $bv2[2]], 
	  [$a11, $a21, $a31]]);
	  return 1 if $m->det != Real(0);
	  return 0;
	  },
);

BEGIN_TEXT
Expand the set below to form a basis for \( \mathbb{R}^3\) by adding vectors to the set.

\[\left\lbrace \begin{pmatrix} $a11 \\ $a21 \\ $a31 \end{pmatrix}
\right\rbrace \]

The vectors to add are  \{$ma->ans_rule(20) \} and  \{$ma->ans_rule(20) \} 
$BR$BR
$BBOLD Note:$EBOLD
enter your answers as \{helpLink('vectors')\}.
END_TEXT

ANS($ma->cmp());

ENDDOCUMENT()
