# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 06/2/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Institution(Howard University)
## Author(Richard Bayne)
## MLT(basis_by_deleting_vectors)
## Level(3)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('')
## TitleText2('')
## AuthorText2('')
## EditionText2('')
## Section2('.')
## Problem2('18')
## KEYWORDS('linear algebra', 'subspaces', 'basis')

## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","18") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector");

$c11 =random(9,13,1);
$c21 =-($c11 - 4);
$mult2 = random(-3,-2,1);
$d11 = $c11*$mult2-1;
$d21 = $c21*$mult2;

$c = ColumnVector($c11,$c21);
$d = ColumnVector($d11,$d21);

$multians = MultiAnswer($c,$d)->with(
  singleAnswer => 1,
  checker => sub {
    my ($correct,$student,$self)=@_;
    my ($stu1,$stu2) = @{$student};
    # make them row vectors
    $stu1=Vector($stu1);
    $stu2=Vector($stu2);
    my $mat = Matrix([
      [@{$stu1->{data}}],
      [@{$stu2->{data}}]
    ]);
    if ( $mat->det->isZero ) {
      Value::Error("Basis vectors must be linearly independent and span the sub space.");
      return 0;
    }
    else {
      return 1;
    }
  }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Find a basis of the given subspace by deleting linearly dependent vectors.  
$PAR
\(S\) = span of \(\Bigg $LBRACE $c\), \($d \Bigg $RBRACE \).  
$PAR
A basis is \(\Bigg $LBRACE\)\{$multians->ans_array \}, \{$multians->ans_array \} \(\Bigg $RBRACE\).
END_TEXT
Context()->normalStrings;

ANS($multians->cmp());

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR$PAR{BBOLD}SOLUTION:${EBOLD} 
$PAR
Since the vectors are not multiples of each other, they are linearly independent, so a basis is
$PAR
\( \left$LBRACE \left[ \begin{array}{c} $c11 \\ $c21 \end{array}\right], \left[ \begin{array}{c} $d11 \\ $d21 \end{array}\right] \right$RBRACE \)
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
