# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 06/2/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Basis and dimension)
## Institution(Howard University)
## Author(Richard Bayne)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('23')
## KEYWORDS('linear algebra', 'subspaces', 'basis')


## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.2","23") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context("Vector");

$a11 = random(1,3,1);
$a21 = random(-5,-2,1);

$v=Vector($a11, $a21);
$v_col = ColumnVector($v);
$ans=ColumnVector(1,0); 

Context()->texStrings;
BEGIN_TEXT
Expand the set below to form a basis for \(\mathbb{R}^2\) by adding a vector to the set. $BR

\(\Bigg $LBRACE $v_col \Bigg $RBRACE\) 

$BR
Adding the vector \{ans_rule(10) \} to the set will make it a basis.
$PAR
Click \{ AnswerFormatHelp("vectors") \} for help with entering vectors.
END_TEXT

Context()->normalStrings;

ANS( $ans->cmp(
 showCoordinateHints => 0,
 checker => sub{
   my ($corr, $stu ) = @_;
   if ($stu->isZero or ($stu->length) != 2 ){
     return 0;
   }
   if ( $v->isParallel($stu) ){
     
   }
   else { 
     return 1;
   }
 }
));

Context()->texStrings;
BEGIN_SOLUTION
Any vector in \(\mathbb{R}^2\) that is linearly independent of \($v\) will do.  For example \($ans\) will work.
END_SOLUTION

ENDDOCUMENT();
