# DESCRIPTION
# Problem from Holt Linear Algebra
# WeBWorK problem written by Richard Bayne, <rbayne@howard.edu>
# ENDDESCRIPTION

## Tagged by reb 06/3/13

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Row, column, and null spaces)
## Institution(Howard University)
## Author(Richard Bayne)
## MLT(basis_row_col_nul_given_rref)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('.')
## Problem1('')
## KEYWORDS('linear algebra', 'subspaces', 'basis', 'nullity', 'rank')
## Problem('3')

## Textbook tags
## \{ textbook_ref_exact("Holt Linear Algebra", "4.3","3") \}

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixReduce.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Matrix");

$k = random(2,4,1);

$a11 = 1;
$a12 = 0;
$a13 = random(-5,-3,1);
$a14 = $a13 + 1;
$a21 = -$k;
$a22 = 1;
$a23 = random(11,15,2);
$a24 = $a23 - 8;
$a31 = 0;
$a32 = 1;
$a33 = $a23 + $k*$a13;
$a34 = $a24 + $k*$a14;

$b11 = $a11;
$b12 = $a12;
$b13 = $a13;
$b14 = $a14;
$b21 = 0;
$b22 = 1;
$b23 = $a33;
$b24 = $a34;
$b31 = 0;
$b32 = 0;
$b33 = 0;
$b34 = 0;

$mA = Matrix([
  [$a11,$a12,$a13,$a14],
  [$a21,$a22,$a23,$a24],
  [$a31,$a32,$a33,$a34]
]);

$mB = Matrix([
  [$b11,$b12,$b13,$b14],
  [$b21,$b22,$b23,$b24],
  [$b31,$b32,$b33,$b34]
]);

# column space basis
$C1 = ColumnVector($mA->column(1));
$C2 = ColumnVector($mA->column(2));

# row space basis
$R1 = $mB->row(1);
$R2 = $mB->row(2);

# null space basis
$N1 = ColumnVector(-$b13, -$b23, 1,0); 
$N2 = ColumnVector(-$b14, -$b24, 0,1); 

$colans = MultiAnswer($C1,$C2)->with(
  singleResult =>1,
  checker => sub {
    my ($correct,$student,$self)=@_;
    my ($cor1,$cor2)=@{$correct};
    my ($stu1,$stu2)=@{$student};
    # convert to rows
    $cor1 = Vector($cor1);
    $cor2 = Vector($cor2);
    $stu1 = Vector($stu1);
    $stu2 = Vector($stu2);
    $cor = Matrix([
      [@{$cor1->{data}}],
      [@{$cor2->{data}}]
    ]);
    $stu = Matrix([
      [@{$stu1->{data}}],
      [@{$stu2->{data}}]
    ]);
    $cor = rref($cor);
    $stu = rref($stu);
    $diff = $cor-$stu;
    if ($diff->isZero ){
      return 1;
    }
    else {
      if ( $stu->row(2)->isZero ){
        Value::Error("Basis vectors must be linearly independent.");
      }
      else { 
        Value::Error("The vectors are independent but they don't span the column space.");
      }
      return 0;
    }
  }
);

$rowans = MultiAnswer($R1,$R2)->with(
  singleResult =>1,
  checker => sub {
    my ($correct,$student,$self)=@_;
    my ($cor1,$cor2)=@{$correct};
    my ($stu1,$stu2)=@{$student};
    $cor = Matrix([
      [@{$cor1->{data}}],
      [@{$cor2->{data}}]
    ]);
    $stu = Matrix([
      [@{$stu1->{data}}],
      [@{$stu2->{data}}]
    ]);
    $cor = rref($cor);
    $stu = rref($stu);
    $diff = $cor-$stu;
    if ($diff->isZero ){
      return 1;
    }
    else {
      if ( $stu->row(2)->isZero ){
        Value::Error("Basis vectors must be linearly independent.");
      }
      else { 
        Value::Error("The vectors are independent but they don't span the row space.");
      }
      return 0;
    }
  }
);


$nullans = MultiAnswer($N1,$N2)->with(
  singleResult =>1,
  checker => sub {
    my ($correct,$student,$self)=@_;
    my ($cor1,$cor2)=@{$correct};
    my ($stu1,$stu2)=@{$student};
    # convert to rows
    $cor1 = Vector($cor1);
    $cor2 = Vector($cor2);
    $stu1 = Vector($stu1);
    $stu2 = Vector($stu2);
    $cor = Matrix([
      [@{$cor1->{data}}],
      [@{$cor2->{data}}]
    ]);
    $stu = Matrix([
      [@{$stu1->{data}}],
      [@{$stu2->{data}}]
    ]);
    $cor = rref($cor);
    $stu = rref($stu);
    $diff = $cor-$stu;
    if ($diff->isZero ){
      return 1;
    }
    else {
      if ( $stu->row(2)->isZero ){
        Value::Error("Basis vectors must be linearly independent.");
      }
      else { 
        Value::Error("The vectors are independent but they don't span the null space.");
      }
      return 0;
    }
  }
);


TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Find bases for the column space, the row space, and the null space of matrix A. You should verify that the Rank-Nullity Theorem holds. An equivalent echelon  form  of matrix A is given to make your work easier.
$PAR
A = \($mA \sim $mB\)

$PAR
A basis for the column space of \(A\) is  \(\Bigg $LBRACE \)\{$colans->ans_array\}, \{$colans->ans_array\}\(\Bigg $RBRACE\).
$PAR
A basis for the row space of  \(A\) is $BR
\(\Bigg $LBRACE\)\{$rowans->ans_array\},$BR
\{$rowans->ans_array\}\(\Bigg $RBRACE \)
$PAR
A basis for the null space of \(A\) is  \(\Bigg $LBRACE \)\{$nullans->ans_array\}, \{$nullans->ans_array\}\(\Bigg $RBRACE\).

END_TEXT
Context()->normalStrings;

ANS($colans->cmp());
ANS($rowans->cmp());
ANS($nullans->cmp());

Context()->texStrings;
BEGIN_SOLUTION
$SOL $PAR
A basis for the column space, determined from the pivot columns 1 and 2, is
$PAR
 \(\left$LBRACE \left [ \begin{array}{c} $a11 \\ $a21 \\ $a31 \end{array} \right] , \left [ \begin{array}{c} $a12 \\ $a22 \\ $a32 \end{array} \right]\right$RBRACE \) .
$PAR
A basis  for the row space is determined from the nonzero rows of the echelon form:
$PAR
\(\left$LBRACE \left [ \begin{array}{c} $b11 & $b12 & $b13 & $b14 \end{array} \right] , \left [ \begin{array}{c} $b21 & $b22 & $b23 & $b24 \end{array} \right]\right$RBRACE \).
$PAR
Solve \( A{\bf x} = {\bf 0}\), to obtain \({\bf x} = s_1\left[ \begin{array}{c} -$b13 \\ -$b23 \\ 1 \\ 0 \end{array} \right] + s_2\left[ \begin{array}{c} -$b14 \\ -$b24 \\  0 \\ 1 \end{array} \right] \), and so the basis for the nullspace is:
$PAR
\( \left $LBRACE \left[ \begin{array}{c} -$b13 \\ -$b23 \\ 1 \\ 0 \end{array} \right], \left[ \begin{array}{c} -$b14 \\ -$b24 \\ 0 \\ 1 \end{array} \right] \right $RBRACE\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
