# DESCRIPTION
# Determine why the system is not in echelon form and then solve.
# ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Systems of linear equations)
## DBsection(Systems with 3 variables)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra')
## AuthorText1('Holt')
## EditionText1('1e')
## Section1('1.1')
## Problem1('32')
## KEYWORDS('linear equations', 'system', 'systems')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "contextFraction.pl",
  "parserImplicitPlane.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Matrix");
sub dispSyst {
  my $A = shift;
  my $b = shift;
  my @d = @_;
  my $dS = "\begin{array}{";
  for (my $i = 0; $i < $d[1]; $i++) {
    $dS = $dS."rc";
  }
  $dS = $dS."r}";
  for (my $i = 0; $i < $d[0]; $i++) {
    my $jnonzero = 0;
    for (my $j = 0; $j < $d[1]; $j++) {
      $a[$i][$j] = $A->element($i+1,$j+1);
      if ($j == 0) {
        if ($a[$i][$j] == 1) {
          $dS = $dS." x_1 ";   $jnonzero = 1;
        } elsif ($a[$i][$j] == -1) {
          $dS = $dS."-x_1 ";  $jnonzero = 1;
        } elsif ($a[$i][$j] != 0) {
          $dS = $dS."$a[$i][$j] x_1";  $jnonzero = 1;
        }
      } else {
        if ($a[$i][$j] == 0) {
          $dS = $dS." & & ";
        } else {
          my $pm = ($a[$i][$j] > 0) ? "+" : "-";
          if ($jnonzero == 0 && $a[$i][$j] > 0) { $pm = "" };
          my $tmp = ($a[$i][$j] < 0) ? -$a[$i][$j] : $a[$i][$j];
          my $subsc = $j+1;
          if ($tmp == 1) {
            $dS = $dS." & $pm & x_{$subsc}";
          } else {
            $dS = $dS." & $pm & $tmp x_{$subsc}";
          }
          $jnonzero = 1;
        }
        if ($j == $d[1] - 1 && $i != $d[0] - 1) {
          $dS = $dS." & = & ".$b->element($i+1,1)." \\";
        } elsif ($j == $d[1] - 1 && $i == $d[0] - 1) {
          $dS = $dS." & = & ".$b->element($i+1,1);
        }
      }
    }
  }
  $dS = $dS."\end{array}";
  return $dS;
}

do {
  $A = Matrix([0,0,random(2,4)*random(-1,1,2)], [0,-1,non_zero_random(-5,5)], [random(2,7), random(2,7), random(2,7)]);
  for ($i = 0; $i< 3; $i++) {
    for ($j = 0; $j < 3; $j++) {
      $a[$i][$j] = $A->element($i+1,$j+1);
    }
  }
  ($x[0],$x[1],$x[2]) = (random(-19,19), non_zero_random(-19,19), random(-19,19));
  for ($i = 0; $i < 3; $i++) {
    $b[$i] = $a[$i][0]*$x[0] + $a[$i][1]*$x[1] + $a[$i][2]*$x[2];
  }
  $tmp = max(abs($b[0]), abs($b[1]), abs($b[2]));
} until ($tmp < 15 && ($b[1] == 0 || $b[2] == 0));

$dS = dispSyst($A,Matrix([$b[0]],[$b[1]],[$b[2]]),(3,3));

Context("ImplicitPlane");
Context()->variables->add(x1 => ["Real", TeX=>"x_1"], x2 => ["Real", TeX=>"x_1"], x3 => ["Real", TeX=>"x_3"], s1 => ["Real", TeX=>"s_1"], s2 => ["Real", TeX=>"s_1"], s3 => ["Real", TeX=>"s_3"]);
Context()->strings->add("none"=>{}, "NA"=>{}, A=>{}, B=>{}, C=>{}, D=>{});

$letters = List(String("B"));

for ($i = 2; $i >= 0; $i--) {
  $eq[2-$i] = ImplicitPlane("$a[$i][0] x1 + $a[$i][1] x2 + $a[$i][2] x3 = $b[$i]")->reduce;
}

$E = Matrix([0,0,1],[0,1,0],[1,0,0]);  $B = $E*$A;
$dEform = dispSyst($B,Matrix([$b[2]],[$b[1]],[$b[0]]),(3,3));

$vec = MultiAnswer($x[0],$x[1],$x[2])->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($s1,$s2,$s3) = @{$student};
    my ($c1,$c2,$c3) = @{$student};
    if ($c1 == $s1 && $c2 == $s2 && $c3 == $s3) {
      return [1,1,1];
    } 
    return [0,0,0];
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
The linear system $PAR \[$dS\] $PAR is not in echelon form. 
$PAR
Begin by choosing which of the following statements are correct. If there is more than one reason why the system is not in echelon form, type the letters as a comma separated list.
$PAR
${BBOLD}A.${EBOLD} The system is ${BITALIC}not${EITALIC} in echelon form because a variable is the leading variable of two or more equations. $BR
${BBOLD}B.${EBOLD} The system is ${BITALIC}not${EITALIC} in echelon form because the system is not organized in a descending ${LQ}stair step${RQ} pattern so that the index of the leading variables increases from the top to bottom.$BR
${BBOLD}C.${EBOLD} The system is ${BITALIC}not${EITALIC} in echelon form because not every equation has a leading variable. $PAR
Correct Letter(s): \{ans_rule(10)\}
$PAR
Now write the system in echelon form. 
$PAR
Equation 1: \{ans_rule(20)\} $BR
Equation 2: \{ans_rule(20)\} $BR
Equation 3: \{ans_rule(20)\} $PAR
Finally, solve the system. Use ${BBOLD}x1${EBOLD}, ${BBOLD}x2${EBOLD}, and ${BBOLD}x3${EBOLD} to enter the variables \(x_1\), \(x_2\), and \(x_3\). If necessary, use ${BBOLD}s1${EBOLD}, ${BBOLD}s2${EBOLD}, etc. to enter the free variables \(s_1\), \(s_2\), etc. $PAR
\((x_1, x_2, x_3) = \bigg(\) \{ans_rule(10)\}, \{ans_rule(10)\}, \{ans_rule(10)\} \(\bigg)\)
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($letters->cmp());
ANS($eq[0]->cmp());
ANS($eq[1]->cmp());
ANS($eq[2]->cmp());
ANS(Compute($x[0])->cmp());
ANS(Compute($x[1])->cmp());
ANS(Compute($x[2])->cmp());


###########################################################################
# solution
###########################################################################
Context()->texStrings;
$e1 = nicestring([$a[1][1],$a[1][2]],["x_2","($x[2])"])."= $b[1]";
$e2 = nicestring([$a[0][0],$a[0][1],$a[0][2]],["x_1","($x[1])","($x[2])"])."= $b[2]";
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD}
The system is not in echelon form because ${BBOLD}B${EBOLD} holds. Interchanging rows 1 and 3, gives: $PAR \[ $dEform \] $PAR
which is now echelon form. From this, the third equation of the echelon \(\Rightarrow x_3 = $x[2]\). Substitute into equation 2, \($e1 \Rightarrow x_2 = $x[1]\). Finally, substitute into equation 1, \($e2 \Rightarrow x_3 = $x[0]\).
END_SOLUTION
Context()->normalStrings;

COMMENT('Created in 2013.
$BR
Desc: Determine why the system is not in echelon form and then solve.');
ENDDOCUMENT();
