# DESCRIPTION
# Use a CAS to write the system in reduced echelon form and then solve.
# ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Echelon form)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(mat_consistant)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra')
## AuthorText1('Holt')
## EditionText1('1e')
## Section1('1.2')
## Problem1('65')
## KEYWORDS('linear equations', 'system', 'systems', 'matrices')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "parserPopUp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
sub dispMat {
  my $A = shift;
  my @d = $A->dimensions;
  my @a = ();
  my $dA = "\left[\begin{array}{";
  for (my $i = 0; $i < $d[1]; $i++) {
    $dA = $dA."r";
  }
  $dA = $dA."}";
  for (my $i = 0; $i < $d[0]; $i++) {
    for (my $j = 0; $j < $d[1]; $j++) {
      $a[$i][$j] = $A->element($i+1,$j+1);
      if ($j != $d[1] - 1) {
        $dA = $dA." $a[$i][$j] & ";
      } elsif ($j == $d[1] - 1 && $i != $d[0] -1) {
        $dA = $dA." $a[$i][$j] \\";
      } else {
        $dA = $dA." $a[$i][$j]";
      }
    }
  }
  $dA = $dA."\end{array}\right]";
  return ($dA, @a);
}

do { 
  $under25 = 1; 
  $E1 = Matrix(
    [1,non_zero_random(-2,2),non_zero_random(-3,3),random(4,7)*random(-1,1,2)], 
    [0, 1, random(2,4)*random(-1,1,2),non_zero_random(-3,3)], 
    [0,0,1,random(2,5)*random(-1,1,2)], 
    [0,0,0,1]
  );
  $E2 = Matrix(
    [1,0,0,0],
    [non_zero_random(-2,2),1,0,0],
    [random(-1,1,2)*random(2,3), random(-1,1,2),1,0], 
    [random(-1,1,2)*random(2,3),non_zero_random(-2,2),2,1]
  );
  $E3 = Matrix(
    [0,0,1,0],
    [0,0,0,1],
    [1,0,0,0],
    [0,1,0,0]
  );
  $A = $E1*$E2*$E3;
  ($dA,@a) = dispMat($A);

  for ($i =0; $i < 4; $i++) {
    for ($j = 0; $j < 4; $j++) {
      $under25 = 0 if abs($a[$i][$j]) > 25;    
    }
  }
} until ($under25 == 1);

Context("Matrix");
$ech = Matrix([1,0,0,0],[0,1,0,0],[0,0,1,0], [0,0,0,1]);
$pop = PopUp(["select", "yes", "no"], "no");
$soln = MultiAnswer(0,0,0,0,0,0)->with(
  singleResult=>0,
  allowBlankAnswers => 1,
  checker => sub {    
    my ($cor, $stud, $self) = @_;
    my ($c1,$c2,$c3,$c4,$c5,$c6) = @{$cor};
    my ($s1,$s2,$s3,$s4,$s5,$s6) = @{$stud};
    @a = (Value::classMatch($s1,'String'),Value::classMatch($s2,'String'),Value::classMatch($s3,'String'),Value::classMatch($s4,'String'),Value::classMatch($s5,'String'),Value::classMatch($s6,'String'));
    return [$a[0],$a[1],$a[2],$a[3],$a[4],$a[5]];
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Given that the matrix 
$PAR\[ $dA \]$PAR
is the augmented matrix for a linear system, use technology to perform the row operations needed to transform the matrix to reduced echelon form. Then determine if the system is consistent and if it is, find all solutions to the system.
$PAR
Reduced echelon form: \{ans_rule(50)\}
$PAR
Is the system consistent? \{$pop->menu()\}
$PAR
Solution: \((x_1,x_2,x_3) = \bigg(\) \{$soln->ans_rule(4)\} \(+\) \{$soln->ans_rule(4)\} \(s_1\), \{$soln->ans_rule(4)\} \(+\) \{$soln->ans_rule(4)\} \(s_1\), \{$soln->ans_rule(4)\} \(+\) \{$soln->ans_rule(4)\} \(s_1\bigg)\)
$PAR ${BBOLD}Help:${EBOLD} To enter a matrix use $BBOLD [[ ],[ ]] $EBOLD. For example, to enter the \(2\times 3\) matrix
$PAR\[
\left[\begin{array}{ccc}
1 & 2 & 3 \\ 6 & 5 & 4
\end{array}\right]
\]$PAR you would type [[1,2,3],[6,5,4]], so each inside set of [ ] represents a row. If there is no free variable in the solution, then type \(0\) in each of the answer blanks directly before each \(s_1\). For example, if the answer is \((x_1,x_2,x_3) = (5,-2,1)\), then you would enter \((5 + 0s_1, -2 + 0s_1, 1 + 0 s_1)\).  If the system is inconsistent, you do not have to type anything in the ${LQ}Solution${RQ} answer blanks.
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($ech->cmp());
ANS($pop->cmp());
ANS($soln->cmp());


###########################################################################
# solution
###########################################################################
Context()->texStrings;
BEGIN_SOLUTION
${BBOLD}SOLUTION:${EBOLD} 
Using a computer algebra system, the row echelon form is
$PAR \[$ech.\] $PAR
Since the last row corresponds to \(0 = 1\), the linear system is inconsistent, and there are no solutions.
END_SOLUTION
Context()->normalStrings;

COMMENT('Created in 2013.
$BR
Desc: Use a CAS to write the system in reduced echelon form and then solve.');
ENDDOCUMENT();
