##DESCRIPTION
## Find eigenvectors given eigenvalues
##ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Institution(Oklahoma State University)
## Author(Jiri Lebl)
## Level(3)
## TitleText1('Notes on Diffy Qs: Differential Equations for Engineers')
## AuthorText1('Jiri Lebl')
## EditionText1('5')
## Section1('3.4')
## Problem1('')
## MO(1)
## KEYWORDS('Eigenvector', 'Eigenvalue')

##################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

#####################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

Context("Matrix");

$cnt = 0;

do {
	do {
		$lambda1 = non_zero_random(-5,5,1);
		$lambda2 = non_zero_random(-5,5,1);
	} while ($lambda1 == $lambda2);

	do {
		$a = non_zero_random(-5,6,1);
		$d = random(1,6,1);
	} while ((abs($a*$d-1) >= 10) or ($a*$d-1 == 0));

	$E = Matrix([[$a,$a*$d-1],[1,$d]]);
	$Einv = $E->inverse;

	$A = $E*Matrix([[$lambda1,0],[0,$lambda2]])*$Einv;

	$cnt = $cnt+1;
} while((abs($A->element(1,1)) >= 21) or
        ($A->element(1,1) == 0) or
        (abs($A->element(2,1)) >= 21) or
        ($A->element(2,1) == 0) or
        (abs($A->element(1,2)) >= 21) or
        ($A->element(1,2) == 0) or
        (abs($A->element(2,2)) >= 21) or
        ($A->element(2,2) == 0));

$v1 = Matrix([[$a],[1]]);
$v2 = Matrix([[$a*$d-1],[$d]]);

#####################################

Context()->texStrings;
BEGIN_TEXT
The matrix \($A\) has
eigenvalues
\( \lambda_1 = $lambda1 \) and \( \lambda_2 = $lambda2 \). 
Find the corresponding eigenvectors.

$PAR
\( {\vec{v}_1} =  \)
\{ $v1->ans_array(5) \}
and
\( \ {\vec{v}_2} =  \)
\{ $v2->ans_array(5) \}
\{ AnswerFormatHelp("matrices") \}
END_TEXT
Context()->normalStrings;

#####################################


ANS($v1->cmp(
  checker => sub {
    my ($correct,$student,$ansHash) = @_;
    return ($A * $student == $lambda1 * $student ? 1 : 0);
  }
));

ANS($v2->cmp(
  checker => sub {
    my ($correct,$student,$ansHash) = @_;
    return ($A * $student == $lambda2 * $student ? 1 : 0);
  }
));

COMMENT('Given distinct real eigenvalues, find eigenvectors.');

ENDDOCUMENT(); 
