# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus', 'integral', 'differential equations')
## Tagged by glr 02/08/09

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Introduction to Differential Equations')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('11.1')
## Section2('11.1')
## Problem1('13')
## Problem2('15')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')
## HHSection1('What is a Differential Equation')
## HHSection2('What is a Differential Equation')



DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
"parserMultiAnswer.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( C=>"Real", n=>"Real" );

$showPartialCorrectAnswers = 1;

$a = random(2,4,1);
$x0 = random(2,5,1);
$y0 = random(10,50,5);

$ans0 = Formula("0");
$ans1 = Formula("C x^($a)");

$multipart = MultiAnswer($ans0, $ans1)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f2stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
);

Context("Numeric");

TEXT(beginproblem());
Context()->texStrings;
Context()->texStrings;
BEGIN_TEXT

${BBOLD}(a)$EBOLD
If \( y=Cx^n \) is a solution to the differential equation
\[ x\frac{dy}{dx} - $a y = 0, \]
What are the possible solutions?
$BR
\( y = \) \{$multipart->ans_rule(35)\}, or
\( y = \) \{$multipart->ans_rule(35)\}.
$BR
${BITALIC}(Give your answer(s) as specifically as possible,
plugging in value(s) for \(C\) or \(n\) as needed to ensure
that the answer(s) are solutions.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
If the solution additionally satisfies  \(y=$y0\) when \(x=$x0\),
what is the solution?
$BR
\(y = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;
Context()->normalStrings;

ANS( $multipart->cmp() );
ANS( Compute("$y0*(x/$x0)^($a)")->cmp() );

Context()->texStrings;
Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
If \(y=C x^n\) is a solution to the given differential equation,
then we must have
\[
x \frac{d \ (C x^n)}{dx} - $a (C x^n) = 0,
\]
so
\[
x(C n x^{n-1}) - $a (Cx^n) = 0,
\]
or
\[
Cnx^n - $a Cx^n = C(n - $a) x^n = 0.
\]
Thus, if \(C = 0\) we have a solution, \(y = 0\); and if \(C\ne0\)
then we must have \(n=$a\), \(C\) may have any value (other than zero),
and the solution is \(y = C x^{$a}\).

$PAR
${BBOLD}(b)$EBOLD
If \(y=$y0\) for \(x=$x0\), we cannot have \(C=0\). Thus, by part (a),
we get \(n=$a\).  The solution to the differential equation is
\(y = C x^{$a}\).  Thus, if \(y = $y0\) when \(x = $x0\),
\[ $y0 = C ($x0)^{$a}, \]
so that \(C = \frac{$y0}{($x0)^{$a}}\), and
\(y = \frac{$y0}{($x0)^{$a}}\cdot x^{$a}\).

END_SOLUTION
Context()->normalStrings;
Context()->normalStrings;

ENDDOCUMENT();
