## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Linear independence)
## DBsection(Testing for linear independence)
## Date(August 2015)
## Institution(University of Michigan)
## Author(David Sher)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Otto Bretscher')
## EditionText1('5')
## Section1('3.2')

## New problem for section 3.2, number 2

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Vector");

$a=non_zero_random(-6,6,1);
$b=non_zero_random(-6,6,1);
$c=non_zero_random(-4,4,1);
$d=non_zero_random(-4,4,1);
$e=non_zero_random(-3,3,1);
$f=non_zero_random(-3,3,1);
$g=non_zero_random(2,5,1);
$h=non_zero_random(2,5,1);

## we have to have bd - ae != 0 for these to have a chance at
##    linear independence
if ( $b*$d - $a*$e == 0 ) {
    $a = -1*$a;
}

$v1=ColumnVector($a,$b,$c);
$v2=ColumnVector($d,$e,$f);
$vcomb=$g*$v1+$h*$v2;

$c1=$a*$g+$d*$h;
$c2=$b*$g+$e*$h;

$answ = Compute($c*$g+$f*$h);
## really, almost any value will work
$ansx = Compute($answ + 1);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Let \(\vec v_1=$v1\), \(\vec v_2=$v2\), and \(\vec v_3=\begin{bmatrix}$c1\\ $c2\\ k\end{bmatrix}\). First find a value of \(k\) for which the vectors \(\vec v_1,\) \(\vec v_2,\) and \(\vec v_3\) are linearly dependent:
$BR
\(k=\) \{ $answ->ans_rule()\}.
$BR
Now find a value of \(k\) for which the vectors \(\vec v_1,\) \(\vec v_2,\) and \(\vec v_3\) are linearly independent:
$BR
\(k=\) \{ $answ->ans_rule()\}.

END_TEXT
Context()->normalStrings;

ANS( $answ->cmp(
		checker => sub {
		    my ( $cor, $stu, $ans ) = @_;
		    # this is really the product ($b*$d - $a*$e)*(the below),
		    #    but we've guaranteed that the first term isn't zero
		    return ($c*$g + $f*$h - $stu) == 0;
		} ) );
ANS( $ansx->cmp(
		checker=>sub {
		    my ( $cor, $stu, $ans ) = @_;
		    # this is really the product ($b*$d - $a*$e)*(the below),
		    #    but we've guaranteed that the first term isn't zero
		    return ($c*$g + $f*$h - $stu) != 0;
		} ) );

ENDDOCUMENT();
