# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Simplify the following expression
#          x^m
#    (  ---------    ) ^ k
#       a y^n z^ p
#
# m, n, p,  are integers on [2,10]; k is an integer on [2,4]
# and a is integer on [2,3]
#
# Last edited: Carl Yao 6/28/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions')
## KEYWORDS('exponent','simplify','fraction')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

#          x^m
#    (  ---------    ) ^ k
#       a y^n z^ p
#
# m, n, p,  are integers on [2,10]; k is an integer on [2,4]
# and a is integer on [2,3]

Context("RationalFunction-Strict");
Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});


$var1 = 'x';
$var2 = 'y';
$var3 = 'z';
Context()->variables->are($var1=>'Real',$var2=>'Real',$var3=>'Real');

$a = random(2,3,2);
$m = random(2,10,1);
$n=$m;
$n = random(2,10,1) while($n==$m);
$p=$n;
$p = random(2,10,1) while($p==$n);
$k = random(2,3,1);

$total1 = $m*$k;
$total2 = $n*$k;
$total3 = $p*$k;
$denom = $a**$k;
$ans = Formula("$var1^($total1)/($denom*$var2^($total2)*$var3^($total3))");


##############################################
TEXT(beginproblem());

BEGIN_PGML
Use the properties of exponents to simplify the following
   
    [`
       \left(\displaystyle\frac{[$var1]^{[$m]}}{[$a][$var2]^{[$n]}[$var3]^{[$p]}}\right)^{[$k]}=
    `] [__________________________]

Be careful to use grouping parenthesis appropriately- remember that 1/(2x) is different from 1/2x.
END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 


BEGIN_PGML_SOLUTION
The quotient to power rule says that

    [`
          \left(\displaystyle\frac{a}{b}\right)^m = \displaystyle\frac{a^m}{b^m}
    `]

Also remember that 

     [`
          (a^m)^n = a^{mn}
     `]

Using these two tools together allows us to write

     [`
        \begin{aligned}
       \left(\displaystyle\frac{[$var1]^{[$m]}}{[$a][$var2]^{[$n]}[$var3]^{[$p]}}\right)^{[$k]}=
                       & = \displaystyle\frac{[$var1]^{[$m]\cdot[$k]}}{[$a]^{[$k]}[$var2]^{[$n]\cdot[$k]}[$var3]^{[$p]\cdot[$k]}}\\
                       & = \displaystyle\frac{[$var1]^{[$total1]}}{[$denom][$var2]^{[$total2]}[$var3]^{[$total3]}}\\
        \end{aligned}
      `]


END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
