# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Simplify the following expression
#       a x^m
#    (  -----  ) ^ p
#       b y^n
#
# a, b, m, n are integers on [2,10]; p is an integer on [2,3]
#
# Last edited: Hughes 9/6/13, Carl Yao 6/28/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions')
## KEYWORDS('exponent','simplify','fraction')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("RationalFunction-Strict");

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});

$var1 = 'x';
$var2 = 'y';
Context()->variables->are($var1=>'Real',$var2=>'Real');

$m = random(2,10,1);
$n = $m;
$n = random(2,10,1) while($n==$m);
$p = random(2,3,1);

#       a x^m
#    (  -----  ) ^ p
#       b y^n

$total1 = $m*$p;
$total2 = $n*$p;

$a = random(3,9,2)*(-1);
$b = $a;
while(gcd($a**$p,$b**$p)!=1)
{
     $a = random(3,9,2)*(-1);
     $b = random(2,10,2);
}

$numerator = $a**$p;
$denom = $b**$p;
$ans = Formula("$numerator*$var1^($total1)/($denom*$var2^($total2))");


##############################################
TEXT(beginproblem());

BEGIN_PGML
Use the properties of exponents to simplify the following
   
    [`
       \left(\displaystyle\frac{[$a][$var1]^{[$m]}}{[$b][$var2]^{[$n]}}\right)^{[$p]}=
    `] [__________________________]

Be careful to use grouping parenthesis appropriately- remember that 1/(2x) is different from 1/2x.
END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 



BEGIN_PGML_SOLUTION
The quotient to power rule says that

    [`
          \left(\displaystyle\frac{a}{b}\right)^m = \displaystyle\frac{a^m}{b^m}
    `]

Also remember that 

     [`
          (a^m)^n = a^{mn}
     `]

Using these two tools together allows us to write

     [`
        \begin{aligned}
       \left(\displaystyle\frac{[$a][$var1]^{[$m]}}{[$b][$var2]^{[$n]}}\right)^{[$p]}
                       & = \displaystyle\frac{([$a])^{[$p]}[$var1]^{[$m]\cdot[$p]}}{[$b]^{[$p]}[$var2]^{[$n]\cdot[$p]}}\\
                       & = \displaystyle\frac{[$numerator][$var1]^{[$total1]}}{[$denom][$var2]^{[$total2]}}\\
        \end{aligned}
      `]


END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
