# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Simplify the following expression
#       a x^m
#    (  -----  ) ^ n
#       b
#
# a and b are integers on [2,10] and n is an integer on [2,3]
#
# Last edited: Hughes 9/7/13, Carl Yao 6/28/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions')
## KEYWORDS('exponent','simplify','fraction')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextLimitedPolynomial.pl",
  "answerHints.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("LimitedPolynomial-Strict");
Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});


$var = 'x';
$m = random(2,10,1);
$n = random(2,3,1);
$a = random(3,9,2);
$b = $a;
$b = random(2,10,2) while(gcd($a**$n,$b**$n)!=1);

$total = $m*$n;
$numerator = $a**$n;
$denom = $b**$n;
$ans = Formula("$numerator*$var^($total)/$denom");


##############################################
TEXT(beginproblem());

BEGIN_PGML
Use the properties of exponents to simplify the expression.
   
    [`
       \left(\displaystyle\frac{[$a][$var]^{[$m]}}{[$b]}\right)^{[$n]}=
    `] [__________________________]

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 



BEGIN_PGML_SOLUTION
The quotient to power rule says that

    [`
          \left(\displaystyle\frac{a}{b}\right)^m = \displaystyle\frac{a^m}{b^m}
    `]

Also remember that 

     [`
          (a^m)^n = a^{mn}
     `]

Using these two tools together allows us to write

     [`
        \begin{aligned}
            \left(\displaystyle\frac{[$a][$var]^{[$m]}}{[$b]}\right)^{[$n]}
                       & = \displaystyle\frac{[$a]^{[$n]}[$var]^{[$m]\cdot[$n]}}{[$b]^{[$n]}}\\
                       & = \displaystyle\frac{[$numerator][$var]^{[$total]}}{[$denom]}\\
        \end{aligned}
      `]


END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
