# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
#
#  Rewrite the following expression using only positive exponents.
#
#      x^m
#    ( --- )^-n
#       a
#
# We choose m is integer on [2,15] 
# a is integer on [2,5]
# n is integer on [2,4]
# 
# Last updated: Hughes 9/4/13, Carl Yao, 7/03/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','fraction','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("RationalFunction-Strict");
$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});

Context()->flags->set(limits=>[0.95,1.05]);


#      x^m
#    ( --- )^-n
#       a
#
# We choose m is integer on [2,15] 
# a is integer on [2,5]
# n is integer on [2,4]

# exponents
$m = random(2,15,1);
$n = random(2,4,1);
$a = random(2,4,1);

# total
$total = ($m*$n);
$numerator = $a**$n;
$ans = Formula("$numerator/($var^($total))");

##############################################
TEXT(beginproblem());

BEGIN_PGML
Simplify the following expression, and write your answer using only _positive_ exponents.

    [`
                \displaystyle\left(\frac{[$var]^{[$m]}}{[$a]}\right)^{[$n*(-1)]}=
     `][____________]

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 


BEGIN_PGML_SOLUTION

We use the properties of exponents to help us here, remembering that [`\left(a^m\right)^{n}=a^{mn}`]

    [`
           \begin{aligned}
                 \left(\frac{[$var]^{[$m]}}{[$a]}\right)^{[$n*(-1)]}
                      &= \frac{[$var]^{[$m]\cdot([$n*(-1)])}}{[$a]^{[$n*(-1)]}}\\
                      &= \frac{[$var]^{[$m*$n*(-1)]}}{[$a]^{[$n*(-1)]}}\\
                                    & = [$ans]
           \end{aligned}
    `]

You might have preferred to do this problem a little differently:


    [`
           \begin{aligned}
                 \left(\frac{[$var]^{[$m]}}{[$a]}\right)^{[$n*(-1)]}
                      &= \left(\frac{[$a]}{[$var]^{[$m]}}\right)^{[$n]}\\
                      &= \frac{[$a]^{[$n]}}{[$var]^{[$m*$n]}}\\
                                    & = [$ans]
           \end{aligned}
    `]
END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
