# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
#
#  Rewrite the following expression using only positive exponents.
#
#    (x^m y^n)^-p
#
# We choose m and n integer on [2,20] and p integer on [2,5].
# 
# Last updated: Hughes 9/4/13, Carl Yao, 7/03/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','fraction','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("RationalFunction-Strict");
$var1 = RandomVariableName(type=>'variable');
$var2 = $var1;
$var2 = RandomVariableName(type=>'variable') while($var2 eq $var1);
Context()->variables->are($var1=>'Real',$var2=>'Real');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});





#
#    (x^m y^n)^-p
#
# We choose m and n integer on [2,20] and p integer on [2,5].

# exponents
$m = random(3,15,1);
$n = random(3,15,1);
$p = random(2,5,1);

# total
$total1 = ($m*$p);
$total2 = ($n*$p);
$ans = Formula("1/($var1^($total1)*$var2^($total2))");

##############################################
TEXT(beginproblem());

BEGIN_PGML
Simplify the following expression, and write your answer using only _positive_ exponents.

    [`
                \left([$var1]^{[$m]}[$var2]^{[$n]}\right)^{[$p*(-1)]}=
     `][____________]

You'll need to be careful with grouping parenthesis in this problem because 1/2[$var1] is not the same as 1/(2[$var1]).
END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 


BEGIN_PGML_SOLUTION

We use the properties of exponents to help us here, remembering that [``\frac{a^m}{a^n} = a^{m-n}``] and [`\left(a^m\right)^{n}=a^{mn}`]

    [`
           \begin{aligned}
                 \left([$var1]^{[$m]}[$var2]^{[$n]}\right)^{[$p*(-1)]}
                      &= [$var1]^{[$m]\cdot([$p*(-1)])}[$var2]^{[$n]\cdot([$p*(-1)])}\\
                      &= [$var1]^{[$m*$p*(-1)]}[$var2]^{[$n*$p*(-1)]}\\
                                    & = [$ans]
           \end{aligned}
    `]

END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
