# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
#
#  Rewrite the following expression using only positive exponents.
#
#   b * a * x^m
#   -----------
#       a * x^n
#
# We make sure that n>m so that x is raised to a negative power.
# 
# Last updated: Hughes 9/5/13, Carl Yao, 7/01/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','fraction','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("RationalFunction-Strict");
$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});


# exponents
$n = random(3,50,1);
$m = random(2,$n-1,1);

# coefficients
$a = random(2,10,1);
$b = (random(1.5,9.5,1)+0.5)*random(-1,1,2);

# total exponent
$total = (($n)*(-1)+$m)*(-1);
$totalCoeff = $a*$b;
$ans = Formula("$b/($var^($total))");

##############################################
TEXT(beginproblem());

BEGIN_PGML
Rewrite the expression simplified and using only positive exponents.

    [`
        \displaystyle\frac{[$totalCoeff][$var]^{[$m]}}{[$a][$var]^{[$n]}}=
     `][____________]

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 

BEGIN_PGML_SOLUTION

We use the properties of exponents to help us here, remembering that [``\frac{a^m}{a^n} = a^{m-n}``]

    [`
           \begin{aligned}
                 \displaystyle\frac{[$totalCoeff][$var]^{[$m]}}{[$a][$var]^{[$n]}} &= [$b][$var]^{[$m]-[$n]}\\
                                                       & =[$b][$var]^{[$total*(-1)]}\\
                                    & = [$ans]
           \end{aligned}
    `]

END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
