# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
# Here is an expression with negative exponents.
#
#     x^-m
#     ----
#     y^-n
#
# Rewrite the expression using only positive exponents.
#
# m and n are integers on [2,20].
# 
# Last updated: Hughes 9/7/13, Carl Yao, 7/01/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','fraction','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("RationalFunction-Strict");
Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});


$var1 = RandomVariableName(type=>'variable');
$var2=$var1;
$var2 = RandomVariableName(type=>'variable') while($var2 eq $var1);
Context()->variables->are($var1=>'Real',$var2=>'Real');

# the exponents
$m = random(2,20,1);
$n=$m;
$n = random(2,20,1) while($n==$m);


$ans = Formula("$var2^$n/$var1^$m");

##############################################
TEXT(beginproblem());

BEGIN_PGML
Here is an expression with negative exponents.

    [`
        \displaystyle\frac{[$var1]^{[$m*(-1)]}}{[$var2]^{[$n*(-1)]}}=
     `][____________]

Rewrite the expression using only positive exponents.

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 


BEGIN_PGML_SOLUTION

We evaluate the expression by remembering that [`x^{-n}`] is the same thing as [``\frac{1}{x^n}``] for any non-zero,
real value of [`x`].

The first thing we do is to rewrite the expression using fractions- then we remember that when 
we divide by a fraction it is the same thing as multiplying by its reciprocal. 

    [``
        \begin{aligned}
         \displaystyle\frac{[$var1]^{[$m*(-1)]}}{[$var2]^{[$n*(-1)]}}
            & = \displaystyle\frac{\displaystyle\frac{1}{[$var1]^{[$m]}}}{\displaystyle\frac{1}{[$var2]^{[$n]}}}\\
            & = \displaystyle\frac{1}{[$var1]^{[$m]}}\cdot \frac{[$var2]^{[$n]}}{1}\\
            & = [$ans]
        \end{aligned}
    ``]

Effectively we 'flip' the fraction since both variables are raised to a negative exponent.

END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
