# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
#
#  Rewrite the following expression using only positive exponents.
#
#   a * x^m
#   --------
#   b * x^n
#
# We make sure that n>m so that x is raised to a negative power. 
#
# a and b are both odd numbers on [3,15].
# 
# Last updated: Hughes 9/4/13, Carl Yao, 7/01/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','fraction','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("RationalFunction-Strict");
$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});

### Context()->variables->set($var=>{limits=>[1,10]});

# exponents
$n = random(3,50,1);
$m = random(2,$n-1,1);

# coefficients
$a = random(3,15,2)*random(-1,1,2);
$b = $a;
$b = random(3,15,2) while( gcd($b,$a)!=1 ); 

# total exponent
$total = (($n)*(-1)+$m)*(-1);
$ans = Formula("$a/($b*$var^($total))");



##############################################
TEXT(beginproblem());

BEGIN_PGML

Simplify the following expression, and write your answer using only _positive_ exponents.

    [`
        \displaystyle\frac{[$a][$var]^{[$m]}}{[$b][$var]^{[$n]}}=
     `][____________]

Note that you may have to be careful with grouping parenthesis in this problem; 1/2[$var]
is not the same as 1/(2[$var]).

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 


#Move negative symbol to the front of fraction in display.
if ($a<0) {
   $absA = -$a;
   $output = "-\frac{$absA}{$b $var^{$total}}";
}
else {
   $output = "\frac{$a}{$b $var^{$total}}";
}

BEGIN_PGML_SOLUTION

We use the properties of exponents to help us here, remembering that [``\frac{a^m}{a^n} = a^{m-n}``]

    [``
           \begin{aligned}
                 \displaystyle\frac{[$a][$var]^{[$m]}}{[$b][$var]^{[$n]}} &= \frac{[$a]}{[$b]}[$var]^{[$m]-[$n]}\\
                                                       & =\frac{[$a]}{[$b]}[$var]^{[$total*(-1)]}\\
                                    & = [$output]
           \end{aligned}
    ``]

END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
