## DESCRIPTION
## exponents
## ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## BookChapter('Rules for Exponents and the Reasons for Them')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Algebra: Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('6.1')
## Problem1('70')
## KEYWORDS('exponents')

##############################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(t=>"Real");

$a = random(3,5,1);
$b = 3;#random(3,5,1);

$fac1 = Formula("$a*100*t^2")->reduce;
$fac2 = Formula("e^(2*$b*t)")->reduce;

$multians = MultiAnswer($fac1,$fac2)->with(
  singleResult => 1,  
  separator => " * ",
  tex_separator => " \cdot ",

  checker => sub {
    my $correct = shift; my $student = shift; my $self = shift;
    my ($F,$G) = @{$correct};
    my ($f,$g) = @{$student};

    Value::Error('Neither factor can be constant')
    unless $f->isFormula && $g->isFormula;

    Value::Error('Your product does not equal the original (it is incomplete)') 
    unless $F*$G == $f*$g;
#    return 0 unless $F*$G == $f*$g;

    my $context = Context()->copy;
    $context->flags->set(no_parameters=>0);
    $context->variables->add('a'=>'Parameter');
    my $a = Formula($context,'a');
    $f = Formula($context,$f);
    my $result = ($a*$F == $f || $a*$G == $f);
    Value::Error('Factor as the product of a polynomial and an exponential function') 
    unless ($result == 1);
    return $result;

  }

);



##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Using laws of exponents, rewrite the following expression as a product.
$BR
$BR
\( $a \big( 10 t e^{$b t} \big)^2 \) = 
\{$multians->ans_rule(10)\}
\( \cdot \) 
\{$multians->ans_rule(10)\}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


#############################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );


;
ENDDOCUMENT();
