## DESCRIPTION
## exponents
## ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## BookChapter('Rules for Exponents and the Reasons for Them')
## DBsubject(Arithmetic)
## DBchapter(Integers)
## DBsection(Exponentiation)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Algebra: Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('6.1')
## Problem1('78')
## KEYWORDS('exponents')

##############################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric");

$a = random(3,5,1);
$b = random(3,5,1);
$c = random(2,3,1);

$fac1 = Formula("$c^$a")->reduce;
$fac2 = Formula("$c^$b")->reduce;

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);

##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Using laws of exponents, rewrite the following expression as a product.
$BR
$BR
\( $c^{$a + $b} \) = 
\{$multians->ans_rule(6)\}
\( \cdot \) 
\{$multians->ans_rule(6)\}
\{ AnswerFormatHelp("numbers") \}
END_TEXT
Context()->normalStrings;


#############################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );


;
ENDDOCUMENT();
