## DESCRIPTION
## exponents
## ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## BookChapter('Rules for Exponents and the Reasons for Them')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## MLT(ExponentsFillBoxes)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Algebra: Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('6.1')
## Problem1('72')
## KEYWORDS('exponents')

##############################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(x=>"Real",n=>"Real");

$a = random(2,5,1);
$b = random(3,4,1);
$c = random(2,3,1);

$fac1 = Formula("($a)^($c*n)")->reduce;
$fac2 = Formula("x^($b*$c*n)")->reduce;

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);

##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Rewrite the following expression without the existing 
parentheses.
Simplify your answer as much as possible, and 
assume that all variables are positive.
\{ AnswerFormatHelp("formulas") \}
$BR
$BR
\( \big( $a x^{$b} \big)^{$c n} \) = 
\{$multians->ans_rule(6)\}
\( \cdot \) 
\{$multians->ans_rule(6)\}
END_TEXT
Context()->normalStrings;


#############################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );


;
ENDDOCUMENT();
