## DESCRIPTION
##  Find non-diagonal square root
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Matrix algebra)
## Date(15/03/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(3)
## MO(1)
## Static(0)




##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
);


TEXT(beginproblem());

COMMENT('non-uniqueness of square roots');

Context("Matrix");

my $a = random(2,10);

$M = Matrix(
 [ $a**2 ,    0  ],
 [    0  , $a**2 ]
);

$A = Matrix(
 [ $a ,  1 ],
 [ 0 , -$a ]
);


#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Find a $BBOLD non-diagonal $EBOLD matrix \(A\) so that 
\( A^2 = $M\)
$PAR
\(\quad A = \) \{ $A->ans_array(2); \}


END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation


ANS( $A->cmp(checker=>sub {
 my ($c,$s,$hash) = @_;
 my $C = Matrix($c);
 my $S = Matrix($s);
 if (($S->element(1,2) == 0) && ($S->element(2,1) == 0)) {
   Value->Error("Give a non-diagonal answer!");
 }
 if ($S*$S != $C*$C) {
   Value->Error("\(A^2\) is incorrect!");
 }
 return 1;
}) );



COMMENT('');

ENDDOCUMENT(); 
