## DESCRIPTION
##  Find matrix with nilpotence 3
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Matrix algebra)
## Date(15/03/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(3)
## MO(1)
## Static(0)




##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
);


TEXT(beginproblem());

COMMENT('Find 3x3 matrix with nilpotency 3 -- shift matrix');

Context("Matrix");

$A = Matrix(
 [ 0, 1, 0 ],
 [ 0, 0, 1 ],
 [ 0, 0, 0 ]
);


#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Find a nonzero \(3\times 3\) matrix \(A\) so that 
\( A^3 \) $BBOLD is$EBOLD zero but \( A^2 \) $BBOLD is not$EBOLD zero. 
$PAR
\(\quad A = \) \{ $A->ans_array(2); \}


END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation


ANS( $A->cmp(checker=>sub {
 my ($c,$s,$hash) = @_;
 my $S = Matrix($s);
 if ($S->isZero) {
   Value->Error("\(A\) should not be zero!");
 }
 if (($S*$S)->isZero) {
   Value->Error("\(A^2\) should not be zero!");
 }
 if (($S*($S*$S))->isZero) {
   return 1;
 }
 Value->Error("\(A\) should be nilpotent!");
 return 0;
}) );



COMMENT('');

ENDDOCUMENT(); 
