##DESCRIPTION
## solve by integrating factor and separation of variables 
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Integrating factor)
## Date(1/7/2010)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(3)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('2.3')
## Problem1('')
## KEYWORDS('linear' 'differential' 'equation')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "parserFormulaUpToConstant.pl",
  "MathObjects.pl",
  "answerHints.pl",
  "contextLimitedPowers.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;
##############################################################
#
#  Setup
#
#

# Set up a context that allows differential notation
$diff = Context("Numeric");
$diff->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
$diff->variables->are("x"=>"Real","y"=>"Real","y'"=>"Real",C=>"Real");


# Set up a context that forces students to reduce their answers.

$reduce = Context("Numeric")->copy;
$reduce->functions->undefine("exp","ln");
$reduce->{error}{msg}{"Function 'ln' is not allowed in this context"} 
= "e raised to a logarithmic power can be reduced";


$left = Formula($diff,"y'-2/(3x)y");
$right = Formula("0");
$integratingfactor = Formula($reduce,"x^(-2/3)"); #If you change this must also change it in the answer checker
$integratingfactor->{limits} = [1,5];
$newleft = Formula($diff,"x^(-2/3)*y");
$newright = Formula($diff,"x^(-2/3)*$right");
$integrateleft = $newleft;
$integrateright = FormulaUpToConstant("0");
$solution = Formula($diff,"Cx^(2/3)");
$solution->{limits} = [1,5];

Context()->texStrings;
BEGIN_TEXT
According to Ince [pg. 531] the first known use of integrating factors to solve a differential equation was by Fatio de Duiller in June of 1687. He was solving the equation \[3xdy - 2ydx = 0\] which we would write in standard form (using the prime notation) as
$BR
$BCENTER \{$left->ans_rule(15)\} \( =\) \{$right->ans_rule(5)\}. $ECENTER
$BR
For this equation the integrating factor is: \{$integratingfactor->ans_rule(20)\}
$BR
After multiplying both sides by the integrating factor and unapplying the product rule we get the new differential equation:
$BR
$BCENTER \(\frac{d}{dx}\big[\) \{$newleft->ans_rule(20)\} \(\big] = \)  \{$newright->ans_rule(20)\} $ECENTER
$BR 
Integrating both sides we get the algebraic equation
\{$integrateleft->ans_rule(30)\} \(=\) \{$integrateright->ans_rule(30)\}
$BR
Solving for y, the solution to the differential equation is y = \{$solution->ans_rule(20)\} (using C as the constant)
$BR
$BR
[Ince] Ince E L, Ordinary Differential Equations, Longmans, Green and Co, London, 1927.
END_TEXT
Context()->normalStrings;
ANS($left->cmp());
ANS($right->cmp());
ANS($integratingfactor->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $contextIF = $reduce->copy;
      $contextIF->flags->set(no_parameters=>0);
      $contextIF->variables->add('C0'=>'Parameter');
      my $c0 = Formula($contextIF,'C0');
      $student = Formula($contextIF,$student);
      $correct = Formula($contextIF,"$c0*x^(-2/3)");
      $zero = Formula($contextIF,"0");
      return (($correct == $student)&&($zero!=$student));
    }));
 # Any constant multiple of the integrating factor will work.
ANS($newleft->cmp());
ANS($newright->cmp());
ANS($integrateleft->cmp());
ANS($integrateright->cmp());
ANS($solution->cmp());

ENDDOCUMENT();
