##DESCRIPTION
## solve by integrating factor and separation of variables 
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Integrating factor)
## Date(1/7/2010)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(3)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('2.3')
## Problem1('')
## KEYWORDS('linear' 'differential' 'equation')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "compoundProblem.pl",
  "MathObjects.pl",
  "parserFormulaUpToConstant.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;
##############################################################
#
#  Setup
#
#

# Make a context for the derivatives
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("x"=>"Real","y"=>"Real","y'"=>"Real","y''"=>"Real",y=>"Real");

# Make a context so that we can use the constant C outside of ParserUpToConstant
$contextC = Context()->copy;
$contextC->variables->add(C=>"Real");

###########################################################################
##
## This is a multipart problem 
##  EDIT: not anymore!
##$cp = new compoundProblem(
##  parts => 2,                  # the total number of parts in this problem
##  totalAnswers => 11,           # total answers in all parts combined
##  parserValues => 1,           # make MathObjects from student answers
##  allowReset => $isProfessor,  # professors get Reset button for testing
## );
## $part = $cp->part;             # look up the current part


$a = Compute(non_zero_random(-9,9));
$solution = Formula($contextC,"C*e^($a*x)");


##############################################
#
#  Part 1
#

## if ($part == 1){
$standardformleft = Formula("y'-$a*y");
$standardformright = Formula("0");
$integratingfactor = Formula("e^(-$a*x)"); # if you change this you must change the answer checker below
$newleft = Formula("e^(-$a*x)*y");
$newright = Formula("0");
$integrate = Formula("e^(-$a*x)*y");

Context()->texStrings;
BEGIN_TEXT
The differential equation \(\frac{dy}{dx} = $a y\) is a linear differential equation.
$BR
Convert the equation to standard form (use the prime notation for the derivative):
$BR
\{$standardformleft->ans_rule(30)\} = \{$standardformright->ans_rule(20)\}
$BR
The integrating factor is: \{$integratingfactor->ans_rule(20)\}
$BR
After multiplying both sides by the integrating factor and unapplying the product rule we get the new differential equation:
$BR
\(\frac{d}{dx}\big[\) \{$newleft->ans_rule(20)\} \(\big]\) = \{$newright->ans_rule(20)\}
$BR 
Integrating both sides we get algebraic equation
\{$integrate->ans_rule(15)\} \( = C\)
$BR
Solving for y, the solution to the differential equation is y = \{$solution->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers Part 1
#
#
ANS($standardformleft->cmp());
ANS($standardformright->cmp());
ANS($integratingfactor->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $contextIF = Context()->copy;
      $contextIF->flags->set(no_parameters=>0);
      $contextIF->variables->add('C0'=>'Parameter');
      my $c0 = Formula($contextIF,'C0');
      $student = Formula($contextIF,$student);
      $correct = Formula($contextIF,"$c0*e^(-$a*x)");
      return $correct == $student;
    }));
 # Any constant multiple of the integrating factor will work.
ANS($newleft->cmp());
ANS($newright->cmp());
ANS($integrate->cmp());
ANS($solution->cmp());

# }



##############################################
#
#  Part 2
#

## if ($part == 2) {

$separableleft = Formula("1/($a*y)");
$separableright = Formula("1");
$integrateleft = FormulaUpToConstant("ln(abs(y))/$a");
$integrateright = FormulaUpToConstant("x");
Context()->texStrings;
BEGIN_TEXT
$BR
$BR
$BR
$BR

The differential equation \(\frac{dy}{dx} = $a y\) is also a separable equation:
$BR
\{$separableleft->ans_rule(20)\}\(\frac{dy}{dx}\) = \{$separableright->ans_rule(20)\}
$BR
Integrating both sides we get the algebraic equation \{$integrateleft->ans_rule(20)\}= \{$integrateright->ans_rule(20)\}.
$BR
Solving for y we get y = \{$solution->ans_rule(30)\} (use C for the constant in your answer)
END_TEXT
Context()->normalStrings;
##############################################################
#
#  Answers Part 2
#
#
ANS($separableleft->cmp());
ANS($separableright->cmp());
ANS($integrateleft->cmp());
ANS($integrateright->cmp());
ANS($solution->cmp());

# }
COMMENT('This problem was originally a multipart problem that took so long to load it become a one part problem. Uncommenting the parts code will put it back. In the first part we solve the equation by integrating factors, in the second part we use separation of variables.');
ENDDOCUMENT();
