# DESCRIPTION
# Factor difference of squares quadratic
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Factoring')
## DBsection( 'Factoring: special forms')
## KEYWORDS('factoring', 'quadratic')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('5.5')
## Problem1('82')
## Author('')
## Institution('Utah')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",            # Union College utilities
  "MathObjects.pl",
  "PGcourse.pl",           # Customization file for the course
"parserMultiAnswer.pl",
);

TEXT(beginproblem);

###################################
# Setup
Context("Numeric");
Context()->variables->add(z=>'Real');

$fac1 = Formula("(z+ 2)");
$fac2 = Formula("(z- 2)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
###################################
# Main text

Context()->texStrings;
BEGIN_TEXT
Factor the difference of two squares. $BR \( z^2- 4= \)

\{$multians->ans_rule(10)\}
\( * \) 
\{$multians->ans_rule(10)\}
 
END_TEXT
Context()->normalStrings;

###################################
# Answers

ANS( $multians->cmp() ); 

###################################

ENDDOCUMENT();        # This should be the last executable line in the problem.
