## DESCRIPTION
## Calculating double integrals  
## ENDDESCRIPTION

## KEYWORDS('Double Integrals','Iterated Integrals')

## DBsubject('Calculus')
## DBchapter('Multiple Integrals')
## DBsection('Double Integrals over General Regions')
## Date('10/31/09')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('16.2')
## Section2('16.2')
## Problem1('43')
## Problem2('43')

DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl",
  "parserMultiAnswer.pl",
);


TEXT(beginproblem);

##############################################
#  Setup

Context("Numeric")->variables->are(x=>"Real",y=>"Real",dx=>"Real",dy=>"Real");


#-----------------------
#  limits of integration
#  A < x < B
#  C < y < D

$a = random(3,6,1);
$z1 = $a**2;
$z2 = random(4,8,1);
$z3 = $z1 + $z2;


#-----------
#  integrand
#

$f_top = Formula("$z3 - x**2 - y**2") -> reduce;
$f_bot = Formula("$z2");

$f = Formula("$z1 - x**2 - y**2") -> reduce;



#-----------
#  volume
#

#$V = Formula("$a**4 / 96");


#=================================
#  Multiple solutions
#
#  inner and outer differentials
#
#  arrays are (if x inner & y outer integral, if y inner & x outer integral)
# 
@id = (Formula("dx"),Formula("dy"));
@od = (Formula("dy"),Formula("dx"));

# inner limit high, inner limit low, outer limit high, outer limit low (case 0, case 1)
@ilh = (Formula("sqrt($z1 - y**2)"),Formula("sqrt($z1 - x**2)"));
@ill = (Formula("-(sqrt($z1 - y**2))"),Formula("-(sqrt($z1 - x**2))"));
@olh = (Formula("$a"),Formula("$a"));
@oll = (Formula("-($a)"),Formula("-($a)"));


$multipart = MultiAnswer( $f, $id[0], $od[0], $oll[0], $olh[0], $ill[0], $ilh[0] )->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $fstu, $idstu, $odstu, $ollstu, $olhstu, $illstu, $ilhstu ) = @{$student};
        if ( 
             ( 
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           return 1;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration.");
           return 0.94;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration and order of integration.");
           return 0.47;
        } else {
             return 0;
        }
    }
  );



##############################################
#  Main text

Context()->texStrings;
Context()->flags->set(reduceConstants=>0);
BEGIN_TEXT
Set up a double integral in rectangular coordinates
for calculating the volume 
of the solid under the graph of the function 
\( f(x,y) = $f_top \) and above the plane \( z = $f_bot \).
$BR
$BR
${BITALIC}Instructions:${EITALIC}
Please enter the integrand in the first answer box. 
Depending on the order of integration you choose,  
enter ${BITALIC}dx${EITALIC} and ${BITALIC}dy${EITALIC} 
in either order into the second and third answer boxes 
with only one ${BITALIC}dx${EITALIC} or 
${BITALIC}dy${EITALIC} in each box.  Then, enter the limits of 
integration.
$BR
$BR
\( \displaystyle \int_A^B \int_C^D \) 
\{ $multipart->ans_rule(40) \}
\{ $multipart->ans_rule(5) \}
\{ $multipart->ans_rule(5) \}
$BR
$BR
A = \{ $multipart->ans_rule(20) \} $BR
B = \{ $multipart->ans_rule(20) \} $BR
C = \{ $multipart->ans_rule(20) \} $BR
D = \{ $multipart->ans_rule(20) \}
END_TEXT
Context()->normalStrings;





##################################################
#  Answers

ANS( $multipart->cmp() );
#ANS(($V)->cmp());

$showPartialCorrectAnswers = 1;

COMMENT('Allows integration in either order.');

ENDDOCUMENT();