##DESCRIPTION
##  Middle School problem:
##ENDDESCRIPTION


## ProblemTester('Laura Travis')
## Publisher('Prentice Hall')
## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Applications and models)
## Date(11/17/2012)
## Institution(Mississippi College)
## Author(John Travis)
## Level(2)
## MO(1)
## TitleText1('Pre-Algebra')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('prealgebra','common core','proportions')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "problemRandomize.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#  Allow for the user to try another version of the problem once they get it correct.
ProblemRandomize(when=>"Correct",onlyAfterDue=>0);

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->flags->set(
  tolerance => 0.01,
  tolType => "absolute",
);

#@r = NchooseK(6,1);
$r = random(0,5,1);
@badchips_text = ('four','five','six','seven','eight','nine');
@badchips = (4,5,6,7,8,9);

$a = Real($badchips[$r]);
$b = Real(random(300,900,100));
$d = Real(random(10000,30000,5000));
$bads = $a*$d/$b; 

#parser::Assignment->Allow;

$multians = MultiAnswer($a,$b,Compute("x"),$d)->with(
  singleResult=>1,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($astu,$bstu,$xstu,$dstu) = @{$student};
    my ($acor,$bcor,$xcor,$dcor) = @{$correct};
    if ( ($acor == $astu && $bcor == $bstu && $xcor == $xstu && $dcor == $dstu) or 
         ($dcor == $astu && $bcor == $bstu && $xcor == $xstu && $acor == $dstu) ){ return 1; }
    else {return 0;}
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
A microchip inspector found $badchips_text[$r] defective chips in a batch containing \($b\) chips.  At that rate, determine the number of defective chips in a batch of \($d\) chips.
$PAR
Write a proportion describing this problem in the form \( \frac{a}{b} = \frac{x}{d} \) where \(x\) represents the 
number that you are asked to find.$PAR
$BCENTER
\{$multians->ans_rule(6)\}/\{$multians->ans_rule(6)\} = \{$multians->ans_rule(6)\}/\{$multians->ans_rule(6)\}
$ECENTER
$PAR
The number defective chips expected in the larger batch is \( x = \) \{ans_rule(10)\} 
END_TEXT
Context()->normalStrings;


##############################################################
#
#  Answers
#
#

ANS( $multians->cmp());
ANS( Compute("$bads")->cmp( ) );

ENDDOCUMENT();
