##DESCRIPTION
##
## Created by the UBC Faculty of Applied Science.  
## Project led by Agnes d'Entremont, Jonathan Verrett, and Negar M. Harandi, 
## with support from Paul Lucina and Joseph Yan.
##              
## Contact: agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this
## project provided by UBC Vancouver students via the Teaching
## and Learning Enhancement Fund. We also gratefully acknowledge
## additional funding support from the UBC Department of
## Electrical Engineering, and the UBC Applied Science Dean’s Office.  
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people.
##
##ENDDESCRIPTION

## DBsubject(Engineering Signals and Systems)
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## RESOURCES('UBC-SIG-18-025.png')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

#image scale
$imgScale = .2;

#image aspect ratio
$width = $imgScale*2742;
$height = $imgScale*964;

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"PGauxiliaryFunctions.pl",
	"niceTables.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
Context()->variables->add(s=>"Real");

#given
$alpha = random(100, 200,1);
$beta = random(2, 8,1);
$a = random(2, 8,1);

$mcYes = PopUp(["?","Yes","No"], "Yes");
$mcNo = PopUp(["?","Yes","No"], "No");

#computation 
$hs = Formula("$alpha/(s^2 + ($a + $beta)s + $alpha)");

$disc = ($a + $beta)**2 - 4*$alpha;

if($disc>=0){
	$sqr = sqrt($disc);
}
else{
	$sqr = sqrt(abs($disc))*i;
}

$s1 = (-$a - $beta + $sqr)/(2);
$s2 = (-$a - $beta - $sqr)/(2);


##############################################################
#
#  Text 
#
#

BEGIN_PGML

For the feedback system shown in the figure below,

*a)* determine the overall transfer function, [`H(s)`]. Assume [`a = [$a]`], [`\alpha = [$alpha]`], [`\beta = [$beta]`].  

[`H(s)=`] [________________________]

[@ image("UBC-SIG-18-025.png", width=>[$width], height=>[$height], tex_size=>500) @]*

*b)* For the input [`u(t)`], is this system BIBO stable? [@ $mcYes->menu() @]*

_Part *b* will only be marked correct if the answer to part *a* is correct._

END_PGML

ANS($hs->cmp);
$ans_hash1 = $hs->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$aCorrect= $ans_hash1->{score};

if (Re($s1) < 0 && Re($s2) < 0) {
	ANS($mcYes->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $aCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
}
else{
	ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $aCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
}



##############################################################

ENDDOCUMENT();