##DESCRIPTION
##
## Created by the UBC Faculty of Applied Science.  
## Project led by Agnes d'Entremont, Jonathan Verrett, and Negar M. Harandi, 
## with support from Paul Lucina and Joseph Yan.
##              
## Contact: agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this
## project provided by UBC Vancouver students via the Teaching
## and Learning Enhancement Fund. We also gratefully acknowledge
## additional funding support from the UBC Department of
## Electrical Engineering, and the UBC Applied Science Dean’s Office.  
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people.
##
##ENDDESCRIPTION

## DBsubject(Engineering Signals and Systems)
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## RESOURCES('UBC-SIG-18-051.png')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

#image scale
$imgScale = .18;

#image aspect ratio
$width = $imgScale*3927;
$height = $imgScale*1589;

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"parserFunction.pl"
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Matrix");
Context()->variables->add(s=>"Real");

#given
$a = random(2, 9, 1);
$b = random(2, 9, 1);
$c = random(2, 9, 1);
$alpha = random(2, 9, 1);
$beta= random(2, 20, 1);
$gamma = random(2, 9, 1);

$mcYes = PopUp(["?","Yes","No"], "Yes");

#computation 
$aA = Matrix("[[-$alpha,1,0],[-$beta,0,1],[-$gamma,0,0]]");
$aB = Matrix("[[$a],[$b],[$c]]");
$aC = Matrix("[1,0,0]");
$aD = Matrix("[[0]]");

$hs = Formula("($a s^2 + $b s + $c)/(s^3 + $alpha s^2 + $beta s + $gamma )");

$cA = Matrix("[[0,3-$beta,1-2*$beta],[-1/5,1/5*($alpha - 2*$gamma),2/5*($alpha - 2*$gamma)],[3/5,1/5*(-3 * $alpha + $gamma),2/5 * (-3 * $alpha + $gamma)]]");
$cB = 1/5*Matrix("[[5*$b],[-$a + 2 *$c],[3 * $a - $c]]");
$cC = Matrix("[0,1,2]");
$cD = Matrix("[[0]]");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

Enter your matrices with nested square brackets or you will receive an error message (e.g., \[\[a,b\],\[c,d\]\] to represent [`\begin{bmatrix}a \hspace{0.1in} b \\c \hspace{0.1in} d\end{bmatrix}`] or \[\[a\],\[b\]\] to represent [`\begin{bmatrix}a \\b \end{bmatrix}`])

The figure below shows the block-diagram for the observer canonical form representation of an LTI system. Assume that [`x_1(t)`], [`x_2(t)`] and [`x_3(t)`] are the the three state variables, [`w(t)`] is the input and [`y(t)`] is the output. In the system, [`a=[$a]`], [`b=[$b]`], [`c=[$c]`], [`\alpha=[$alpha]`], [`\beta=[$beta]`], and [`\gamma=[$gamma]`].

[@ image("UBC-SIG-18-051.png", width=>[$width], height=>[$height], tex_size=>500) @]*

*a)* Find the corresponding state-space representation of the system and enter the [`[A,B,C,D]`] matrices below.

[`A=`] [___________________]

[`B=`] [___________________]

[`C=`] [___________________]

[`D=`] [___________________]

*b)* Write the transfer function for the system, [`H(s)=\frac{Y(s)}{W(s)}`]

[`H(s)=`] [___________________]

*c)* Use the transformation matrix [`T=\begin{bmatrix}0 & 1 & 2\\1 & 0 & 0\\0 & 3 & 1\end{bmatrix}`] to obtain a different set of state-variables [`v(t)=Tx(t)`]. Enter the new set of [`[\tilde{A},\tilde{B},\tilde{C},\tilde{D}]`] below. 

[`\tilde{A}=`] [___________________]

[`\tilde{B}=`] [___________________]

[`\tilde{C}=`] [___________________]

[`\tilde{D}=`] [___________________]

*d)* Is the new representation using [`[\tilde{A},\tilde{B},\tilde{C},\tilde{D}]`] equivalent to the one you found in part *a*? [@ $mcYes->menu() @]*

_Part *d* will only be marked correct if the answers to part *c*are correct._

END_PGML

ANS($aA->cmp);
ANS($aB->cmp);
ANS($aC->cmp);
ANS($aD->cmp);

ANS($hs->cmp);

ANS($cA->cmp);
ANS($cB->cmp);
ANS($cC->cmp);
ANS($cD->cmp);

$ans_hash6 = $cA->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(6)});
$caCorrect= $ans_hash6->{score};
$ans_hash7 = $cB->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(7)});
$cbCorrect= $ans_hash7->{score};
$ans_hash8 = $cC->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(8)});
$ccCorrect= $ans_hash7->{score};
$ans_hash9 = $cD->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(9)});
$cdCorrect= $ans_hash9->{score};

ANS($mcYes->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $caCorrect==1 && $cbCorrect==1 && $ccCorrect==1 && $cdCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));

##############################################################

ENDDOCUMENT();