##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')

##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"niceTables.pl",
	"parserFunction.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->variables->add(t=>"Real");
Context()->variables->remove("x");

#given
$mc = PopUp(["?","<",">","=","<=",">="], "<");
$b_0 = random(2, 11, 1);
$b_1 = random(2, 11, 1);
$a = random($b_0, $b_0 + $b_1, 0.5);
$b = random($b_0 + $b_1, $b_0 + $b_1 + 20,  0.5);
$c = random(-8, -2, 1);
$d = random(2, 10, 1);
$e = 2.7182818284590452353602874713527;

#computation 
$q = Formula("0");
$y1a = ($e**($c*($a-$b_0))-1)/$c;
$y1b = (($e**($c*($b-$b_0)))-($e**($c*($b-$b_0-$b_1))))/$c;
$y2 = Formula("e^($c t)/(1 - e^($c $d))");

##############################################################
#
#  PGML 
#
#

BEGIN_PGML


The impulse response of an LTI system is given by: [`h(t)=e^{qt}u(t)`], where [`q`] is a real number. 

*a)* What is the condition on parameter [`q`] so that the system is BIBO stable?

[`q=`] [@ $mc->menu() @]* [_____]

*b)* Now suppose that [`q=[$c]`]. Calculate the output of the system [`y_1(t)`] at [`t=[$a]`] and [`t=[$b]`] when the input is given by [`x(t)=u(t-[$b_0])-u(t-[$b_0+$b_1])`].

[`y_1([$a])=`] [____]  
[`y_1([$b])=`] [____]  

*c)* Find the expression that describes the output of the system [`y_2(t)`] for [`0 < t < [$d]`] when [`q=[$c]`] and the input is [``x(t)=\sum_{k=-\infty}^{\infty}\delta (t-[$d]k)``].

[`y_2(t)=`] [____]  for [`0 \le t < [$d]`]

_Hint: Use the shifting property of the delta function and the geometric series._

END_PGML

ANS($mc->cmp);
$ans_hash1 = $mc->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$symCorrect= $ans_hash1->{score};

ANS($q->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $symCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));

ANS(num_cmp($y1a));
ANS(num_cmp($y1b));
ANS($y2->cmp);

##############################################################

ENDDOCUMENT();