##DESCRIPTION
##
## Created by the UBC Faculty of Applied Science.  
## Project led by Agnes d'Entremont, Jonathan Verrett, and Negar M. Harandi, 
## with support from Paul Lucina and Joseph Yan.
##              
## Contact: agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this
## project provided by UBC Vancouver students via the Teaching
## and Learning Enhancement Fund. We also gratefully acknowledge
## additional funding support from the UBC Department of
## Electrical Engineering, and the UBC Applied Science Dean’s Office.  
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people.
##
##ENDDESCRIPTION

## DBsubject(Engineering Signals and Systems)
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"parserFunction.pl"
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
Context()->variables->add(n=>"Real");
parserFunction("u(n)"=>"sin(2.71828*n/pi)+1.57*n^2");

#given
$alpha = random(0.1, .99, 0.01);
$beta = random(2, 9, 1);
$a = random(2, 12, 1);

$mcYes = PopUp(["?","Yes","No"], "Yes");
$mcNo = PopUp(["?","Yes","No"], "No");

if($alpha < 1){
	$mc = $mcYes;
}
else{
	$mc = $mcNo;
}

#computation 
$hn = Formula("$beta*$alpha^n u(n)");
$yn = Formula("$beta ($alpha^n u(n) + $alpha^(n-1)u(n-1) + $alpha^(n-2)u(n-2))");
$ypn = Formula("$beta *(1 + $alpha^(n+1))/(1 - $alpha)");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

An LTI system is represented by the difference equation [`y[n] = [$alpha] y[n-1] + [$beta] x[n]`].

*a)* Find the impulse response of the system and enter it as a simplified expression in terms of the discrete unit-step function, [`u[n]`]. 

[`h[n]=`] [_____]  

*b)* Is the system BIBO stable? [@ $mc->menu() @]*

*c)* Determine the response of the system, [`y[n]`], to the input:

[``x[n]=\begin{cases}1 & if & 0\le n \le 2 \\ 0 & if & otherwise \end{cases}``]

[`y[n]=`] [_____]  

_In your answers, enter [`z(n)`] for a discrete-time function [`z[n]`]. WebWork is unable to parse a function that uses square brackets._

_Part *b* will only be marked correct if part *a* is correct._

END_PGML

ANS($hn->cmp);

$ans_hash1 = $hn->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$aCorrect= $ans_hash1->{score};

ANS($mc->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $aCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
ANS($yn->cmp);


#*d)* Now suppose that the input to the system is [`x[n]=u[n] - u[n-[$a]]`], use convolution sum to find the output, [`y'[n]`], for [`0\le n\le [$a]`].

#[`y'[n]=`] [_____]  

#ANS($ypn->cmp);


##############################################################

ENDDOCUMENT();