##DESCRIPTION
##
## Created by the UBC Faculty of Applied Science.  
## Project led by Agnes d'Entremont, Jonathan Verrett, and Negar M. Harandi, 
## with support from Paul Lucina and Joseph Yan.
##              
## Contact: agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this
## project provided by UBC Vancouver students via the Teaching
## and Learning Enhancement Fund. We also gratefully acknowledge
## additional funding support from the UBC Department of
## Electrical Engineering, and the UBC Applied Science Dean’s Office.  
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people.
##
##ENDDESCRIPTION

## DBsubject(Engineering Signals and Systems)
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"PGauxiliaryFunctions.pl",
	"niceTables.pl",
	"parserImplicitEquation.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("ImplicitEquation");
Context()->variables->add(yp=>"Real");
Context()->variables->add(ypp=>"Real"); 

#given
$mcYes = PopUp(["?","Yes","No"], "Yes");
$mcNo = PopUp(["?","Yes","No"], "No");

$a = random(2, 9,1);
$b = random(2, 9,1);
$c = random(2, 9,1);
$f = random(2, 9,1);
$y_0 = non_zero_random(-6, 6,1);
$dy_0 = non_zero_random(0, 6,1);

#computation 
$disc = $b*$b - 4 * $a * $c;

if($disc>=0){
	$sqr = sqrt($disc);
	$s1 = (-$b + $sqr)/(2*$a);
	$s2 = (-$b - $sqr)/(2*$a);
}
else{
	$s1 = (-$b )/(2*$a);
	$s2 = (-$b )/(2*$a);
}

$y_0 = Formula("$y_0");
$dy_0 = Formula("$dy_0");
$eqn = ImplicitEquation("$a ypp + $b yp + $c y = $f x");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

Given the input singal [`x(t)=u(t)`], the Laplace transform of the output, [`y(t)`], of a continuous-time system is [`Y(s)=\frac{[$a*$y_0] s^2 + [$b*$y_0 + $a*$dy_0] s + [$f]}{s([$a]s^2 + [$b]s + [$c])}`]. 

*a)* Find the ordinary differential equation that describes this system as well as the corresponding initial conditions, [`y(0)`] and [`\frac{dy(0)}{dt}`] given that the numerator of transfer function has order zero.

[____________________]  
[`y(0)=`][______]  
[`\frac{dy(0)}{dt}=`][______]  

_Enter "yp" for [`\frac{dy(t)}{dt}`], "ypp" for [`\frac{d^2y(t)}{dt^2}`], [`x`] for [`x(t)`], and [`y`] for [`y(t)`]._

*b)* Is the system BIBO stable? [@ $mcYes->menu() @]*

_Part *b* will only be marked correct if all the answers to part *a* are correct._

END_PGML

ANS($eqn->cmp);
$ans_hash1 = $eqn->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$a1Correct= $ans_hash1->{score};

ANS($y_0->cmp);
$ans_hash2 = $y_0->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(2)});
$a2Correct= $ans_hash2->{score};

ANS($dy_0->cmp);
$ans_hash3 = $dy_0->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(3)});
$a3Correct= $ans_hash3->{score};

if ($s1 < 0 && $s2 < 0) {
	ANS($mcYes->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $a1Correct==1&& $a2Correct==1&& $a3Correct==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
}
else{
	ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $a1Correct==1&& $a2Correct==1&& $a3Correct==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
}

##############################################################

ENDDOCUMENT();