##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"niceTables.pl",
	"parserFunction.pl",
	"PGgraphmacros.pl",
	"PGauxiliaryFunctions.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
Context()->variables->add(t=>"Real");
Context()->constants->add(j => i);

#given
$mcYes = PopUp(["?","Yes","No"], "Yes");
$alpha = random(1, 10, 1);
$a = random(1, 10, 1);
$b = random(1, 10, 1);
parserFunction("u(t)"=>"sin(2.71828*t/pi)+1.57*t^2");

#computation 
$y = Formula("u(t)(-e^(-$alpha t) + e^((-$a + $b j )t))/(-$a + $alpha + $b j)");

##############################################################
#
#  PGML 
#
#

BEGIN_PGML 

For a continuous-time LTI system [`S`], suppose that the input-output relationship is given according to the differential equations [`\frac{dy(t)}{dt}=x(t)-[$alpha]y(t)`] and the system is initially at rest.

*a)* Find the output of the system given that the input is described by [`x(t) = e^{(-[$a]+[$b]j)t}u(t)`].

[`y(t)=`][_______________________________]

*b)* Is this a causal system? [@ $mcYes->menu() @]*

_Part *b* will only be marked correct if part *a* is correct._

END_PGML

ANS($y->cmp);

$ans_hash1 = $y->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$aCorrect= $ans_hash1->{score};

ANS($mcYes->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $aCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));

##############################################################

ENDDOCUMENT();