##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"PGauxiliaryFunctions.pl",
	"niceTables.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
Context()->variables->add(s=>"Real");
Context()->variables->add(t=>"Real");

#given
$mcYes = PopUp(["?","Yes","No"], "Yes");
$mcNo = PopUp(["?","Yes","No"], "No");

$alpha = random(2, 9,1);
$beta = random(2, 9,1);
$gamma = random(2, 9,1);

#computation 
$disc = $beta*$beta - 4 * $alpha * $gamma;

if($disc>=0){
	$sqr = sqrt($disc);
}
else{
	$sqr = sqrt(abs($disc))*i;
}

$s1 = (-$beta + $sqr)/(2*$alpha)+0*i;
$s2 = (-$beta - $sqr)/(2*$alpha)+0*i;

$hs = Formula("1/($alpha s^2 + $beta s + $gamma)");
$poles = List($s1,$s2);
$yss = Formula("1/$gamma");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

A system is represented by the ordinary differential equation

[``[$alpha] \frac{d^2 y(t)}{dt^2} + [$beta] \frac{d y(t)}{dt}+  [$gamma] y(t)= x(t)``]

where [`x(t)`] is the input and [`y(t)`] is the output.

*a)* Find the transfer function of the system, [`H(s)=\frac{Y(s)}{X(s)}`].

[`H(s)=`][____________]

*b)* Find the poles of the system. Separate your answers with commas.

[`s=`][____________]

*c)* Is the system BIBO stable?

[@ $mcNo->menu() @]*

*d)* Suppose that the input [`x(t)=u(t)`] and that the system is at rest. Find the steady state response of the system, [`y_{ss}(t)`].

[`y_{ss}(t)=`][____________]

_Part *c* will only be marked correct if part *b* is correct._

END_PGML

ANS($hs->cmp);
ANS($poles->cmp);

$ans_hash2 = $poles->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(2)});
$bCorrect= $ans_hash2->{score};

if (Re($s1) < 0 && Re($s2) < 0) {
	ANS($mcYes->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $bCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
}
else{
	ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $bCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));
}

ANS($yss->cmp);

##############################################################

ENDDOCUMENT();