##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## RESOURCES('UBC-SIG-18-009a.png', 'UBC-SIG-18-009b.png')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"niceTables.pl",
	"parserFunction.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->variables->add(t=>"Real");
Context()->variables->add(T=>"Real");
Context()->variables->add(A=>"Real");
Context()->variables->add(B=>"Real");
Context()->variables->remove("x");
parserFunction("x(t)"=>"sin(6*2.71828*t/pi)+2.87*t^6");
parserFunction("x1(t)"=>"sin(2*2.71828*t/pi)+3.14*t^4");
parserFunction("x2(t)"=>"sin(2.71828*t/pi)+1.57*t^2");

#image
$imgScaleA = .2;
$imgScaleB = .2;
$widthA = $imgScaleA*1760;
$heightA = $imgScaleA*769;
$widthB = $imgScaleB*1833;
$heightB = $imgScaleB*1495;

#given
$mcYes = PopUp(["?","Yes","No"], "Yes");
$mcNo = PopUp(["?","Yes","No"], "No");
do{$alpha = random(-9,9,1);}while($alpha == 1 or $alpha == -1 or $alpha == 0);
do{$a = random(-9,9,1);}while($a == 1 or $a == -1 or $a == 0);
do{$b = random(-9,9,1);}while($b == 1 or $b == -1 or $b == 0);

#computation 
$y1 = Formula("x(t - T - $a) + x($b - t + T)");
$y2 = Formula("x(t - T - $a) + x($b - t - T)");
$y3 = Formula("A(x1(t - $a) + x1($b - t))+B(x2(t - $a) + x2($b - t))");
$y4 = Formula("A(x1(t - $a) + x1($b - t))+B(x2(t - $a) + x2($b - t))");

##############################################################
#
#  PGML 
#
#

BEGIN_PGML

The continuous time system [`S`] is described by [``y(t)=x(t-[$a]) + x([$b]-t)``], where [`x(t)`] is the input and [`y(t)`] is the output of the system.

*a)* In the two diagrams below, find the outputs [`y_1`] and [`y_2`] in terms of the input signal [`x`]. Here, [`D_{\tau}`] shows a delay system that introduces a delay of [`\tau`].

[@ image( "UBC-SIG-18-009a.png", width=>[$widthA], height=>[$heightA], tex_size=>500) @]*

[`y_1(t)=`] [____________________________]  
[`y_2(t)=`] [____________________________] 

_Use "T" to represent [`\tau`] in your solutions._

END_PGML

ANS($y1->cmp);
ANS($y2->cmp);
$ans_hash1 = $y1->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$a1Correct= $ans_hash1->{score};
$ans_hash2 = $y2->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(2)});
$a2Correct= $ans_hash2->{score};

BEGIN_TEXT

$PAR

END_TEXT

BEGIN_PGML

*b)* Is the system [`S`] time-invariant? [@ $mcNo->menu() @]*

END_PGML

ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && ($a1Correct==1) && ($a2Correct==1)) {
			return 1;
		}else {
			return 0;
		}	
	}));

BEGIN_TEXT

$PAR

END_TEXT

BEGIN_PGML

*c)* In the two diagrams below, find the outputs [`y_3`] and [`y_4`] in terms of the input signals [`x_1`] and [`x_2`].

[@ image( "UBC-SIG-18-009b.png", width=>[$widthB], height=>[$heightB], tex_size=>500) @]*

[`y_3=`] [____________________________]  
[`y_4=`] [____________________________] 

_Use "x1" to represent [`x_1`] and "x2" to represent [`x_2`] in your solutions._

END_PGML

ANS($y3->cmp);
ANS($y4->cmp);
$ans_hash4 = $y3->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(4)});
$c1Correct= $ans_hash4->{score};
$ans_hash5 = $y4->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(5)});
$c2Correct= $ans_hash5->{score};

BEGIN_TEXT

$PAR

END_TEXT

BEGIN_PGML

*d)* Is the system [`S`] linear? [@ $mcYes->menu() @]*

*e)* Is the system [`S`] causal? [@ $mcNo->menu() @]*

*f)* Is the system [`S`] memoryless? [@ $mcNo->menu() @]*


_Part *b* will only be marked correct if part *a* is correct. Part *d* will only be marked correct if part *c* is correct. Parts *e* to *f* will only be marked correct if parts *a* and *c* are both correct._

END_PGML

ANS($mcYes->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && ($c1Correct==1) && ($c2Correct==1)) {
			return 1;
		}else {
			return 0;
		}	
	}));

ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && ($a1Correct==1) && ($a2Correct==1) && ($c1Correct==1) && ($c2Correct==1)) {
			return 1;
		}else {
			return 0;
		}	
	}));

ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && ($a1Correct==1) && ($a2Correct==1) && ($c1Correct==1) && ($c2Correct==1)) {
			return 1;
		}else {
			return 0;
		}	
	}));


##############################################################

ENDDOCUMENT();