##DESCRIPTION
##
## Created by the UBC Faculty of Applied Science.  
## Project led by Agnes d'Entremont, Jonathan Verrett, and Negar M. Harandi, 
## with support from Paul Lucina and Joseph Yan.
##              
## Contact: agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this
## project provided by UBC Vancouver students via the Teaching
## and Learning Enhancement Fund. We also gratefully acknowledge
## additional funding support from the UBC Department of
## Electrical Engineering, and the UBC Applied Science Dean’s Office.  
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people.
##
##ENDDESCRIPTION

## DBsubject(Engineering Signals and Systems)
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## RESOURCES('UBC-SIG-18-028.png')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

#image scale
$imgScale = .2;

#image aspect ratio
$width = $imgScale*1680;
$height = $imgScale*985;

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"PGauxiliaryFunctions.pl",
	"parserFunction.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
parserFunction("u(t)"=>"sin(2.71828*t/pi)+1.57*t^2");
Context()->variables->add(s=>"Real");
Context()->constants->add(j => i);
Context()->variables->add(k=>"Real");
Context()->variables->add(t=>"Real");

#given
$alpha = random(2, 10, 1);
$a = random(2, 10, 1);
$mcNo = PopUp(["?","Yes","No"], "No");

#computation 
$ans1 = Formula("$alpha t (u(t) - u(t - $a))");
$ans2 = Formula("$alpha(1/s^2 (1 - e^(-$a s)) - $a / s e^(-$a s) )");
$ans3 = Formula("$alpha*$a j / (2 pi k)");
$ans5 = Formula("$alpha*$a / 2");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

A periodic signal, [`x(t)`] is given in the figure below, where [`a=[$a]`], and [`\alpha=[$alpha]`].

>>[@ image("UBC-SIG-18-028.png", width=>[$width], height=>[$height], tex_size=>500) @]*<<

*a)* Find an equation for [`x_c(t)`], the signal that describes one cycle of [`x(t)`], in terms of the unit step function [`u(t)`].  
[`x_c(t)=`][__________________]

*b)* Find the Laplace transform, [`X_c(s)`] of the signal in part *a*.  
[`X_c(s)=`][__________________]

*c)* Calculate the Fourier Series coefficients of the signal [`x(t)`], [`X_k`] for [`k\neq 0`] using the Laplace transform from part *b*.  
[`X_k=`][______]

*d)* Is it possible to find the Fourier Series coefficient, [`X_0`] using the Laplace transform method? [@ $mcNo->menu() @]*  
*e)* Compute the Fourier Series coefficient, [`X_0`], using the integral definition.  
[`X_0=`][______]

_Part *d* will only be marked correct if part *c* is correct._

END_PGML
ANS($ans1->cmp);
ANS($ans2->cmp);

ANS($ans3->cmp);
$ans_hash3 = $ans3->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(3)});
$cCorrect= $ans_hash3->{score};

ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $cCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));

ANS($ans5->cmp);


##############################################################

ENDDOCUMENT();