# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 02/09/09

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear, constant coefficients, homogeneous (complex roots))
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(repeated_second_complex_ivp)
## Level(3)
## Static(1)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('11.11')
## Problem1('3')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('11.11')
## Problem2('3')
## KEYWORDS('calculus', 'integral', 'differential equations', 'higher derivatives')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real', A=>'Real', B=>'Real', C=>'Real' );
$showPartialCorrectAnswers = 1;

## make sure that this problem, Q19.pg and Q23.pg, if used, have
##    different equations
SRAND( $psvn );

$a = random(2,5,1);
## the roots are -$a +/- i
##    so (r + $a)^2 = -1, or r^2 + 2 $a r + $a^2 + 1 = 0
$twoa  = 2*$a;
$asq   = $a*$a;
$asqp1 = $asq + 1;

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find the (real-valued) general solution
to the differential equation.
\( y'' + $twoa y' + $asqp1 y = 0 \)

$PAR
\(y(t) = \) \{ ans_rule(35) \}
$BR
${BITALIC}(Use constants A, B, etc., for any constants
in your solution formula.)$EITALIC

END_TEXT
Context()->normalStrings;

## this is only good if students don't tweak the constants; they
##     could enter y = 3 A e^(-$a t) cos(t) + ... and it should be
##     correct, which this won't catch.  But we'll assume that they
##     won't be that contrarian...
ANS(Compute("e^(-$a*t)*(A*cos(t) + B*sin(t))")->cmp( checker=>sub {
    my ( $cor, $stu, $ansHash ) = @_;
    my $opt = $cor->substitute( A=>'C' );
    $opt = $opt->substitute( B=>'A' );
    $opt = $opt->substitute( C=>'B' );
    return $cor == $stu || $opt == $stu;
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The characteristic equation is \( r^2 + $twoa r + $asqp1 = 0\), 
so \(r=-$a\pm i\).
$PAR
Therefore \(y(t)=A e^{-$a t}\cos t + B e^{-$a t}\sin t\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
