# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 02/09/09

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear, constant coefficients, homogeneous (distinct real roots))
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('11.11')
## Problem1('11')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('11.11')
## Problem2('11')
## KEYWORDS('calculus', 'integral', 'differential equations', 'higher derivatives')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real', A=>'Real', B=>'Real', C=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,9,1);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find the (real-valued) general solution
to the differential equation.
\( z''+ $a z'=0 \)

$PAR
\( z(t) = \) \{ ans_rule(35) \}
$BR
${BITALIC}(Use constants A, B, etc., for any constants
in your solution formula.)$EITALIC

END_TEXT
Context()->normalStrings;

## this is only good if students don't tweak the constants; they
##     could enter y = 3 A e^(-$a t) cos(t) + ... and it should be
##     correct, which this won't catch.  But we'll assume that they
##     won't be that contrarian...
ANS(Compute("A + B*e^(-$a*t)")->cmp( checker=>sub {
    my ( $cor, $stu, $ansHash ) = @_;
    my $opt = $cor->substitute( A=>'C' );
    $opt = $opt->substitute( B=>'A' );
    $opt = $opt->substitute( C=>'B' );
    return $cor == $stu || $opt == $stu;
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

If we try a solution of the form \( z(t) = e^{rt} \) then
\[ r^2 + $a r=0\]
which has  solutions \(r=0\) and \(r=-$a\),
so that the general solution is
\[ y(t) = A + B e^{-$a t}.\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
