# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 02/09/09

## DBsubject(Differential equations)
## DBchapter(Introductory concepts)
## DBsection(Verification of solutions)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('11.11')
## Problem1('27')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('11.11')
## Problem2('27')
## KEYWORDS('calculus', 'integral', 'differential equations', 'higher derivatives')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real', A=>'Real', B=>'Real', C=>'Real' );
$showPartialCorrectAnswers = 1;

$r1 = random(2,5,1);
$r2 = $r1 + random(1,3,1);
## then the characteristic equation is (r - r1)(r - r2)
##    so y'' - ($r1+$r2) y' + $r1*$r2 y = 0
$r1pr2 = ($r1 + $r2);
$r1tr2 = $r1*$r2;

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

If \( y=e^{$r1 t} \) is a solution to the differential 
equation 
\[ \frac{d^2y}{dt^2} - $r1pr2 \frac{dy}{dt} + ky = 0, \]
find the value of the constant  \(k\) and the general solution
to this equation.
$PAR
\( k = \) \{ ans_rule(15) \}
$BR
\( y = \) \{ ans_rule(35) \}
$BR
${BITALIC}(Use constants A, B, etc., for any constants
in your solution formula.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS(Compute( "$r1tr2" )->cmp() );
## this is only good if students don't tweak the constants; they
##     could enter y = 3 A e^(-$a t) cos(t) + ... and it should be
##     correct, which this won't catch.  But we'll assume that they
##     won't be that contrarian...
ANS(Compute("A*e^($r1*t) + B*e^($r2*t)")->cmp( checker=>sub {
    my ( $cor, $stu, $ansHash ) = @_;
    my $opt = $cor->substitute( A=>'C' );
    $opt = $opt->substitute( B=>'A' );
    $opt = $opt->substitute( C=>'B' );
    return $cor == $stu || $opt == $stu;
} ) );

$r1sq = $r1*$r1;
$c = $r1sq + $r1*$r2;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

\( 0 = \frac{d^2}{dt^2}(e^{$r1 t}) - $r1pr2 \frac{d}{dt}(e^{$r1 t})
    + k e^{$r1 t}= $r1sq e^{$r1 t} - $c e^{$r1 t} + k e^{$r1 t} =
    e^{$r1 t}(k - $r1tr2)\).  Since \(e^{$r1 t}\neq 0\),
we must have \(k-$r1tr2=0\).  Therefore \(k=$r1tr2\).
$PAR

The characteristic equation is \(r^2 - $r1pr2 r + $r1tr2 = 0\), so
\(r=$r1\) or $r2.
Therefore \(y(t) = A e^{$r1 t} + B e^{$r2 t}\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
