# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 02/09/09

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Applications)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('11.11')
## Problem1('29')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('11.11')
## Problem2('29')
## KEYWORDS('calculus', 'integral', 'differential equations', 'higher derivatives')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Interval");
$showPartialCorrectAnswers = 1;

$c = random(3,9,1);
## for overdamped, (b^2 - 4ac) > 0; for underdamped, < 0; etc.
$fourc = 4*$c;
## this allows for negative damping, which isn't damped
# $over  = List( Interval( "(-infinity, -sqrt($fourc))" ),
# 	       Interval( "(sqrt($fourc), infinity)" ) );
$over  = List( Interval( "(sqrt($fourc), infinity)" ) );
# $under = List( Interval( "(-sqrt($fourc), sqrt($fourc))" ) );
$under = List( Interval( "(0, sqrt($fourc))" ) );
$crit  = List(
    Interval( "[sqrt($fourc), sqrt($fourc)]" )->with( checkTypes => 0 ),
);
$val = Compute( "sqrt($fourc)" );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

For the differential equation
\[ s'' + b s' + $c s = 0, \]
find the values of \(b\) that make the general solution
overdamped, underdamped, or critically damped.

$PAR
${BITALIC}(For each, give an interval or intervals for b for which
the equation is as indicated.  Thus if the the equation is
overdamped for all b in the range \(-1<b\le 1\) and \(3\le b <\infty\),
enter ${BBOLD}(-1,1], [3,infinity)$EBOLD;
if it is overdamped only for \(b=3\), enter ${BBOLD}[3,3]$EBOLD.)$EITALIC

$PAR
If the equation is overdamped, \( b \in \) \{ ans_rule(45) \}
$BR
If the equation is underdamped, \( b \in \) \{ ans_rule(45) \}
$BR
If the equation is critically damped, \( b \in \) \{ ans_rule(45) \}

END_TEXT
Context()->normalStrings;

ANS($over->cmp() );
ANS($under->cmp() );
ANS($crit->cmp( checker=>sub{
    my ( $cor, $stu, $ansHash ) = @_;
    if ( ref($cor) == ref($stu) && $cor == $stu ) {
	    return 1;
    } elsif ( ref(val) == ref($stu) && $val == $stu ) {
	return 1;
    } else {
	return 0;
    }
}) );


if ( $c == 4 || $c == 9 ) {
    $valdisp = ( $c == 4 ) ? 4 : 6;
} else {
    $valdisp = "2\sqrt{$c}";
}

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Recall that \(s''+bs'+cs=0\) is overdamped if the discriminant
\(b^2-4c>0\),
critically damped if \(b^2-4c=0\), and underdamped if
\(b^2-4c<0\).  In all cases, this is true if the motion is damped:
that is, the real part of the exponential in the solutions is negative.
Since \(b^2-4c = b^2-$fourc\), the solution is
overdamped if \(b > $valdisp\).  If \(b<-$valdisp\), the condition
\(b^2 - 4c > 0\) is satisfied, but the exponentials in the solution
will be positive and the system is therefore not damped.  Similarly,
the system will be critically damped if \(b=$valdisp\) (but not
\(b = -$valdisp\), because that will give exponential growth), and
underdamped if \(0 < b < $valdisp\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
