## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Row, column, and null spaces)
## Institution('TCNJ')
## KEYWORDS('matrix' 'inverse' 'row space')
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('4.3')
## Problem1('')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context("Matrix");

do { $e = non_zero_random(2,4,1);
     $f = non_zero_random(2,4,1);

     $a = non_zero_random(-5*$e,-3*$e,$e);
     $b = non_zero_random(-5*$f,-3*$f,$f);
     $c = non_zero_random(-2*$e,2*$e,$e);
     $d = non_zero_random(-2*$f,2*$f,$f); 
} until ($a*$d+$b*$c != 0 );  # so $M has rank 2

$M = Matrix([
  [$a, $c, 0 ],
  [$b, -$d, $d-$b ],
  [-$a*($e+1)/$e , -$c*($e+1)/$e, 0 ],
  [$b*($f+1)/$f,-$d*($f+1)/$f , ($d-$b)*($f+1)/$f]]);

$ans = Matrix([[$a, $c, 0 ]]);

Context()->texStrings;
BEGIN_TEXT

Find a non-zero vector \(V\) in the row space of the matrix $PAR
\($M\).$PAR

\(V = \)\{$ans->ans_array()\}.

END_TEXT
Context()->normalStrings;

ANS($ans->cmp(
  checker=>sub{
    my ($corr,$stu,$ansHash) = @_;
    if ($stu->isZero){
      return 0;
    }
    my $D = Matrix([
      [$a, $c, 0 ],
      [$b, -$d, $d-$b ],
      [$stu->element(1,1),$stu->element(1,2),$stu->element(1,3)]]);
    if ( $D->det->isZero ){ return 1;}
    else { 
      return 0;
    }
  }
));

ENDDOCUMENT();       # This should be the last executable line in the problem.
