##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Algebra')
## DBchapter('Fundamentals')
## DBsection('Real Numbers')
## Date('6/3/2002')
## Author('')
## Institution('')
## TitleText1('Precalculus')
## EditionText1('3')
## AuthorText1('Stewart, Redlin, Watson')
## Section1('1.1')
## Problem1('22')


    DOCUMENT();        # This should be the first executable line in the problem.

    loadMacros(
      "PGstandard.pl",
      "compoundProblem.pl",
      "Parser.pl",
      "PGunion.pl",
      "parserMultiAnswer.pl",
      "PGcourse.pl",
    );





#loadMacros("PG.pl",
#           "PGbasicmacros.pl",
#           "PGchoicemacros.pl",
#           "PGanswermacros.pl",
#           "PGauxiliaryFunctions.pl",
#           "PGgraphmacros.pl",
#           "MathObjects.pl",
#           "parserMultiAnswer.pl",
#           "parserImplicitEquation.pl");







Context()->variables->are(x=>'Real',a=>'Real',b=>'Real',c=>'Real',m=>'Real');

$auxeq = Compute("a * m^2 + b * m + c");
$D = Compute("b^2 - 4 * a * c");

$r = random(2,9,1);
$s = random(2,9,1);
$p = random(2,9,1);
$q = random(2,9,1);

$a1 = $r * $s;
$b1 = $r * $q - $p *$s;
$c1 = -1 * $p * $q;

$auxeq1 = Compute("$a1 * m^2 + $b1 * m + $c1");

$a2 = $p * $p;
$b2 = 2 * $p * $s;
$c2 = $s * $s;

$auxeq2 = Compute("$a2 * m^2 + $b2 * m + $c2");

$a3 = $q * $q;
$b3 = 2 * $q * $r;
$c3 = $r * $r + $s * $s;

$auxeq3 = Compute("$a3 * m^2 + $b3 * m + $c3");

$m1 = Compute("(-$b1-sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)");
$m2 = Compute("(-$b1+sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)");

$y11 = Compute("e^((-$b1-sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)*x)");
$y12 = Compute("e^((-$b1+sqrt(($b1)^2-4*$a1*$c1))/(2*$a1)*x)");

$a4 = list_random(3, 5, 7,);
$a42 = $a4* $a4;
$b4 = 4 * $a4; 


$y4 = Compute("e^(2*x / $a4)");
$y42 = Compute("e^(4*x / $a4)");
$P = Compute("-$b4 / $a42");
$xy4 = Compute("x*e^(2*x / $a4)");


$multipart = MultiAnswer($y4, $xy4)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


TEXT(beginproblem);
    BEGIN_TEXT

In this WeBWork assignment I want to remind you of a result which will come in useful later in the course. 

Given a second order linear homogeneous differential equation
\[ a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0\] there are times when the method for finding a solution yields only one function.
Since a second order linear homogeneous differential equation must have two linearly independent solutions, we need a method to find the second solution.
$PAR

The following result will help us find the second solution.
$BR

If \( y_1(x)\) is a solution to 
\[ a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\]
then if \( P(x)=\frac{a_1(x)}{a_2(x)}\), a second solution to  
\( a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\) is given by 
\[
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx.
\]
Furthermore, \(y_1\) and \(y_2\) are linearly independent.
$BR $BR

$BBOLD To see how this formula is derived, read Section 4.2 in our text. $EBOLD
$PAR

$BBOLD Example :$EBOLD It is easy to check (and you should) that 
\( y_1(x)=e^{\frac{2}{$a4} x}\) 
is a solution to the differential equation 
\( $a42 \frac{d^2y}{dx^2}- $b4 \frac{dy}{dx}+4y=0.\)  Let's use the formula from above to find the second solution. 

$PAR
In this case
\( y_1^2(x)=\)\{ ans_rule(15) \},
\( P(x)=\frac{a_1(x)}{a_2(x)}=\)\{ ans_rule(15) \}, 
$BR
(For theoretical purposes we drop the constant of integration on the next integral)
\( e^{-\int P(x) dx} = \)\{ ans_rule(15) \} , and 
$BR
(For theoretical purposes we drop the constant of integration on the next integral)
\[
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx=\hbox{\{ ans_rule(15) \}}.
\]

$PAR
So the general solution to 
\[ $a42 \frac{d^2y}{dx^2}- $b4 \frac{dy}{dx}+4y=0\] is
\[ y=c_1 \hbox{\{ ans_rule(15) \}} + c_2 \hbox{\{ ans_rule(15) \}}\]
where \(c_1\) and \(c_2\) are arbitrary constants.


    END_TEXT

ANS($y42->cmp());
ANS($P->cmp());
ANS($y42->cmp());
ANS($xy4->cmp());
ANS($multipart->cmp());
    





##############################################

    ENDDOCUMENT();        # This should be the last executable line in the problem.

