## DBsubject(Abstract algebra)
## DBchapter(Groups)
## DBsection(Cyclic groups)
## Institution(University of Massachusetts Amherst)
## Author(Siman Wong, Patrick O'Brien, Dan Nichols)
## Level(6)
## MO(1)
## KEYWORDS('cyclic groups, generators, subgroups')



DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "algebraMacros.pl",
  "PGcourse.pl"
);




TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

# make a modulus by picking 3 distinct small primes

@primes = (2, 3, 5, 7, 11, 13, 17, 19);

$p1 = splice( @primes, random( 0, @primes - 1, 1 ), 1 );
$p2 = splice( @primes, random( 0, @primes - 1, 1 ), 1 );
$p3 = splice( @primes, random( 0, @primes - 1, 1 ), 1 );

$n = $p1 * $p2 * $p3;


# example correct answer -- just take two prime divisors of n. That way they each have nontrivial gcd with n, but their gcd with each other is 1.
( $exampleX, $exampleY ) = ( $p1, $p2 );

# multianswer. We verify that the student's answers each generate a proper subgroup of Z/n (by making sure they have nontrivial gcd with n) and that together they generate the whole group (by making sure gcd( gcd( x, y ), n ) = 1 )
$xyAnswer = MultiAnswer( $exampleX, $exampleY )->with(
	singleResult => 1,
	checker => sub
	{
		my( $correct, $student, $self ) = @_;
		my( $x, $y ) = @$student;
		
		return ( gcd( $x, $n ) > 1 && gcd( $y, $n ) > 1 && gcd( gcd( $x, $y ), $n ) == 1 ) ? 1 : 0;
	}
);


BEGIN_TEXT


Find one pair \((x, y)\) of elements of \( \{ cyclic( $n ) \} \) such that both \(\langle x \rangle , \langle y \rangle\) are proper subgroups of \( \{ cyclic( $n ) \} \), and that \(\langle x, y \rangle = \{ cyclic( $n ) \} \). $BR $BR

\(x =\) \{ $xyAnswer->ans_rule(6) \} $BR
\(y =\) \{ $xyAnswer->ans_rule(6) \} $BR $BR

HINT: $n is a product of three primes. $BR $BR



END_TEXT




ANS( $xyAnswer->cmp() );



Context()->texStrings;

 SOLUTION(EV3(<<'END_SOLUTION'));
$SOLUTION $BR

\(\left<x\right>\) is a proper subgroup of \( \{ cyclic($n) \}\) if and only if \(\gcd($n, x) > 1\), and \(\left<x,y\right> = \{ cyclic($n) \}\) if
and only if \(\gcd( \gcd(x,y), $n) = 1\).

    END_SOLUTION

COMMENT("Complete solution included.");


Context()->normalStrings;

ENDDOCUMENT();  
