## DESCRIPTION
## Given vectors in R^2, find bases of R^2 for which their coordinate vector is as given.
##   All coefficients and solutions will be small real integers. The ordering of the vectors in the solution basis matter (i.e., an ORDERED basis is requested of the student).
##   This exercise has two parts: the solution to the first one will always be simple (for students who understand the concepts), but the second will require some calculation.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Euclidean spaces')
## DBsection('Coordinate vectors and change of basis')
## Date('March 23, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('3.1: Coordinate Systems')
## Problem1('')
## KEYWORDS('coordinate vector','change of basis')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

do {
    do {
        $P = Matrix([[random(-3,5),random(-3,5)],[random(-3,5),random(-3,5)]]);
    } until (abs($P->det) < 10.1 && abs($P->det) > 0.1);
    
    Context("Vector");
    Context()->lists->set(Vector => {open => "(", close => ")"});
    Context()->parens->set("(" => {type => "Vector", close => ")"});
    
    $p1 = Vector($P->column(1));
    $p2 = Vector($P->column(2));
    
    $a = list_random(-2,-1,1,2,3,4);
    $b = random(-2,4);
    $c = random(-2,4);
    $d = list_random(-2,-1,1,2,3,4);
    
    $v = $a*$p1 + $b*$p2;
    $w = $c*$p1 + $d*$p2;
    
    $sv = Vector($a,$b);
    $sw = Vector($c,$d);
} until ($v->isParallel($w) == 0);

$multians1 = MultiAnswer($v,$w)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if($v == $s0 && $w == $s1) {
            return 1;
        } else {
            return 0;
        }
    }
);

$multians2 = MultiAnswer($p1,$p2)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if($p1 == $s0 && $p2 == $s1) {
            return 1;
        } else {
            return 0;
        }
    }
);

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Suppose \( \mathbf{v} = $v \) and \( \mathbf{w} = $w \). Find an ordered basis \( B \) of \( \mathbb{R}^2 \) for which their coordinate vectors \( [\mathbf{v}]_B \) and \( [\mathbf{w}]_B \) are as follows:
$BR
$BR
\( [\mathbf{v}]_B = (1,0) \) and \( [\mathbf{w}]_B = (0,1) \): \( \Bigg\lbrace \)\{$multians1->ans_array(3)\}, \{$multians1->ans_array(3)\}\( \Bigg\rbrace \).
$BR
\( [\mathbf{v}]_B = $sv \) and \( [\mathbf{w}]_B = $sw \): \( \Bigg\lbrace \)\{$multians2->ans_array(3)\}, \{$multians2->ans_array(3)\}\( \Bigg\rbrace \).
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );

ENDDOCUMENT();
