# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Exponential and logarithmic expressions and functions)
## DBsection(Properties of logarithms)
## Institution(Loyola University Chicago and Fort Lewis College)
## Author(Adam Spiegler and Paul Pearson)
## MLT(SimplifyLogsGivenVariableValues)
## MLTleader(1)
## Level(4)
## MO(1)
## TitleText1('Functions Modeling Change')
## AuthorText1('Connally')
## EditionText1('3')
## Section1(4.Review)
## Problem1(46)
## KEYWORDS('logarithms','natural log','common log','log','ln','properties of logs')


DOCUMENT();       

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);
       
TEXT(beginproblem()); # standard preamble to each problem.


$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->are(u=>'Real',v=>'Real');
Context()->flags->set(useBaseTenLog=>1);


$a = random(2,4,1);
@bases = ($a,5);
@perm = shuffle(2);

$in1 = $a/5;
$ans1 = Compute("u - v");

$raise2 = random(2,3,1);
$p2 = random(-3,-1,1);
$base2 = $bases[$perm[0]];
$base2_pr = $base2**$raise2;
$in2 = $base2_pr*10**$p2;
if ($base2 == $a) {$ans2 = Compute("$p2 + $raise2 u") }
        else {$ans2 = Compute("$p2 + $raise2 v") };

$raise3 = random(2,3,1);
$p3 = random(1,3,1);
$base3 = $bases[$perm[1]];
$base3_pr = $base3**$raise3;
$in3 = $base3_pr*10**$p3;
if ($base3 == $a) {$ans3 = Compute("$p3 + $raise3 u") }
        else {$ans3 = Compute("$p3 + $raise3 v") };

$p4 = random(1,3,2);
if ($p4 == 1){
  $in4 = "10**0.5"; $in4tex = "\sqrt{10}"; 
} else {
  $in4 = "1000**0.5";
  $in4tex = "\sqrt{1000}"};
$ans4 = Compute("$p4/2");

BEGIN_TEXT
Suppose that \( \ u=\log_{10}{($a)} \ \) and \( \ v=\log_{10}{(5)} \). Find possible formulas for the following expressions in terms of \( u \) and/or \( v \). Your answers should NOT involve any \( \log{ \ } \)'s.
$BR
$BR
a) \( \log_{10}{($in1)} = \) \{ ans_rule(20) \}
$BR
$BR
b) \( \log_{10}{($in2)} = \) \{ ans_rule(20) \}
$BR
$BR
c) \( \log_{10}{($in3)} = \) \{ ans_rule(20) \}
$BR
$BR
d) \( \log_{10}{($in4tex)} = \) \{ ans_rule(20) \}
END_TEXT

#ANS(fun_cmp( $ans1,vars=>['u','v']) );
#ANS(fun_cmp( $ans2,vars=>['u','v']) );
#ANS(fun_cmp( $ans3,vars=>['u','v']) );
#ANS(fun_cmp( $ans4,vars=>['u','v']) );


$u = Real("log($a)");
$v = Real("log(5)");

ANS( $ans1->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  Value::Error("Your answer should be in terms of u and v") 
  if ($student->D('u')==0 && $student->D('v')==0 );
  return Real("log($in1)") == $student->eval(u=>$u,v=>$v);
}));

ANS( $ans2->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  Value::Error("Your answer should be in terms of u and v") 
  if ($student->D('u')==0 && $student->D('v')==0 );
  return Real("log($in2)") == $student->eval(u=>$u,v=>$v);
}));

ANS( $ans3->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  Value::Error("Your answer should be in terms of u and v") 
  if ($student->D('u')==0 && $student->D('v')==0 );
  return Real("log($in3)") == $student->eval(u=>$u,v=>$v);
}));

ANS( $ans4->cmp()); 


$disp2 = 10**$p2;
$disp3 = 10**$p3;

SOLUTION(EV3(<<'END_SOLUTION'));
Using the base ten logarithm:
$BR 
a) \( \ $in1 = \frac{$a}{5} \ \ \ \) , so \( \ \ \ \ \ \log{($in1)} = \log{($a/5)} = \log{($a)} - \log{(5)} = u - v \)
$BR $SPACE $BR
b) \( \ $in2 = ($disp2)($base2_pr) = 10^{$p2} \ \cdot $base2^{$raise2} \) , so $BR $SPACE $BR
\( \ \ \ \ \ \log{($in2)} = \log{(10^{$p2} \ \cdot $base2^{$raise2})} \) $BR
\( \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ = \log{(10^{$p2})} + \log{($base2^{$raise2})} \) $BR
\( \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ = $p2 \ \log{(10)} + $raise2 \ \log{($base2)} \) $BR
\( \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ = $ans2 \)
$BR $SPACE $BR
c) \( \ $in3 = ($disp3)($base3_pr) = 10^{$p3} \ \cdot $base3^{$raise3} \) , so $BR $SPACE $BR
\( \ \ \ \ \ \log{($in3)} = \log{(10^{$p3} \ \cdot $base3^{$raise3})} \) $BR
\( \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ = \log{(10^{$p3})} + \log{($base3^{$raise3})} \) $BR
\( \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ = $p3 \ \log{(10)} + $raise3 \ \log{($base3)} \) $BR
\( \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ \ = $ans3 \)
$BR $SPACE $BR
d) \( \ $in4 = \left( 10^{$p4} \right)^{1/2} = 10^{$p4/2} \ \ \ \) , so \( \ \ \ \ \ \log{($in4)} = \log{(10^{$p4/2})} = \frac{$p4}{2} \ \log{(10)} = \frac{$p4}{2} \)
$BR $SPACE $BR
$BR
END_SOLUTION

      
ENDDOCUMENT();   # This should be the last executable line in the problem.
