## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Inner products)
## DBsection(Applications)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('8.5')
## Problem1('')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
"PGcourse.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');

# we choose the matrix A, solution x, and a vector c orthogonal to the columns of A. 
# we set b=Ax+c, then x is the least-squares solution of Ax=b.

@a1 = (1, 1, 1, 1);

@a2 = (1, 1, -1, -1);

@a3 = (1, -1, 1, -1);

@a4 = (1, 1, 1, 1);

$x1 = non_zero_random(-5,5,1);
$x2 = random(-5,5,1);
$x3 = non_zero_random(-5,5,1);
$x4 = non_zero_random(-5,5,1);

$k = non_zero_random(-5,5,1);

@c = ($k, - $k, - $k, $k);

@s = NchooseK(4,4);

$b1 = $a1[$s[0]]*$x1 + $a2[$s[0]]*$x2 + $a3[$s[0]]*$x3 + $a4[$s[0]]*$x4 + $c[$s[0]];
$b2 = $a1[$s[1]]*$x1 + $a2[$s[1]]*$x2 + $a3[$s[1]]*$x3 + $a4[$s[1]]*$x4 + $c[$s[1]];
$b3 = $a1[$s[2]]*$x1 + $a2[$s[2]]*$x2 + $a3[$s[2]]*$x3 + $a4[$s[2]]*$x4 + $c[$s[2]];
$b4 = $a1[$s[3]]*$x1 + $a2[$s[3]]*$x2 + $a3[$s[3]]*$x3 + $a4[$s[3]]*$x4 + $c[$s[3]];

$A = Matrix([[$a1[$s[0]], $a2[$s[0]], $a3[$s[0]], $a4[$s[0]]], 
			 [$a1[$s[1]], $a2[$s[1]], $a3[$s[1]], $a4[$s[1]]], 
			 [$a1[$s[2]], $a2[$s[2]], $a3[$s[2]], $a4[$s[2]]], 
			 [$a1[$s[3]], $a2[$s[3]], $a3[$s[3]], $a4[$s[3]]]]);
$x = Matrix([[$x1],[$x2],[$x3],[$x4]]);
$b = Matrix([[$b1], [$b2], [$b3], [$b4]]);
$Atrans=$A->transpose;
$AtransA=$Atrans*$A;
$Atransb=$Atrans*$b;

Context()->texStrings;
BEGIN_TEXT
Find a least-squares solution \(\vec{x}^*\) of the system 
\[ $A \vec{x} = $b. \]
$BR
\(\vec{x}^* = \) \{ $x->ans_array(10) \}  
END_TEXT
Context()->normalStrings;

ANS( $x->cmp(
         checker=>sub {
	my ( $cor, $stu, $ans ) = @_;
        if ($AtransA*$stu==$Atransb) {
        return 1;
        }
        return 0;
        }
) );
;
ENDDOCUMENT();       # This should be the last executable line in the problem.