## DESCRIPTION
## College Algebra
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Polynomial equations and functions)
## DBsection(Zeros and multiplicities)
## Date(6/15/2017)
## Institution(Red Rocks Community College, Colorado Community College System)
## Author(Craig Faulhaber)
## MO(1)
## KEYWORDS('college algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"contextLimitedPoint.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

###########################
#  Setup
Context("Numeric");

$a=non_zero_random(2,8,1);
$b=non_zero_random(-8,-2,1);
$m=random(2,6);
do { $n = random(2,4,1); } until ( $n != $m );


$f=Formula("(x-$a)^$m*(x-$b)^$n")->reduce;


$multians = MultiAnswer($a, $m, $b, $n)->with(
  singleResult => 0,
  allowBlankAnswers => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2, $stu3, $stu4) = @{$student};
      my ( $a1, $a2, $a3, $a4) = @{$correct};
      if ( 
(   ($stu1 == $a1 && $stu2 == $a2)  && ($stu3 == $a3 && $stu4 == $a4)       ||
(    $stu1 == $a3 && $stu2 == $a4)  && ($stu3 == $a1 && $stu4 == $a2 )    ) )
{return [1,1,1,1];}  

else { return [0,0,0,0];  } 

} );



###########################
#  Main text

BEGIN_PGML
Find the zeros and give the multiplicity of each for  [`f(x)=[$f]`]. 

Note: to get the problem correct, you must get all answers correct.

+ One zero is [`x=`][____]{$multians} and its multiplicity is: [____]{$multians} [@ AnswerFormatHelp("numbers") @]*

+ The other zero is [`x=`][____]{$multians} and its multiplicity is: [____]{$multians} [@ AnswerFormatHelp("numbers") @]*

END_PGML



############################
#  Solution

#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION

COMMENT('MathObject version. Uses PGML.');

ENDDOCUMENT();