# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Prove a function is odd/even.
#
# Last update: Carl Yao 03/28/2018
#
# ENDDESCRIPTION

## DBCCSS('8.F','F-IF')
## DBsubject(Algebra)
## DBchapter(Function)
## DBsection(Domain and Range)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('function')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
  "PCCmacros.pl",
);

##############################################
TEXT(beginproblem());

Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');

Context()->operators->set(
'*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'u-' => {class => 'bizarro::UOP::minus', isCommand => 1},
);

$a = random(2,4,1);
$c = random(2,4,1);

if (random(1,4,1)==1) {
   $func = Compute("$a*x**2+$c")->reduce;
   $ans = Compute("$a*(-x)**2+$c")->reduce;
} elsif (random(1,4,1)==2) {
   $func = Compute("$a*abs(x)+$c")->reduce;
   $ans = Compute("$a*abs(-x)+$c")->reduce;
} elsif (random(1,4,1)==3) {
   $func = Compute("$a/x**2+$c")->reduce;
   $ans = Compute("$a/(-x)**2+$c")->reduce;
} else {
   $func = Compute("$a/x**4+$c*x**2")->reduce;
   $ans = Compute("$a/(-x)**4+$c*(-x)**2")->reduce;
}

$evaluator = $ans->cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      Context()->flags->set(bizarroMul=>1,bizarroNeg=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      my $OK = ($correct == $student);
      Context()->flags->set(bizarroMul=>0,bizarroNeg=>0);
      Value::Error("Your answer is equivalent to the correct answer, but does not work in this proof.") unless $OK;
      return $OK;
});

##############################################

BEGIN_PGML

The following is a proof that the function [`f(x)=[$func]`] is even. Fill in the blank to complete the proof.

Proof:

[`` f(-x) = ``][__________________]{$evaluator}

[``= [$func]``]

[``=f(x)``]

END_PGML

##############################################

BEGIN_PGML_SOLUTION

Proof:

    [``
\begin{aligned}
   f(-x) &= [$ans] \\
   &= [$func] \\
   &= f(x)
\end{aligned}
    ``]
    
Since [`f(-x)=f(x)`], [`f(x)`] is an even function.

END_PGML_SOLUTION


ENDDOCUMENT();
