##DESCRIPTION
## 
## Created by the UBC Faculty of Applied Science.  
## Project led by Agnes d'Entremont, Jonathan Verrett, and Negar M. Harandi. 
##
## Contact: agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this 
## project provided by UBC Vancouver students via the Teaching
## and Learning Enhancement Fund. We also gratefully acknowledge
## additional funding support from BCcampus, the UBC Department of
## Mechanical Engineering, and the UBC Applied Science Dean’s Office.  
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people. 
##
##ENDDESCRIPTION

## DBsubject(Dynamics)
## DBchapter(Planar kinetics of rigid bodies)
## DBsection(Conservation of energy)
## Institution('University of British Columbia')
## Author(UBC Engineering)
## Date()
## Level()
## KEYWORDS('Conservation of energy','mass','spring','energy','kinetic energy','potential energy')
## RESOURCES('UBC-DYN-18-036.png')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
## Beginner


DOCUMENT();

#image scale
$imgScale = .2;

#image aspect ratio
$width = $imgScale*610;
$height = $imgScale*1464;

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"weightedGrader.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

install_weighted_grader();

##############################################################
#
#  Setup
#
#

Context("Numeric");
#given
$g = 9.81; #m/s^2
$L = random(1,2,.1); #m
$k = random(1,5,.1); #kN/m

#computation
$placeholder = Formula("0");

##############################################################
#
#  PGML 
#
#

BEGIN_PGML

Enter your mass: [_____] [`kg`] _(Your mass should be between [`30`] [`kg`] and [`150`] [`kg`].)_

END_PGML

$ans_hash1 = $placeholder->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$studentAns = $ans_hash1->{original_student_ans};  

WEIGHTED_ANS($placeholder->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if ($studentAns <= 150 and $studentAns >= 30) {
			return 1;
		}else {
			return 0;
		}	
	}), 0);

if ($studentAns > 150) {
	$M = 150;
}
elsif ($studentAns < 30) {
	$M = 30;
}
else{
	$M = $studentAns;
}

$Xmax = ($M*$g + sqrt(($M*$g)**2+4*500*$k*$M*$g*$L))/1000/$k;
$Vmax = sqrt(2*$g*$L+$M/$k/1000*$g**2);

BEGIN_TEXT
$PAR
END_TEXT

BEGIN_PGML

Consider a case that you jump on a trampoline from a certain distance. Assume that your movement can be represented by a mass and spring model that is shown in the figure below. Use the following assumptions to find the maximum deflection of the spring and the maximum velocity of the mass.  
	[`\cdot`] Your initial velocity is zero  
	[`\cdot`] Your initial distance from the trampoline is [`[$L]`] [`m`].  
	[`\cdot`] The represented stiffness of the spring is [`[$k]`] [`kN/m`].  
	[`\cdot`] [`M`] is your given mass.  


[@ image("UBC-DYN-18-036.png", width=>[$width], height=>[$height]) @]*

[`x_{max} =`][____] [`m`]  
[`v_{max} =`][____] [`m/s`]

END_PGML

WEIGHTED_ANS( num_cmp($Xmax, reltol => .5), 50 );
WEIGHTED_ANS( num_cmp($Vmax, reltol => 1), 50 );


#.5% .01%

##############################################################

ENDDOCUMENT();