# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Interpretation and applications)
## Institution(Loyola University Chicago)
## Author(Adam Spiegler)
## Level(3)
## MO(1)
## TitleText1('Functions Modeling Change')
## TitleText2('Functions Modeling Change')
## TitleText3('Functions Modeling Change')
## AuthorText1('Connally')
## AuthorText2('Connally')
## AuthorText3('Connally')
## EditionText1('3')
## EditionText2('4')
## EditionText3('5')
## Section1(2.Review)
## Section2('2.4')
## Section3('2.5')
## Problem1(34)
## Problem2('35')
## Problem3('53')
## KEYWORDS('functions','domain','range','inverse','function notation')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "extraAnswerEvaluators.pl",
  "MathObjects.pl",
  "parserFunction.pl",
  "parserMultiAnswer.pl",
  "parserNumberWithUnits.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context("Numeric")->variables->are( r=>'Real' );
parserFunction("f(r)" => "sin(r)");

$showPartialCorrectAnswers = 1;

$num = random(2,5,1);
$A = $num**2;
$ans5 = $num/sqrt(pi);
$r = round($ans5*10000)/10000;

$ans1 = "pi r^2";

$a1 = MultiAnswer( Formula("f(r)"), Compute("$ans1") )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $a ) = @_;
	if ( ( $c->[0] == $s->[0] && $c->[1] == $s->[1] ) ||
	     ( $c->[1] == $s->[0] && $c->[0] == $s->[1] ) ) {
	    return [1,1];
	} elsif ( $c->[0] == $s->[0] || $c->[1] == $s->[0] ) {
	    return [1,0];
	} elsif ( $c->[0] == $s->[1] || $c->[1] == $s->[1] ) {
	    return [0,1];
	} else {
	    return [0,0];
	}
    } );

$ans3 = "pi (r+1)^2";
$ans4 = "pi r^2 + 1";

$ml = new_match_list();
$ml -> qa ( "\( f(r+1) \)",
   "The area of a circle whose radius is 1 cm. more than \(r\).",
            "\( f(r) + 1 \)",
   "One more square cm. than the area of a circle with radius \(r\).");

$ml ->extra("One cm. more than the area of circle with radius \(r\) square cm.",
            "The radius of a circle whose area is one square cm. more than the area of a circle with radius \(r\).");
$ml ->choose(2);
$ml->choose_extra(2);
$ml->makeLast("None of the above.");

Context()->texStrings;
BEGIN_TEXT

The area, in square centimeters, of a circle whose radius is \(r\) cm
is given by \( A = \pi r^2 \).
$PAR
(a) Write this formula using function notation, where \( f \) is the name
of the function.
$BR
\{ $a1->ans_rule(5) \} = \{ $a1->ans_rule(15) \}

$PAR
(b) Evaluate \( f(0) = \) \{ ans_rule(20) \}

$PAR
(c) Evaluate and simplify the formula for \( f(r+1) = \) \{ ans_rule(25) \}

$PAR
(d) Evaluate and simplify the formula for \( f(r)+1 = \) \{ ans_rule(25) \}

$PAR
(e) Evaluate, including units in your answer:
\( f^{-1}($A) = \) \{ ans_rule(30) \} \{ AnswerFormatHelp("units") \}

$PAR
(f) For each of the mathematical expressions below, match one of the
statements A - E below which best explains its meaning in practical terms. $BR
\{ $ml -> print_q \}
\{$ml -> print_a \}


END_TEXT
Context()->normalStrings;

# ANS(fun_cmp( $ans1, vars=>'r' ) );
# ANS(num_cmp( 0 ) );
# ANS(fun_cmp( $ans3, vars=>'r' ) );
# ANS(fun_cmp( $ans4, vars=>'r' ) );
# ANS(num_cmp( $ans5, units=>'cm' ) );

ANS( $a1->cmp() );
ANS( Compute("0")->cmp() );
ANS( Compute( $ans3 )->cmp() );
ANS( Compute( $ans4 )->cmp() );
ANS( NumberWithUnits( $ans5, 'cm' )->cmp() );
ANS(str_cmp( $ml->ra_correct_ans )   ) ;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
(a) \( f(r) = \pi r^2 \)
$PAR
(b) \( f(0) = \pi 0^2 = 0 \)
$PAR
(c) \( f(r+1) = \pi (r+1)^2 \)
$PAR
(d) \( f(r) + 1 = \pi r^2 + 1 \)
$PAR
(e) \( f^{-1}($A) \) means we need to find the radius (in cm),
\(r\), of a circle which has area $A square cm.  Therefore we set
\( f(r) = $A \) and solve for \(r\):
\[\begin{aligned}
 \pi r^2 &= $A \\
r^2 &= \frac{$A}{\pi} \\
r &= \sqrt{ \frac{$A}{\pi} } \\
r &= \frac{$num}{\sqrt{\pi}} \approx $r.
\end{aligned}\]
Since this number is the corresponding radius (in cm) of a circle which
has area $A square cm, the we have \( f^{-1}($A) \approx $r \) cm.
$PAR
(f) The expression \( f(r + 1) \ \) is the area of a circle whose radius is
1 cm more than \(r\). $BR
The expression \( f(r) + 1 \ \) is is the area of a circle of radius \(r\),
plus 1 square centimeter more.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
