# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Graphs)
## Institution(Loyola University Chicago)
## Author(Adam Spiegler)
## MLT(graphicEval)
## MO(1)
## TitleText1('Functions Modeling Change')
## TitleText2('Functions Modeling Change')
## TitleText3('Functions Modeling Change')
## AuthorText1('Connally')
## AuthorText2('Connally')
## AuthorText3('Connally')
## EditionText1('3')
## EditionText2('4')
## EditionText3('5')
## Section1(2.Review)
## Section2(2.Review)
## Section3('2.Review')
## Problem1(40)
## Problem2(63)
## Problem3('66')
## KEYWORDS('functions','domain','range','inverse','function notation')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "extraAnswerEvaluators.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context("Numeric");


$showPartialCorrectAnswers = 1;

$y2 = random(-5,-2,1);
$y2_sq = $y2**2;
$num = random(3,4,1);
if ($num == -$y2) { $y2 = -5 };
$ans2 = -$num;
$ans3 = $num;
$num_sq = $num**2;
$k = $y2 + $num_sq;
$ans1 = $k - $y2**2;

$qr = MultiAnswer( Compute($ans2),Compute($ans3) )->with(
    singleResult=>0,
    checker=>sub{
	my ( $c, $s, $a ) = @_;
	if ( ( $c->[0] == $s->[0] && $c->[1] == $s->[1] ) ||
	     ( $c->[1] == $s->[0] && $c->[0] == $s->[1] ) ) {
	    return [1,1];
	} elsif ( $c->[0] == $s->[0] || $c->[1] == $s->[0] ) {
	    return [1,0];
	} elsif ( $c->[0] == $s->[1] || $c->[1] == $s->[1] ) {
	    return [0,1];
	} else {
	    return [0,0];
	}
    } );

$c = -$y2;
$c_sq = $c**2;
$h[0] = $k - $c_sq;
$c_sub = $c-1;
$csub_sq = $c_sub**2;
$h[1] = $k - $csub_sq;
$ans4 = $h[0] - $h[1];

Context()->texStrings;
BEGIN_TEXT

Let \( h(x)= $k- x^2 \) .
$PAR
(a) Find a point \(P\) on the graph of \( h(x) \) whose \(x\)-coordinate
is \($y2\).
$BR
\(P = ($y2, \) \{ ans_rule(8) \} \( ) \)
$PAR
(b) Find two points \(Q\) and \(R\) on the graph of \( h(x) \) whose
\(y\)-coordinates are \($y2\). $BR
\( Q = ( \) \{ $qr->ans_rule(8) \} \()\) $BR
\( R = ( \) \{ $qr->ans_rule(8) \} \()\) $BR
$PAR
(c) Let \(p = $c\). Calculate the value of \( h(p) - h(p-1) \). $BR
\( h(p) - h(p-1) = \) \{ ans_rule(10) \}

END_TEXT
Context()->normalStrings;

# ANS(num_cmp( $ans1 ) );
# ANS(num_cmp( $ans2 ) );
# ANS(num_cmp( $ans3 ) );
# ANS(num_cmp( $ans4 ) );

ANS( Compute($ans1)->cmp() );
ANS( $qr->cmp() );
ANS( Compute($ans4)->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
(a) To find a point on the graph \( h(x) \) with an \(x\)-coordinate
of $y2, we substitute $y2 for \(x\) in the formula for \( h(x) \).  We obtain
\( h($y2) = $k - ($y2)^2 = $k - $y2_sq = $ans1 \).
Thus, we have the point \(P = ($y2 ,$ans1)\).
$PAR

(b) To find these points, we want to find all the values of \(x\) for
which \( h(x)=$y2 \).  We have
\[\begin{aligned}
$k - x^2 &= $y2 \\
-x^2 &= - $num_sq \\
x^2 &= $num_sq \\
x = \pm $num.
\end{aligned}\]
Thus, the points \(Q = (-$num, $y2)\) and \(R = ($num, $y2)\) both have a
\(y\)-coordinate of $y2.
$PAR

(c) For \( p=$c \), \( h(p)-h(p-1)=h($c)-h($c_sub) \).  Now
\( h($c)=$k-($c)^2=$k-$c_sq=$h[0] \), while
\( h($c_sub)=$k-($c_sub)^2=$k - $csub_sq = $h[1] \), thus,
\( h(p)-h(p-1) = h($c)-h($c_sub)= $h[0]-$h[1]= $ans4 \).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
