# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('vector field', 'multivariable', 'calculus')
## Tagged by glr 06/11/10


## DBsubject('Calculus')
## DBchapter('Vector Calculus')
## DBsection('Vector Fields')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('17.3')
## Problem1('25')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')
## HHSection1('Vector Fields')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
);

Context('Vector2D');
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);

$showPartialCorrectAnswers = 1;

$a = 1;
$b = -1;
$c = -1;
$v = 'x';
$t1 = Compute("x + $a*y")->reduce();
$t2 = Compute("$b + $c*$v^2")->reduce();

$v = Vector( $t1, $t2 );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Give an example of a vector field \( \vec F (x, y)  \) in 2-space
with the properties that \( \vec F \) is perpendicular to
\( \vec G = $v \) at every point.

$PAR
\( \vec F(x,y) = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( $v->cmp( checker=>sub { my ( $c, $s, $an ) = @_;
    return $c.$s == Compute(0) && norm($s) != 0; } ) );


Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

If \(\vec F(x, y) = f(x,y)($h)\) where \(f(x, y)\) is any function,
then \(\vec F\cdot \vec G=0\), which shows that \(\vec F\) is
perpendicular to \(\vec G\).
$PAR
For example \(\vec F(x, y) = (y+\cos x  )($h)\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
