##DESCRIPTION
##  Trigonometry problem: Find two angles (one positive, one negative) coterminal with a	given angle
##ENDDESCRIPTION


## DBsubject(Trigonometry)
## DBchapter(Geometric and algebraic foundations for trigonometry)
## DBsection(Reference angles (using coterminal angles))
## Institution(CUNY City Tech)
## Author(K. Andrew Parker)
## Level(2)
## Language(en)
## MO(1)
## KEYWORDS('trigonometry', 'angle', 'angles', 'angle measure', 'coterminal', 'radian', 'radians')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->flags->set(
  reduceConstants=>0,
  reduceConstantFunctions=>0,
  tolType=>'absolute',
  tolerance=>0.0005
);

$den = list_random(3,4,6);
do { $num = random(1,2*$den-1,1); } until ( gcf($num,$den) == 1 );
$num = -$num;

$angle = ($num==-1)? Formula("(-pi)/$den") :Formula("($num*pi)/$den");

##############################################################
#
#  Text
#
#

BEGIN_PGML

Find two angles (one positive and one negative) that are coterminal with [`` [$angle] ``].

[__________] is a positive angle that is coterminal with [`` [$angle] ``].

[__________] is a negative angle that is coterminal with [`` [$angle] ``].

* You'll have an easier time with this problem if you use fractions and _pi_ instead of decimal approximations.

* If you use decimal approximations, your answer must be correct to at least 3 decimal places.

END_PGML

##############################################################
#
#  Answers
#
#

ANS($angle->cmp(checker=>sub{
  my ($correct, $student, $ans) = @_;
  Value->Error("Your answer for this part must be a positive angle.") if ($student < 0);
  Value->Error("Yes, technically, angles are coterminal with themselves. $check <BR> Please enter a coterminal angle that is _different_ than the one provided.") if ( abs($student-$correct) < 0.05);
  my $diff = Real("$student - $correct");
  my $f = floor($diff/(2*pi));
  my $c = ceil($diff/(2*pi));
  my $error = min( abs($f*2*pi - $diff), abs($c*2*pi - $diff));
  return 0 if ($error > 0.05);
  my $correct = $error > 0.0005 ? 0 : 1;
  Value->Error("You're close! Try entering your answer as a fraction of 'pi'.") unless $correct;
  return $correct;
}));

ANS($angle->cmp(checker=>sub{
  my ($correct, $student, $ans) = @_;
  Value->Error("Your answer for this part must be a negative angle.") if ($student > 0);
  Value->Error("Yes, technically, angles are coterminal with themselves. <BR> Please enter a coterminal angle that is _different_ than the one provided.") unless (abs($student-$correct) > 0.05);
  my $diff = Real("$student - $correct");
  my $f = floor($diff/(2*pi));
  my $c = ceil($diff/(2*pi));
  my $error = min( abs($f*2*pi - $diff), abs($c*2*pi - $diff));
  return 0 if ($error > 0.05);
  my $correct = ($error > 0.0005) ? 0 : 1;
  Value->Error("You're close! Try entering your answer as a fraction of 'pi'.") unless $correct;
  return $correct;
}));

BEGIN_PGML_HINT

Coterminal angles are angles that differ by a full rotation.

In degrees, a full rotation is [`360^{\circ}`].  
In radians, a full rotation is [`2\pi`]. 

Try adding or subtracting full rotations to find angles that are coterminal with [`[$angle]`].

END_PGML_HINT
BEGIN_PGML_SOLUTION

There are many positive coterminal angles with [`[$angle]`].

*To find positive coterminal angles, we add full rotations.*  
You can think of this like spinning around in a full circle.  
You end up facing exactly the same direction you were before you spun around.

_Add one full rotation:_  
[` [$angle] + 2\pi \rightarrow [$angle] + \frac{[@2*$den@]\pi}{[$den]} \rightarrow \frac{[@2*$den+$num@]\pi}{[$den]} `]  

_Add two full rotations:_  
[` [$angle] + 4\pi \rightarrow [$angle] + \frac{[@4*$den@]\pi}{[$den]} \rightarrow \frac{[@4*$den+$num@]\pi}{[$den]} `]  

You could even add more full rotations, there are infinitely many positive coterminal angles with [`[$angle]`].  
In other words, it doesn't matter how many times you spin around,  
_as long as you end up facing the same direction you were before you started._

*To find negative coterminal angles, we instead subtract full rotations.*  
Here, we're just spinning circles in the opposite direction.

_Subtract one full rotation:_  
[` [$angle] - 2\pi \rightarrow [$angle] - \frac{[@2*$den@]\pi}{[$den]} \rightarrow \frac{[@$num-2*$den@]\pi}{[$den]} `]  

_Subtract two full rotations:_  
[` [$angle] - 4\pi \rightarrow [$angle] - \frac{[@4*$den@]\pi}{[$den]} \rightarrow \frac{[@$num-4*$den@]\pi}{[$den]} `]  

Again, we could subtract any number of full rotations.  
There are an infinite number of negative coterminal angles as well.

*In general:*
Let [`k`] be the number of times to spin around.  
[`k`] can be positive or negative (depending on which direction you want to spin).  
[` [$angle] + k(2\pi) \rightarrow [$angle] + \frac{[@2*$den@]k\pi}{[$den]} \rightarrow \frac{([$num]+[@2*$den@]k)\pi}{[$den]} `]  

Do you see how this relates to our answers above?

What happens when [`k`] is 1? Or -1?  
What about when [`k`] is 2? Or -2?

END_PGML_SOLUTION

COMMENT("Funded by US DoE Title V: Opening Gateways grant.");

ENDDOCUMENT();        
