# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/03/09


## DBsubject('ZZZ-Inserted Text')
## DBchapter('ZZZ-Inserted Text')
## DBsection('ZZZ-Inserted Text')
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## Static(1)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('13.1')
## Problem1('17')
## KEYWORDS('vectors', 'coordinate vector')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')

Context("Vector2D");
$showPartialCorrectAnswers = 1;

$v = 3*i + 2*j;
# $v = Compute("<3,2>");

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

\($v\)
$PAR
Parallel: \{ ans_rule(25) \}
$BR
Perpendicular: \{ ans_rule(25) \}

END_TEXT
Context()->normalStrings;

ANS($v->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student->isParallel($correct);
} ) );

ANS($v->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student.$correct == 0;
} ) );


;
ENDDOCUMENT();
