##DESCRIPTION
#
# File Created: 3/27/2009
# Last Modified: 3/27/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear independence)
## Date(3/27/2009)
## Institution(Alfred University)
## Author(Juan Marin)
## Level(2)
## TitleText1('Differential Equations with Boundary-Value Problems')
## AuthorText1('Zill and Cullen')
## EditionText1('7')
## Section1('4.7')
## Problem1('')


DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraders.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);
 
#install_problem_grader(~~&full_partial_grader);

# No partial credit on this problem, so we say:
install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 0;
$seed = random(1,4,1);
SRAND($seed);


## Declare variables
Context()->variables->are(x=>'Real',t=>'Real');


$a = random(2,5,1);
$b = random(6,9,1);
$c = 2 * $a;
$d = $b * $a;
$y1 = Compute("$a");
$y2 = Compute("cos($b x)^2+sin($b x)^2");

$dy1 = $y1->D('x');
$dy2 = $y2->D('x');
$w = Compute("$y1 * $dy2 - $y2 * $dy1")->reduce;


#$f = Formula("(x**2 + x)(x**2 - x)")->reduce;
#$df = $f->D('x');



# Make a new select list
$tf = new_pop_up_select_list();
$tf->ra_pop_up_list(["Yes" => "Yes","No" => "No"]);
# $tf now "contains" the select list object.  
# See matchinglistexample.pg for details


$tf -> qa (
q!Is \(\{y_1,y_2\}\) a fundamental set for 
\($b x^5\sin($a x)D^2 y -$a x^2\cos($b x) Dy =0\)
on \((0,\pi/$c)\)?!,  # each entry has to end with a comma
"No",
);   # remember every statement has to end with a semi-colon.

# Choose two of the question and answer pairs at random.
$tf ->choose(1);  # Using choose(3) would choose all three questions, but the order of the questions and answers would be 
                  # scrambled.


# Now print the text using $ml->print_q for the questions and $ml->print_a to print the answers.





TEXT(beginproblem());

Context()->texStrings;
BEGIN_TEXT
$PAR
It can be shown that \(y_1=$a\)  and \(y_2=\cos^2($b x)+\sin^2($b x)\)  are 
solutions to the differential equation 
\($b x^5\sin($a x)D^2 y -$a x^2\cos($b x) Dy =0\)
on \((0,\pi/$c)\).

$PAR 
What does the  Wronskian of \(y_1, y_2\) equal on \((0,\pi/$c)\)? 
$PAR
\(W(y_1,y_2)\) =  \{ ans_rule()\} on \((0,\pi/$c)\).

$BR
\{ $tf-> print_q \}
 
END_TEXT


Context()->normalStrings;


 ANS( $w->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return abs($student) == abs($correct);
  } ) );



#ANS(fun_cmp("c*$w", vars=>"x",params=>['c']));

#ANS(fun_cmp("abs($h*e^(-$c x))", vars=>"x"));
#ANS(fun_cmp("c*e^(-$b x)", vars=>"x",params => ['c']));


# Enter the correct answers to be checked against the answers to the students.

ANS(str_cmp( $tf->ra_correct_ans )   ) ;


ENDDOCUMENT();
