##DESCRIPTION
#
# File Created: 3/27/2009
# Last Modified: 3/27/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear independence)
## Date(3/27/2009)
## Institution(Alfred University)
## Author(Juan Marin)
## Level(2)
## TitleText1('Differential Equations with Boundary-Value Problems')
## AuthorText1('Zill and Cullen')
## EditionText1('7')
## Section1('4.7')
## Problem1('')


DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraders.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);
 
#install_problem_grader(~~&full_partial_grader);

# No partial credit on this problem, so we say:
install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 0;
$seed = random(1,4,1);
SRAND($seed);


## Declare variables
Context()->variables->are(x=>'Real',t=>'Real');


$u = random(2,5,1);
$v = random(6,9,1);

$b = $v+$u+1;
$c = $v*$u;
$e = random(2,9,1);
$be = $b*$e;
$ce = $c*$e;
$p = 1-$u;
$q = 1-$v;


$f = Compute("x^{$p}");
$g = Compute("x^{$q}");
$h = Compute("$e");

$df = $f->D('x');
$dg = $g->D('x');
$dh = $h->D('x');

$ddf = $df->D('x');
$ddg = $dg->D('x');
$ddh = $dh->D('x');


$w = Compute(" $f * ($dg * $ddh - $dh * $ddg) 
              -$g * ($df * $ddh - $dh * $ddf)
              +$h * ($df * $ddg - $dg * $ddf)
           ")->reduce;




# Make a new select list
$tf = new_pop_up_select_list();
$tf->ra_pop_up_list(["Yes" => "Yes","No" => "No"]);
# $tf now "contains" the select list object.  
# See matchinglistexample.pg for details


$tf -> qa (
q!Is \(\{y_1,y_2,y_3\}\) a fundamental set for \(x^2D^3y+$b x D^2y+ $c Dy =0\)
on \((0,\infty)\)?!,  # each entry has to end with a comma
"yes",
);   # remember every statement has to end with a semi-colon.


# Choose two of the question and answer pairs at random.
$tf ->choose(1);  # Using choose(3) would choose all three questions, but the order of the questions and answers would be 
                  # scrambled.


# Now print the text using $ml->print_q for the questions and $ml->print_a to print the answers.





TEXT(beginproblem());

Context()->texStrings;
BEGIN_TEXT
$PAR
It can be shown that \(y_1=x^{$p},\ y_2=x^{$q}\) and \(y_3=$e\) are 
solutions to the differential equation 
\(x^2D^3y+$b x D^2y+ $c Dy =0\)
on \((0,\infty)\).
$PAR 
What does the  Wronskian of \(y_1, y_2, y_3\) equal? 
$PAR

\(W(y_1, y_2, y_3 )\) =  \{ ans_rule()\} on \((0,\infty)\).

$BR
\{ $tf-> print_q \}
 
END_TEXT


Context()->normalStrings;


 ANS( $w->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return abs($student) == abs($correct);
  } ) );



#ANS(fun_cmp("c*$w", vars=>"x",params=>['c']));

#ANS(fun_cmp("abs($h*e^(-$c x))", vars=>"x"));
#ANS(fun_cmp("c*e^(-$b x)", vars=>"x",params => ['c']));


# Enter the correct answers to be checked against the answers to the students.

ANS(str_cmp( $tf->ra_correct_ans )   ) ;


ENDDOCUMENT();
