#DESCRIPTION
#  Series: Ratio test
#ENDDESCRIPTION

## KEYWORDS('series','ratio test')

## Author('Paul Pearson')
## Institution('Fort Lewis College')
## DBsubject('Calculus')
## DBchapter('Infinite Sequences and Series')
## DBsection('Absolute Convergence and the Ratio and Root Tests')           
## AuthorText1('')
## TitleText1('Calculus')
## EditionText1('')
## Section1('')
## Problem1('')


###############################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"unionTables.pl",
"parserPopUp.pl",
"answerHints.pl",
"PGgraders.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());


###############################
#  Setup

Context("Numeric")->variables->are(n=>"Real");

$a = random(2,9,1);
$b = random(2,9,1);
$c = random(5,20,1);

# TeX
$series = "\sum_{n=$c}^{\infty} \frac{$a^n}{($b n)^2}";
$fraction = "\lim_{n \to \infty} \left| \frac{ a_{n+1} }{ a_{n} } \right| = \lim_{n \to \infty}";


$num1 = Formula("$a*n^2");
$den1 = Formula("n^2+2n+1"); 

$num2 = Formula("$a");
$den2 = Formula("1+2/n+1/(n^2)");


$multians = MultiAnswer($num1, $den1)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $stu1, $stu2 ) = @{$student};

      if ( ( Formula("$b^2 * $a^(n+1) * n^2") == $stu1 && 
             Formula("$b^2 * $a^n * (n+1)^2") == $stu2 ) ||
           ( Formula("$a^(n+1) * n^2") == $stu1 &&
             Formula("$a^n * (n+1)^2") == $stu2 ) ||
           ( Formula("$b^2 * $a * n^2") == $stu1 &&
             Formula("$b^2 * (n+1)^2") == $stu2 )
         )
      {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];
      } elsif (($num1 == $stu1 && $den2 == $stu2) || 
               ($num2 == $stu1 && $den1 == $stu2)) {
          $ansHash->setMessage(1,"Check your algebra");
          $ansHash->setMessage(2,"Check your algebra");
          return [0,0];
      }
    
      
      if (($num1 == $stu1 && $den1 == $stu2) ||
          ($num2 == $stu1 && $den2 == $stu2) ) {
          return [1,1];
      } elsif ($num1 == $stu1 || $num2 == $stu1) {
          return [1,0];
      } elsif ($den1 == $stu2 || $den2 == $stu2) {
          return [0,1];
      } else {
          return [0,0];
      }
  }
);



$limit = Formula("$a");
$popup = PopUp(["Choose","Converges","Diverges","Inconclusive"],"Diverges");



# 
#  Display the fraction and answer blanks nicely
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
  $showfraction =
  "\[ $fraction ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $showfraction =
  ColumnTable(
  "\( \displaystyle $fraction \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}
Context()->normalStrings;


#####################################
#  Main text

BEGIN_TEXT
Use the ratio test to determine whether 
\( \displaystyle $series \) converges or diverges.
$BR
$BR
(a) Find the ratio of successive terms.  Write your answer as
a fully simplified fraction.
For \( n \geq $c \), 
$showfraction
$BR
(b) Evaluate the limit in the previous part.  Enter \( \infty \) 
as ${BITALIC}infinity${EITALIC} and \( -\infty \) 
as ${BITALIC}-infinity.${EITALIC}  If the limit does not exist, 
enter ${BITALIC}DNE.${EITALIC}
$BR
\( \displaystyle \lim_{n\to\infty} \left| \frac{a_{n+1}}{a_{n}} \right| \, \) = 
\{ ans_rule(20) \}
$BR
$BR
(c) By the ratio test, does the series
converge, diverge, or is the test inconclusive?
\{ $popup->menu() \}
END_TEXT


##########################################
#  Answer Evaluation

$showPartialCorrectAnswers=1;
install_problem_grader(~~&custom_problem_grader_fluid);
$ENV{'grader_numright'} = [2,4];
$ENV{'grader_scores'} = [0.4,1];
$ENV{'grader_message'} = "You can earn " .
"40% partial credit for 2 - 3 correct answers.";


ANS( $multians->cmp() );

ANS( $limit->cmp() );
ANS( $popup->cmp() );



COMMENT('MathObject version');
ENDDOCUMENT();



ANS( $num1->cmp( checker=>~~&mychecknum )->withPostFilter(AnswerHints( 
  [Formula("$b^2 * $a^(n+1) * n^2"),Formula("$a^(n+1) * n^2")] 
  => "Simplify your answer further (factor and cancel)",
)) 
);

sub mycheckden {
  my ($correct, $student, $ansHash) = @_;
  return $student == $correct || $student == $den2;
}

ANS( $den1->cmp( checker=>~~&mycheckden )
->withPostFilter(AnswerHints( 
  => "Simplify your answer further (factor and cancel)",
))
);

#ANS( $num->cmp() 
#->withPostFilter(AnswerHints( 
#  [Formula("$b^2 * $a^(n+1) * n^2"),Formula("$a^(n+1) * n^2")] 
#  => "Simplify your answer further (factor and cancel)",
#))
#);

#ANS( $den1->cmp() 
#->withPostFilter(AnswerHints( 
#  [Formula("$b^2 * $a^n * (n+1)^2"),Formula("$a^n * (n+1)^2")] 
#  => "Simplify your answer further (factor and cancel)",
#))
#);

