# DESCRIPTION
# Antiderivatives
# ENDDESCRIPTION

## Tagged by glr 02/08/09

## DBsubject(Calculus - single variable)
## DBchapter(Integrals)
## DBsection(Antiderivatives)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('Calculus')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('antiderivative')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


###################################
#  Setup

Context("Numeric")->variables->are(x=>"Real",C=>'Real');
Context()->variables->set(x=>{limits=>[1,4]});

$a = random(2,9,1);

$n1 = random(2,5,1);
do { $n2 = random(2,5,1); } until ($n2 != $n1);
do { $n3 = random(2,9,1); } until ($n3 != $n2 && $n3 != $n1);
do { $n4 = random(2,9,1); } until ( gcd($n3,$n4) == 1 );

$a1 = $a * $n1;
$m1 = $n1 + 1;
$m2 = $n2 - 1;
$m3 = $n3 - $n4;


$F = Formula("$a / x^($n1) + x^($n2) - x^($n3/$n4)");
$f = $F->D('x');

$Ftex = "\frac{$a}{x^{$n1}} + x^{$n2} - \sqrt[$n4]{x^{$n3}}";
$ftex = "- \frac{$a1}{x^{$m1}} + $n2 x^{$m2}  - \frac{$n3}{$n4} x^{$m3/$n4}";



$answer1 = $f;
$answer2 = $F+C;

$ans_eval1 = $answer1->cmp();
$ans_eval2 = $answer2->cmp( checker => sub {
      my ( $correct, $student, $self ) = @_;
      if ($self->{_filter_name} ne 'produce_equivalence_message') {
        my $context = Context()->copy;
        $context->flags->set(no_parameters=>0);
        $context->variables->add('C0'=>'Parameter');
        $student = Formula($context,$student);
        $correct = Formula($context,"C0 + ($correct)");
      }
      return $correct == $student;
    } )->withPostFilter(AnswerHints( 
  $F => "Is your answer the most general antiderivative?",
));




###################################
#  Main text part 1
Context()->texStrings;
BEGIN_TEXT
${BBOLD}Part 1 of 2:${EBOLD}
$BR
$BR
This is a multi-part problem.  If \( \displaystyle F(x) = $Ftex \), find \( F^{\prime}(x) \).
$BR
$BR
\( F^{\prime}(x) = \) \{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;


######################################
#  Begin do not edit

ANS($ans_eval1 );

$ans_hash1 = $ans_eval1->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});

if (1 == $ans_hash1->{score} || 
#$envir{permissionLevel}>=10 ||
$envir{permissionLevel} eq ''
) {


#  End do not edit
#######################################

Context()->texStrings;
BEGIN_TEXT
$PAR
$HR
${BBOLD}Part 2 of 2:${EBOLD} 
$BR
$BR
Find the most general antiderivative or indefinite integral.  
Use a capital letter \( C \) to denote an arbitrary constant. 
$BR
$BR
\( \displaystyle \int \left( $ftex \right) \, dx  = \)
\{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;

ANS( $ans_eval2 );

}

COMMENT("This is a multi-part problem in which the next part is revealed only after the previous part is correct.");

ENDDOCUMENT();
