# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Applications - Newton's law of cooling)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(FO-AN-DE)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('11.5')
## Section2('11.5')
## Section3('11.5')
## Section4('11.5')
## Problem1('11')
## Problem2('15')
## Problem3('29')
## Problem4('42')
## KEYWORDS('calculus', 'integral', 'differential equations', 'modeling', 'exponentials')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->flags->set(reduceConstants=>0);
Context()->variables->are(H=>"Real", t=>"Real", k=>"Real", C=>"Real");
$showPartialCorrectAnswers = 1;

$t0 = random(64,70,1);
$de = Compute("k*($t0 - H)");
$h0 = random(32,50,1);
$h1 = $h0 + random(5,10,1);
$n  = random(3,5,1);

$gensol = Compute("C*e^(-k*t) + $t0");

$tlater = Compute("($h0-$t0)*e^($n*ln(($h1-$t0)/($h0-$t0))) + $t0");

## this grades for k of either +/- sign
$mp = MultiAnswer($de, $gensol)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $stuEqn, $stuSol ) = @{$student};
        my ( $eqn, $sol ) = @{$correct};

	my $eqn2 = Compute("-k*($t0 - H)");
	my $sol1a = Compute("-C*e^(-k*t) + $t0");
	my $sol2 = Compute("C*e^(k*t) + $t0");
	my $sol2a = Compute("-C*e^(k*t) + $t0");

	if ( ( $eqn == $stuEqn &&
	       ($sol == $stuSol || $sol1a == $stuSol) ) ||
	     ( $eqn2 == $stuEqn &&
	       ($sol2 == $stuSol || $sol2a == $stuSol) ) ) {
	    return [1,1];
	} elsif ( ( $eqn == $stuEqn &&
		    ($sol2 == $stuSol || $sol2a == $stuSol) ) ||
		  ( $eqn2 == $stuEqn &&
		    ($sol == $stuSol || $sol1a == $stuSol) ) ) {
	    $self->setMessage(2,"Your solution doesn't match your " .
			      "differential equation.");
	    return [1,0];
	} else {
	    if ( $eqn == $stuEqn || $eqn2 == $stuEqn ) {
		return [1,0];
	    } elsif ( $sol == $stuSol || $sol2 == $stuSol ||
		      $sol1a == $stuSol || $sol2a == $stuSol ) {
		return [0,1];
	    } else {
		return [0,0];
	    }
	}
    }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider a \($h0^\circ \)F object placed in \( $t0^\circ \)F room.
$PAR
${BBOLD}(a)$EBOLD
Write a differential equation for \( H \), the temperature of the
object at time \( t \), using \(k>0\) for the constant of proportionality,
and write your equation in terms of \(H\), \(k\), and \(t\).
$BR
\( H' = \) \{ ans_rule(35) \}

$PAR
${BBOLD}(b)$EBOLD
Give the general solution for your differential equation.  Simplify your
solution in terms of an unspecified constant \(C\), which
appears as the coefficient of an exponential term, and the growth factor \( k \).
$BR
\( H = \) \{ ans_rule(35) \}
$BR
${BITALIC}(Your answer may involve the constant of proportionality 
\(k\).)$EITALIC

$PAR
${BBOLD}(c)$EBOLD
The temperature of the object is \( $h0^\circ \)F initially, and
\( $h1^\circ \)F one hour later. Find the temperature of the object
after \( $n \) hours.
$BR
\( H($n) = \) \{ ans_rule(15) \} degrees F

END_TEXT
Context()->normalStrings;

ANS($de->cmp() );
ANS($gensol->cmp( checker=>sub {
     my ( $cor, $stu, $ans ) = @_;
     $cor2 = $cor->substitute( C=>"-1*C" );
     return ($cor == $stu) || ($cor2 == $stu); } ) );
    
ANS($tlater->cmp() );

$c0 = $t0 - $h0;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
Letting \(k\) be the constant of proportionality,
by Newton's Law of Cooling, we have
\[
\frac{dH}{dt} = k ($t0 - H).
\]
Alternately, if our \(k\) is negative, we could have
\(H' = k(H - $t0)\).

$PAR
${BBOLD}(b)$EBOLD
We solve this equation by separating variables
\[
\int \frac{dH}{$t0 - H} = \int k \;dt, \quad\hbox{so}\quad
-\ln \left |$t0 - H \right| = k t + A.
\]
Exponentiating both sides and letting \(C = \pm e^{A}\), 
\[ $t0 - H = \pm e^{A - k t} = C e^{-k t}, \quad\hbox{and}\quad
H = $t0 - C e^{-k t}.\]
If we take \(C = \mp e^{A}\), we could write this as
\(H = $t0 + C e^{-k t}\).

$PAR
${BBOLD}(c)$EBOLD
We are told that \(H = $h0\) when \(t = 0\); this tells us that
\[ $h0 = $t0 - C e^{-k(0)}, \]
so that \(C = $c0\).
$PAR
Knowing \(C\), we can solve for \(k\) using the fact that
\(H = $h1\) when \(t = 1\):
\[ $h1 = $t0 - $c0 e^{-k (1)},\quad\hbox{so}\quad
   k = -\ln\left(\frac{$h1-$t0}{$c0}\right). \]

Then, to find \(H\) when \(t=$n\), we plug this \(k\) and \(t=$n\)
into our equation:
\[
H($n) = $t0 - $c0 e^{k($n)} \approx $tlater^\circ\;\hbox{F}.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
