##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Algebra')
## DBchapter('Fundamentals')
## DBsection('Real Numbers')
## Date('6/3/2002')
## Author('')
## Institution('')
## TitleText1('Precalculus')
## EditionText1('3')
## AuthorText1('Stewart, Redlin, Watson')
## Section1('1.1')
## Problem1('22')


    DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "compoundProblem.pl",
  "Parser.pl",
  "PGunion.pl",
  "parserMultiAnswer.pl",
  "parserFormulaUpToConstant.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);

Context()->variables->are(x=>'Real',a=>'Real',b=>'Real',c=>'Real',m=>'Real');

$a = list_random(2, 3, 5, 7,);
$b = $a;
$u = list_random(11, 13, 17, 19);
$w = $u;


$A = $u * $w;
$B = -($b * $u + $a * $w) + $A ;
$C = $a * $b;

$p = $a/$u;
$q = $b/$w;

$y1 =random(-10,10);
$Dy1 =random(-10,10);


$Yh1 = Compute("x^($p)")->reduce;
$Yh12 = Compute("x^(2*$p)")->reduce;
$Yh2 = Formula("x^($p)ln(x)")->reduce;

$P = Compute("$B/($A*x)");
$g = Compute("x^(-$B/$A)");

$multipart = MultiAnswer($Yh1, $Yh2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

$C1 = $y1;
$C2 = $Dy1 - $p * $C1;


$yp = Formula("$C1 x^($p)+$C2 x^($p)ln(x)")->reduce;


TEXT(beginproblem);
    BEGIN_TEXT

In this WeBWork assignment I want to remind you of a result which will come in useful later in the course. 

Given a second order linear homogeneous differential equation
\[ a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0\] there are times when the method for finding a solution yields only one function.
Since a second order linear homogeneous differential equation must have two linearly independent solutions, we need a method to find the second solution.
$PAR

The following result will help us find the second solution.
$BR

If \( y_1(x)\) is a solution to 
\[ a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\]
then if \( P(x)=\frac{a_1(x)}{a_2(x)}\), a second solution to  
\( a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0\) is given by 
\[
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx.
\]
Furthermore, \(y_1\) and \(y_2\) are linearly independent.
$BR $BR

$BBOLD To see how this formula is derived, read Section 4.2 in our text. $EBOLD
$PAR

$BBOLD Problem 1 :$EBOLD It is easy to check (and you should) that 
\(y_1(x)=x^{\frac{$a}{$u}}\) 
is a solution to the differential equation 
\( $A x^2 y''+ $B x y' + $C y = 0\) for \( x \in (0,\infty)\).  
Let's use the formula from above to find the second solution. 

$PAR
In this case
\( y_1^2(x)=\)\{ ans_rule(15) \},
\( P(x)=\frac{a_1(x)}{a_2(x)}=\)\{ ans_rule(15) \}, 
$BR
(For theoretical purposes we drop the constant of integration on the next integral)
\( e^{-\int P(x) dx} = \)\{ ans_rule(15) \} , and 
$BR
(For theoretical purposes we drop the constant of integration on the next integral)
$BR
$BCENTER
\(y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx=\)\{ ans_rule(15) \}.
$ECENTER
$PAR
So the general solution to 
\( $A x^2 y''+ $B x y' + $C y = 0\) for \( x \in (0,\infty)\) is
$BR
$BCENTER
\(y(x)=c_1 \) \{ ans_rule(15) \} \( + c_2 \)\{ ans_rule(15) \}
$ECENTER
$BR
where \(c_1\) and \(c_2\) are arbitrary constants.
$PAR

$BBOLD Problem 2: $EBOLD
The unique solution to 
the initial value problem   
\[ $A x^2 y''+$B x y' + $C y =0,\ \ y(1) = $y1,\ \ y'(1) = $Dy1.\] 
is the function
\(y(x)=\)\{ans_rule(35)\} 
for \( x \in (0,\infty).\)
END_TEXT

ANS($Yh12->cmp());
ANS($P->cmp());
ANS($g->cmp());
#ANS($Yh1->cmp());
ANS($Yh2->cmp());
ANS($multipart->cmp());
ANS( $yp->cmp() );





##############################################

ENDDOCUMENT();  # This should be the last executable line in the problem.
























##DESCRIPTION
#
# File Created: 3/27/2009
# Last Modified: 3/27/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##KEYWORDS()
##

## DBsubject('')
## DBchapter('')
## DBsection('')
## Date('3/27/2009')
## Author('Juan Marin')
## Institution('Alfred University')
## TitleText1('Differential Equations with Boundary-Value Peoblems')
## EditionText1('7')
## AuthorText1('Zill and Cullen')
## Section1('4.7-CauchyEulerEquations')
## Problem1('')


DOCUMENT();

# Load whatever macros you need for the problem
 
## Show partial correct answers
$showPartialCorrectAnswers = 1;
## Display the problem information
TEXT(beginproblem());




  Context("Interval");
  # to allow open or closed intervals, uncomment
  #    the following line.
  # Context()->flags->set(ignoreEndpointTypes=>1);


$int2 = Compute("(0,inf)");

## Declare variables
Context()->variables->are(x=>'Real',t=>'Real');


$a = list_random(2, 3, 5, 7,);
$b = $a;
$u = list_random(11, 13, 17, 19);
$w = $u;


$A = $u * $w;
$B = -($b * $u + $a * $w) + $A ;
$C = $a * $b;

$p = $a/$u;
$q = $b/$w;

$y1 =random(-10,10);
$Dy1 =random(-10,10);


$Yh1 = Formula("x^($p)")->reduce;
$Yh2 = Formula("x^($p)ln(x)")->reduce;



$multipart = MultiAnswer($Yh1, $Yh2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );



$C1 = $y1;
$C2 = $Dy1 - $p * $C1;


$yp = Formula("$C1 x^($p)+$C2 x^($p)ln(x)")->reduce;




Context()->texStrings;
BEGIN_TEXT
$BBOLD (1) $EBOLD
 Let \(C_1\) and \(C_2\) be arbitrary constants.  The general solution to 
the homogeneous differential equation  
\( $A x^2 y''+ $B x y' + $C y = 0\) is the function
\(y(x)=C_1\ y_1(x)+C_2\ y_2(x)=C_1\) \{$multipart->ans_rule(15)\} \(+C_2\)\{$multipart->ans_rule(15)\}.
$PAR
$BBOLD (2) $EBOLD
The unique solution to 
the initial value problem   
\[ $A x^2 y''+$B x y' + $C y =0,\ \ y(1) = $y1,\ \ y'(1) = $Dy1.\] 
is the function
\(y(x)=\)\{ans_rule(35)\} 
for \( x \in\) \{ ans_rule(15) \}.
$BR
$BBOLD
For \( -\infty \) type -inf and for \( \infty \) type inf.
$EBOLD

END_TEXT


Context()->normalStrings;
ANS($multipart->cmp());
ANS( $yp->cmp() );
ANS( $int2->cmp() );

ENDDOCUMENT();
