##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('PDE', 'laplace', 'equation', 'dirichlet')

## DBsubject('Calculus')
## DBchapter('Partial Differential Equations')
## DBsection('Laplace's Equation')
## Date('7/29/2011')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('12.5')
## Problem1('')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "Alfredmacros.pl",
  "answerHints.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("X"=>"Real","X''"=>"Real","Y"=>"Real","Y''"=>"Real","lambda"=>"Real","a"=>"Real","n"=>"Real");
Context()->variables->set(lambda => {TeX => '{\lambda}'});
Context()->functions->undefine("acos","acosh");
Context()->{error}{msg}{"Function 'acosh' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cosh function, you need to put the multiplication symbol between the a and the cosh";
Context()->{error}{msg}{"Function 'acos' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cos function, you need to put the multiplication symbol between the a and the cos";
$DEcontext = Context()->copy;

Context("Numeric");
Context()->variables->add(a=>'Real',b=>'Real',n => ['Real', limits=>[1,20], 
    resolution=>1]);
Context()->functions->undefine("cos");
$coefficient = Context()->copy;

Context("Numeric");
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add(y=>"Real", a=>'Real',b=>'Real',n=>'Real',"A"=>"Real","B"=>"Real");


$PDE = Formula($DEcontext,"X''*Y+X*Y''");
## Tranlate PDE to ODE with IC

$Xdiff1 = Formula($DEcontext,"X''/X");
$Ydiff1 = Formula($DEcontext,"Y''/-Y");

$Xdiff2 = Formula($DEcontext,"X''+lambda*X");


$left = Formula($DEcontext,"0");
$right = Formula($DEcontext,"a");
$homogeneousleft = Formula($DEcontext,"0");
$homogeneousright = Formula($DEcontext,"0");

$multians = MultiAnswer($left,$homogeneousleft,$right,$homogeneousright)->with(
  singleResult => 0,
  checker => sub 
    {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2,$stu3,$stu4 ) = @{$student};
      my ( $c1, $c2, $c3, $c4 ) = @{$correct};
      if ( ($c1 == $stu1) && ($c3 == $stu3) ) 
         {if ( ($c2 == $stu2) && ($c4 == $stu4) )
            {return [1,1,1,1];}
          elsif ($c2 == $stu2) {return [1,1,1,0]}
          elsif ($c4 == $stu4) {return [1,0,1,1]}
          else {return [1,0,1,0]};
         }
      elsif ( ($c3 == $stu1) && ($c1 == $stu3) )
         {if ( ($c4 == $stu2) && ($c2 == $stu4) )
             {return [1,1,1,1];}
         elsif ($c4 == $stu2) {return [1,1,1,0]}
         elsif ($c2 == $stu4) {return [1,0,1,1]}
         else {return [1,0,1,0]};
         }
      else {return [0,0,0,0];}
  }
);

#### Solving for X
$cn = Formula("cos(n*pi*x/a)");
$sn = Formula("sin(n*pi*x/a)");
$lambda = Formula("(n*pi/a)^2");

###### Solve for Y
$Ydiff2 = Formula($DEcontext,"Y''-(n*pi/a)^2*Y");
$Y = Formula("A*cosh(n*pi*y/a)+B*sinh(n*pi*y/a)");


$f = Formula("x");
$A0 = Formula("a/(2*b)");

$An = Formula($coefficient,"2/(sinh(n*pi*b/a))a*(-1)^(n+1)/(n*pi)");
$An->{test_points}=[[1,1,1,1],[1,1,-1,1],[1,1,-2,1],[1,1,2,1],[1,1,-3,1],[1,1,3,1],[1,1,4,1]];

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
$BBOLD Note: $EBOLD Use the prime notation for derivatives, so the derivative of \(X\) is written as \(X^\prime\). Do NOT use \(X^\prime(x)\). The variable \(\lambda\) is typed lambda.
$BR
$BR
Solve Laplace's equation with homogeneous vertical boundary conditions:
$BR
$BR
$BCENTER
\(\displaystyle
\frac{\partial^2 u}{\partial x^2} +\frac{\partial^2 u}{\partial y^2}=0, 0< x < a, 0 < y < b\)
$BR
$BR
\(\displaystyle
u(0,y)=0, \hskip 10pt u(a,y)=0, 0 < y < b\)
$BR
$BR
\(\displaystyle
u(x,0) = f(x), \hskip 10pt u(x,b) = g(x),0<x<a\)
$ECENTER
$BR
$BR
We try for a separable solution \(u(x,y) = X(x)Y(y)\), plugging \(XY\) into the PDE for \(u\) we get 
$BR
$BCENTER
\{ans_rule\} \(=0\)
$ECENTER
$BR
The PDE can be separated into an ODE in X and an ODE in Y(placing constants in the DE for Y):
$BR
$BCENTER
\{$Xdiff1->ans_rule\} = \{$Ydiff1->ans_rule\} = \(-\lambda\)
$ECENTER
$BR
$BR
The boundary conditions in the PDE translate into initial conditions for the differential equation in X:
$BR

Since these differential equations are independent of each other, they can be separated 
$BR
DE in X: \{$Xdiff2->ans_rule\} \( = 0\)
$BR IC's in X: 
$BR
$BCENTER
\(X(\)\{$multians->ans_rule(1)\}\() = \)\{$multians->ans_rule(1)\}
$BR
\(X(\)\{$multians->ans_rule(1)\}\() = \)\{$multians->ans_rule(1)\}
$ECENTER
$BR
$BR
We have solved this S-L problem before, without the unknown constant, the solution is \(X(x) = \) \{$sn->ans_rule\}.
$BR
Therefore \(\lambda = \) \{$lambda->ans_rule\} and the differential equation in Y becomes:
$BR 
DE in Y: \{$Ydiff2->ans_rule\} \( = 0\)
$BR
Using A and B as the constants, the solution for Y is Y(y) = \{$Y->ans_rule(30)\}.
$BR
So the PDE admits a series solution of the form
$BR
\(u(x,y) = \displaystyle\sum\limits_{n=1}^\infty\left( A_n \cosh\left(\frac{n\pi}{a}y\right) +B_n \sinh\left(\frac{n\pi}{a}y\right)\right) \sin\left( \frac{n\pi}{a}x\right) \). 
$BR
Now we can apply the lateral boundary conditions:
$BR
\(u(x,0) = \displaystyle\sum\limits_{n=1}^\infty A_n\sin\left( \frac{n\pi}{a}x\right) =f(x)\)
$BR
\(u(x,b) = \displaystyle\sum\limits_{n=1}^\infty\left( A_n \cosh\left(\frac{n\pi}{a}b\right) +B_n \sinh\left(\frac{n\pi}{a}b\right)\right) \sin\left( \frac{n\pi}{a}x\right)=g(x) \)
$BR
The first series we recognize as a Fourier series on [0,a], so that 
$BR
$BCENTER
\{   BeginTable(center=>0).
     Row(['\(A_n = \)',ans_rule(1).$BR.'----------'.$BR.ans_rule(1),ans_rule(1).$BR.'\(\displaystyle\int\)'.$BR.ans_rule(1),'\(f(x)\)',ans_rule,'dx',
     ],separation=>2).
   EndTable();
\}
$ECENTER
$BR
$BR
The second series is also a Fourier series with coefficient 
$BR
$BCENTER
\{   BeginTable(center=>0).
     Row(['\(\displaystyle A_n \cosh\left(\frac{n\pi}{a}b\right) +B_n \sinh\left(\frac{n\pi}{a}b\right)= \)',ans_rule(1).$BR.'----------'.$BR.ans_rule(1),ans_rule(1).$BR.'\(\displaystyle\int\)'.$BR.ans_rule(1),'\(g(x)\)',ans_rule,'dx',
     ],separation=>2).
   EndTable();
\}
$ECENTER
$BR
Which we solve for \(B_n\)
$BR
$BCENTER
\{   BeginTable(center=>0).
     Row(['\(B_n= \)','\(\hskip 40pt 1\)'.$BR.'-----------------------------------'.$BR.ans_rule(),'\(\Biggl(\)',ans_rule(1).$BR.'----------'.$BR.ans_rule(1),ans_rule(1).$BR.'\(\displaystyle\int\)'.$BR.ans_rule(1),'\(g(x)\)',ans_rule,'\(dx - A_n\)',ans_rule,'\(\Biggr)\)',
     ],separation=>2).
   EndTable();
\}
$ECENTER
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
#####



# Translate to ODE
#
ANS($PDE->cmp);
ANS($Xdiff1->cmp);
ANS($Ydiff1->cmp);
ANS($Xdiff2->cmp);
ANS( $multians->cmp() ); 

## Solve the equation in X
ANS($sn->cmp() );
ANS($lambda->cmp() );

#### Solve for Y
ANS($Ydiff2->cmp);
ANS($Y->cmp);

#### 
# Compute A_n
ANS(Compute("2")->cmp); 
ANS(Compute("a")->cmp);
ANS(Compute("a")->cmp);
ANS(Compute("0")->cmp);
ANS($sn->cmp);

#Compute B_n
ANS(Compute("2")->cmp); 
ANS(Compute("a")->cmp);
ANS(Compute("a")->cmp);
ANS(Compute("0")->cmp);
ANS($sn->cmp);

# Solve for B_n
ANS(Compute("sinh(n*pi*b/a)")->cmp());
ANS(Compute("2")->cmp); 
ANS(Compute("a")->cmp);
ANS(Compute("a")->cmp);
ANS(Compute("0")->cmp);
ANS($sn->cmp);
ANS(Compute("cosh(n*pi*b/a)")->cmp());

ENDDOCUMENT();        
