## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','matrix','multiplication')
## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter('Abstract vector spaces')
## DBsection('Vector space axioms')
## Date('July 2013')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextArbitraryString.pl",
  "parserFunction.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context('ArbitraryString');

$a = "\alpha";
$b = "\beta";

$x = "\mathbf{u}";
$y = "\mathbf{v}";
$z = "\mathbf{w}";

$zero = "\mathbf{0}";

Context()->texStrings;
BEGIN_TEXT
A ${BBOLD}vector space${EBOLD} over \( \mathbb{R} \) is a set
\( V \) of objects (called vectors) together with two operations,
addition and multiplication by scalars (real numbers), that satisfy 
the following \( 10 \) axioms.  The axioms must hold for ${BBOLD}all${EBOLD} vectors
\( $x, $y, $z \) in \( V \) and for ${BBOLD}all${EBOLD} scalars \( $a, $b \) in \( \mathbb{R} \).
$BR
$BR
\{ BeginList('OL') \}
$ITEM (Closed under addition:) The sum of \( $x \) and \( $y \), denoted \( $x + $y \), is in \( V \).
$ITEMSEP
$ITEM (Closed under scalar multiplication:) The scalar multiple of \( $x \) by \( $a \), denoted \( $a $x \), is in \( V \).
$ITEMSEP
$ITEM (Addition is commutative:) \( $x + $y = $y + $x \).
$ITEMSEP
$ITEM (Addition is associative:) \( ($x + $y) + $z = $x + ($y + $z) \).
$ITEMSEP
$ITEM (A zero vector exists:) There exists a vector \( $zero \) in \( V \) such that \( $x + $zero = $x \).
$ITEMSEP
$ITEM (Additive inverses exist:) For each \( $x \) in \( V \), there exists a \( $y \) in \( V \) such that \( $x + $y = $zero \).  (We write \( $y = -$x \).)
$ITEMSEP
$ITEM (Scaling by \( 1 \) is the identity:) \( 1 $x = $x \).
$ITEMSEP
$ITEM (Scalar multiplication is associative): \( $a ($b $x) = ($a $b) $x \).
$ITEMSEP
$ITEM (Scalar multiplication distributes over vector addition:) \( $a ($x + $y) = $a $x + $a $y \).
$ITEMSEP
$ITEM (Scalar addition is distributive:) \( ($a + $b) $x = $a $x + $b $x \).
$ITEMSEP
\{ EndList('OL') \}

$HR
$BR

Let \( V \) be the set of functions \( f : \mathbb{R} \to \mathbb{R} \).  
For any two functions \( f, g \) in \( V \), define the sum \( f+g \) to be the 
function given by \( (f+g)(x) = f(x)+g(x) \) for all real numbers \( x \).  
For any real number \( c \) and any function \( f \) in \( V \), define
scalar multiplication \( c f \) by \( (cf)(x) = c f(x) \) for all real numbers \( x \).
$BR
$BR
Answer the following questions as partial verification that \( V \) is a vector space.
$BR
$BR
(Addition is commutative:)  Let \( f \) and \( g \) be any vectors in \( V \).  Then
\( f(x) + g(x) = \) \{ ans_rule(15) \} for all real numbers \( x \) since adding 
the real numbers \( f(x) \) and \( g(x) \) is a commutative operation.
$BR
$BR
END_TEXT
Context()->normalStrings;

ANS(Compute("g(x)+f(x)")->cmp(checker=>sub {
  my ($correct,$student,$ans) = @_;
  $correct = $correct->value; # get perl string from String object
  $student = $student->value; # ditto
  $correct =~ s/ //g; # chomp all whitespace
  $student =~ s/ //g; # ditto
  if ($student eq $correct) { return 1; } else { return 0; } 
}));

###############################################

Context('Numeric')->variables->add('c'=>'Real');
parserFunction(f => "sin(x+1/pi)-e^2");
parserFunction(g => "cos(x+1/e)-pi^2");

Context()->texStrings;
BEGIN_TEXT
(A zero vector exists:) The zero vector in \( V \) is the function \( f \) given by
\( f(x) = \) \{ ans_rule(5) \} for all \( x \).
$BR
$BR
(Additive inverses exist:) The additive inverse of the function \( f \) in \( V \) is 
a function \( g \) that satisfies \( f(x) + g(x) = 0 \) for all real numbers \( x \).
The additive inverse of \( f \) is the function \( g(x) = \) \{ ans_rule(10) \} for all \( x \).
$BR
$BR
END_TEXT
Context()->normalStrings;

ANS( Compute("0")->cmp);

ANS(Compute("-f(x)")->cmp);

#################################################

Context('ArbitraryString');

Context()->texStrings;
BEGIN_TEXT
(Scalar multiplication distributes over vector addition:)
If \( c \) is any real number and \( f \) and \( g \) are two vectors in \( V \), then
\( c(f+g)(x) = c( f(x) + g(x) ) = \) \{ ans_rule(20) \}
END_TEXT
Context()->normalStrings;

ANS(Compute("cf(x)+cg(x)")->cmp(checker=>sub {
  my ($correct,$student,$ans) = @_;
  $correct = $correct->value; # get perl string from String object
  $student = $student->value; # ditto
  $correct1 = "cf(x)+cg(x)";
  $correct2 = "c*f(x)+cg(x)";
  $correct3 = "cf(x)+c*g(x)";
  $correct4 = "c*f(x)+c*g(x)";
  $student =~ s/ //g; # ditto
  if (($student eq $correct1) or ($student eq $correct2) or ($student eq $correct3) or ($student eq $correct4)) { return 1; } else { return 0; } 
}));


COMMENT('MathObject version.');
ENDDOCUMENT();
