## DBsubject(Abstract algebra)
## DBchapter(Fields and polynomials)
## DBsection(Fields)
## Institution(University of Massachusetts Amherst)
## Author(Daniel Nichols and Siman Wong)
## Level(6)
## MO(1)
## KEYWORDS('fields, zero-divisors')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "answerHints.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "parserRadioButtons.pl",
  "PGcourse.pl"
);

Context( 'Numeric' );

TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

$showHint = 2;

# pick random d not equal to 0 or 1
$d = random( -30, 36, 1 ) until $d && $d != 1;

# determine if d is square
$dSquare = $d > 0 && int( sqrt( $d ) ) == sqrt( $d );

# make radio buttons for asking whether or not the ring is a field
$isFieldAnswer = $dSquare ? 'No' : 'Yes';
$isFieldRadio = RadioButtons( [ 'Yes', 'No' ], $isFieldAnswer, last => [ ] );

$sCorrect = $dSquare ? Compute( sqrt( $d ) ) : Compute( 1 / ( 1 - $d ) );
$tCorrect = $dSquare ? Compute( 1 ) : Compute( 1 / ( 1 - $d ) );


$stAnswer = MultiAnswer( $sCorrect, $tCorrect )->with(
	singleResult => 1,
	checker => sub
	{
		my( $correct, $student, $self ) = @_;
		my( $s, $t ) = @$student;
		
		# if A( d ) is not a field, check to make sure that s,t satisfy s^2 - dt^2 = 0
		if( $dSquare )
		{
			return ( $s || $t ) && ( $s ** 2 - $d * $t ** 2 == 0 ) ? 1 : 0;
		}
		# otherwise, compare student answer to the known inverse ( s, t ) = ( 1 / ( 1 - d ), 1 / ( 1 - d ) )
		else
		{
			return ( $s == 1 / ( 1 - $d ) && $t == 1 / ( 1 - $d ) ) ? 1 : 0;
		}

	}
);


TEXT(EV3(<<'EOT'));

Denote by \( I \) the \( 2 \times 2 \) identity matrix, and let $BR $BR

\{ mbox( [ '\( D = \)', display_matrix([ [ '\( 0 \)', '\( $d \)' ], [ '\( 1 \)', '\( 0 \)' ] ]) ] ) \} 

Denote by \( A( $d ) \) the set of matrices of the form \( sI + tD \), where \( s, t \) are rational numbers. It is a fact that \( A( $d ) \) is a commutative ring with a multiplicative identity under matrix addition and multiplication. $BR $BR

(a) Is \( A( $d ) \) a field? $BR
\{ $isFieldRadio->buttons \} $BR $BR

(b) If \( A( $d ) \) is a field, enter an ordered pair of rational numbers \( ( s, t ) \) such that \( sI + tD \) is the multiplicative inverse of \( I - D \). If \( A( $d ) \) is not a field, enter an ordered pair of non-zero rational numbers \( ( s, t ) \) such that \( sI + tD \) is a zero-divisor. $BR
\{ $stAnswer->ans_rule( 3 ) \} \( I + \) \{ $stAnswer->ans_rule( 3 ) \} \( D \) $BR

EOT

BEGIN_HINT
A commutative ring with multiplicative identity is a field if and only if every non-zero element has a multiplicative inverse. To solve this problem, express in terms of \( 2 \times 2 \) matrices what it means for an element of \( A( $d ) \) to have an inverse. This leads to a system of 4 equations, one for each entry of the \( 2 \times 2 \) matrices. Now, try to determine when this system has a unique solution.
END_HINT


ANS( $isFieldRadio->cmp );

ANS( $stAnswer->cmp );


# solution is split into two cases, depending on whether or not d is square

if( $dSquare )
{
SOLUTION(EV3(<<'END_SOLUTION'));
$SOLUTION $BR

A commutative ring with multiplicative identity is a field if and only if every non-zero element has a multiplicative inverse, in which case the inverse must be unique. A non-zero element in \( A( $d ) \) is given by \( sI + tD \) where at least one of \( s, t \) is non-zero. To see if it has a multiplicative inverse, we try to solve the matrix equation $BR $BR

\[ (*)\quad  I = ( sI + tD ) ( uI + wD ). \] $BR

Expand the right side and we get $BR $BR

\[ I = ( su )I + ( sw + tu )D + ( tw )D^2. \] $BR

But \( D^2 \) is the diagonal matrix with \( $d \) on the diagonals, so this simplifies to $BR $BR

\[ I = ( su + $d tw )I + ( sw + tu )D. \] $BR

Of course the matrix \( D \) cannot be a scalar multiple of \( I \), so that means $BR $BR

\[ \begin{aligned}
	su +& $d tw &= 1 \\
	sw +& tu  &= 0
\end{aligned} \] $BR

Note that we can view this as a system of linear equations in the variables \( u, w \), in which case we can rewrite this using matrix notation as $BR $BR

\{ mbox( [
	'\( (**)\quad \)',
	display_matrix( [ [ '\( s \)', '\( t $d \)' ], [ '\( t \)', '\( s \)' ] ] ),
	display_matrix( [ [ '\( u \)' ], [ '\( w \)' ] ] ),
	'\( = \)',
	display_matrix( [ [ '\( 1 \)' ], [ '\( 0 \)' ] ] )
] ) \}

By general facts from linear algebra, we know that this system has a unique solution if and only if the determinant of the matrix above is not zero, i.e. if and only if \( s^2 - $d t^2 \neq 0 \), or equivalently, \( $d \neq ( s/t )^2 \) for some rational numbers \( s, t \). $BR $BR

Since \( $d = \{ sqrt( $d ) \} ^2 \), the system does not have a unique solution and so \( A( $d ) \) is not a field. The zero-divisors are described by the following analog of \( (*) \): $BR $BR

\[ 0 = ( sI + tD ) ( uI + wD ). \] $BR

Expand this matrix equation as before and we end up with the matrix equation $BR $BR

\{ mbox( [ display_matrix( [ [ '\( s \)', '\( t $d \)' ], [ '\( t \)', '\( s \)' ] ] ), display_matrix( [ [ '\( u \)' ], [ '\( w \)' ] ] ), '\( = \)', display_matrix( [ [ '\( 0 \)' ], [ '\( 0 \)' ] ] ) ] ) \}

Setting for example \( s = t = 1 \), we find that one solution is \( w = 1, u = ( \sqrt{$d} + 1 )/2 \). $BR $BR

Suppose we were dealing with \( A( d ) \) where \( d \) is NOT a perfect square. Can you find the inverse of \( I - D \) in this ring? $BR

END_SOLUTION
}

else
{
SOLUTION(EV3(<<'END_SOLUTION'));
$SOLUTION $BR

A commutative ring with multiplicative identity is a field if and only if every non-zero element has a multiplicative inverse, in which case the inverse must be unique. A non-zero element in \( A( $d ) \) is given by \( sI + tD \) where at least one of \( s, t \) is non-zero. To see if it has a multiplicative inverse, we try to solve the matrix equation $BR $BR

\[ (*)\quad  I = ( sI + tD ) ( uI + wD ). \] $BR

Expand the right side and we get $BR $BR

\[ I = ( su )I + ( sw + tu )D + ( tw )D^2. \] $BR

But \( D^2 \) is the diagonal matrix with \( $d \) on the diagonals, so this simplifies to $BR $BR

\[ I = ( su + $d tw )I + ( sw + tu )D. \] $BR

Of course the matrix \( D \) cannot be a scalar multiple of \( I \), so that means $BR $BR

\[ \begin{aligned}
	su +& $d tw &= 1 \\
	sw +& tu  &= 0
\end{aligned} \] $BR

Note that we can view this as a system of linear equations in the variables \( u, w \), in which case we can rewrite this using matrix notation as $BR $BR

\{ mbox( [
	'\( (**)\quad \)',
	display_matrix( [ [ '\( s \)', '\( $d t \)' ], [ '\( t \)', '\( s \)' ] ] ),
	display_matrix( [ [ '\( u \)' ], [ '\( w \)' ] ] ),
	'\( = \)',
	display_matrix( [ [ '\( 1 \)' ], [ '\( 0 \)' ] ] )
] ) \}

By general facts from linear algebra, we know that this system has a unique solution if and only if the determinant of the matrix above is not zero, i.e. if and only if \( s^2 - $d t^2 \neq 0 \), or equivalently, \( $d \neq ( s/t )^2 \) for some rational numbers \( s, t \). $BR $BR

Since \( $d \) is not a perfect square, the system has a unique solution and so \( A( $d ) \) IS a field. Using \( (**) \) we can readily determine the inverse of any non-zero element of \( A( $d ) \). $BR $BR

Suppose we were dealing with \( A( d ) \) where \( d \) is a perfect square. Can you find the zero-divisors in this ring? $BR

END_SOLUTION
}


COMMENT( "Complete solution included, hint after 3 attempts. Problem involves solving system of 2 linear equations in 2 variables" );


ENDDOCUMENT(); 
