## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Complex eigenvalues)
## Date(01/31/2012)
## Institution(Rochester and Fort Lewis College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.6')
## Problem1('4')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());



#####################################
#  Setup 1

Context("Complex");

# dilation
$a = non_zero_random(-4,4,1);

$A11 =  3 * $a;  $A12 =  2 * $a;
$A21 = -5 * $a;  $A22 = -3 * $a;

$lambda1 =  Complex("$a i");
$v1top = Complex("3 + i");
$v1bot = Complex("-5");

$lambda2 = Complex("-$a i");
$v2top = Complex("3 - i");
$v2bot = Complex("-5");

$multians1 = MultiAnswer($lambda1, $v1top, $v1bot, $lambda2, $v2top, $v2bot)->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='0'>
    <tr>
      <td> %s, </td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td>&nbsp;&nbsp; %s,</td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  
    "%s, \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack, 
     %s, \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",
    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};

        Value::Error("Your vectors are parallel") if
           $c[1]->typeMatch($s[1]) 
        && $c[2]->typeMatch($s[2])
        && $c[4]->typeMatch($s[4])
        && $c[5]->typeMatch($s[5])
        && ($s[1] * $s[5] == $s[2] * $s[4]);

        if (   $c[0]->typeMatch($s[0])
            && $c[1]->typeMatch($s[1])
            && $c[2]->typeMatch($s[2])
            && ($s[1] != 0 || $s[2] != 0)
            && ($A11 * $s[1] + $A12 * $s[2] == $s[0] * $s[1])
            && ($A21 * $s[1] + $A22 * $s[2] == $s[0] * $s[2])
           )
        { $flag1 = 1; }

        if (   $c[3]->typeMatch($s[3])
            && $c[4]->typeMatch($s[4])
            && $c[5]->typeMatch($s[5])
            && ($s[4] != 0 || $s[5] != 0)
            && ($A11 * $s[4] + $A12 * $s[5] == $s[3] * $s[4])
            && ($A21 * $s[4] + $A22 * $s[5] == $s[3] * $s[5])
           )
        { $flag2 = 1; }

        if ($flag1==1 && $flag2 ==1) {
           return 1;
        } elsif ($flag1 == 1) {
           Value::Error("The eigenpair { lambda1, v1 } is correct");
        } elsif ($flag2 == 1) {
           Value::Error("The eigenpair { lambda2, v2 } is correct");
        }else{
              Value::Error("Neither of the eigenpairs is correct");
        }
       
    }
);






#####################################
#  Main text 1

Context()->texStrings;
BEGIN_TEXT
Consider the linear system
$PAR
\[ 
\displaystyle
\boldsymbol{ \vec{y}^{\,\prime} }
=
\left\lbrack
\begin{array}{rr}
$A11 & $A12 \\
$A21 & $A22
\end{array}
\right\rbrack
\boldsymbol{\vec{y}}.
\]
$BR
$BR
Find the eigenvalues and eigenvectors for the coefficient matrix.
$BCENTER
\{ 
mbox(
[
'\( \lambda_1 = \) ',
$multians1->ans_rule(5),
', \( \boldsymbol{\vec{v}_1} =  \) ', 
display_matrix([[$multians1->ans_rule(5)],[$multians1->ans_rule(5)]]),
', and \( \lambda_2 = \) ',
$multians1->ans_rule(5),
', \( \boldsymbol{\vec{v}_2} =  \) ', 
display_matrix([[$multians1->ans_rule(5)],[$multians1->ans_rule(5)]]),
],
allowbreaks=>'yes'
) 
\}
$ECENTER
END_TEXT
Context()->normalStrings;


##################################
#  Answer evaluation 1

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );


COMMENT('Complex eigenvalues +/- ai and complex eigenvectors.');

ENDDOCUMENT(); 
