## DESCRIPTION
## Generates a 2nd degree constant coefficient homogeneous recurrence relation.
## The characteristic equation will have one repeated nonzero real root.
## A multianswer parser will check both homogeneous solutions in any order entered.
## The problem has the form: a_n = p*a_(n-1) + q*a_(n-2), a_0 = c, a_1 = d
## ENDDESCRIPTION

## DBsubject(Discrete Math)
## DBchapter(Sequences)
## DBsection(Recurrence Relations)
## Date(06/03/2026)
## Institution(Louisiana Tech University)
## Author(Jason Terry)
## Level(3)
## KEYWORDS('discrete math','recurrence relation','2nd degree','homogeneous','constant coefficient','distinct roots')

DOCUMENT();

## Load libraries
loadMacros("PGstandard.pl",
	"MathObjects.pl",
	"parserMultiAnswer.pl",
	"PGML.pl",
	"PGcourse.pl",
	"contextFraction.pl");

## Set the context: no decimal input, absolute tolerance, variable "n", and integer test points for answer checker
Context("Fraction-NoDecimals");
Context()->flags->set( tolerance => 0.01, tolType => "absolute");
Context()->variables->are(n=>["Real",limits=>[0,15],resolution=>1]);

## Randomly generate a nonzero root of the characteristic equation
$r = non_zero_random(-6,6,1);

## Randomly generate initial conditions
$c = random(-4,4,1);
$d = random(-4,4,1);

## Compute the coefficients of the problem based on root
$p = 2*$r;
$q = -$r**2;

## Compute the particular solution constants based on initial conditions
$c1 = $c;
$c2 = Compute("($d - $c*$r)/$r");

## Construct the solution of the recurrence relation
$n = Formula("n");
$y1 = ($r**$n)->reduce;
$y2 = $n*($r**$n)->reduce;
$y = ($c1*$y1 + $c2*$y2)->reduce;

## Construct the answer parser for the general solution
$ans = MultiAnswer($y1, $y2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if (
           ($f1 == $f1stu && $f2 == $f2stu ) ||
           ($f1 == $f2stu && $f2 == $f1stu ) 
         ) {
          return 1;
      } else {
          return 0;
      }
  }
);


## Display problem
TEXT(beginproblem());

BEGIN_PGML
For the recurrence relation:

    [`` a_{n} = [$p]a_{n-1} + [$q]a_{n-2} ``], 

The general solution has the form:

	[`` a_{n} = c_{1} ``] [____________________]{$ans} [`` + c_{2} ``] [____________________]{$ans}

*Note: The arbitrary constants are placed for you, so your answer functions should have no other constants included (except 1). Both functions must be correct to receive credit.*

Given the initial conditions [`` a_{0} = [$c] ``] and [`` a_{1} = [$d] ``], solve for [`` c_{1} ``] and [`` c_{2} ``]. The final solution is:

	[`` a_{n} = ``] [________________________________________]{$y}

END_PGML

BEGIN_PGML_SOLUTION

The general solution is:

	[`` a_{n} = c_{1} \cdot [$y1] + c_{2} \cdot [$y2] ``]

Given the initial conditions [`` a_{0} = [$c] ``] and [`` a_{1} = [$d] ``], the final solution is:

	[`` a_{n} = [$y] ``]

END_PGML_SOLUTION

ENDDOCUMENT();