#DESCRIPTION
##Type of
#ENDDESCRIPTION

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
#for currency use Context("Currency") then Currency($A);
#Then, in the text use $DOLLAR $a
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
Context("Numeric");
    $ma = MultiAnswer(-5,3)->with(
         singleResult => 1,
         checker => sub {
           my ($correct, $student, $self) = @_;
           my ($a,$b) = @{$student};
           return $a < $b && abs($a) > abs($b);
         }
);
BEGIN_TEXT
Write two numbers so that $PAR
The first number is less than the second number, and $PAR
The absolute value of the first number is greater than the absolute value of the second number$PAR

\{$ma->ans_rule(3)\} \{$ma->ans_rule(3)\}

END_TEXT

ANS($ma->cmp);
BEGIN_SOLUTION
$PAR Solution $PAR
The only way this can happen is if the first number is negative so you just need to start by choosing$BR
two positive numbers and using the smaller as your second number and the multiply the larger by -1 $BR
and using it as your first number.$BR
Note that -5,3 is just a sample correct answer required by WeBWorK in this type of problem.

END_SOLUTION

;
ENDDOCUMENT();



## DBsubject(Algebra)
## DBchapter(Absolute value expressions and functions)
## DBsection(Solving equations with absolute values)
## Level(2)
## Static(1)
## MO(1)

