## DBCCSS()
## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection()
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('exponent','divide','polynomial','monomial')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl"
);

##############################################

Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);


#Monomial a*b*x^m
$a = non_zero_random(2, 5);
$b = non_zero_random(2, 5);
$m = random(3, 8);
$f1 = Formula("$a*$b*x^$m")->reduce;

#Answer is cx + d/x^n + x^p/b
$c = non_zero_random(-6, 6);
$d = non_zero_random(-6, 6);
do{
$n = random(2, 5);
}until($n < $m);
$p = random(2, 5);

#For Numerator:  abcx^(m+1) + abdx^(m-n) + ax^(p+m)
$c1 = $a*$b*$c;
$c2 = $a*$b*$d;
$c3 = $a;
$e1 = $m+1;
$e2 = $m-$n;
$e3 = $m+$p;
$f2 = Formula("$c1*x^$e1 + $c2*x^$e2 + $c3*x^$e3")->reduce;

#Answer
$ans = Formula("$c*x + $d/x^$n + x^$p/$b");

##############################################
TEXT(beginproblem());

BEGIN_PGML

Find the quotient.  Write your answer as three terms with positive exponents.  Do not use decimals.

    [``\frac{[$f2]}{[$f1]} ``][________________]{$cmp}

END_PGML

############################################################

ANS($ans -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));


############################################
BEGIN_PGML_SOLUTION
To divide the expression,

[``\frac{[$f2]}{[$f1]}``]

we have to divide every term by [`[$f1]`].

[``\frac{[$c1]x^{[$e1]}}{[$f1]} + \frac{[$c2]x^{[$e2]}}{[$f1]} + \frac{[$c3]x^{[$e3]}}{[$f1]}``]

We can look for common factors.

[``\frac{[$a]\cdot [$b] \cdot [$c] \cdot x^[$m] \cdot x}{[$a]\cdot [$b] \cdot x^[$m]} + \frac{[$a] \cdot [$b] \cdot [$d] x^[$e2]}{[$a]\cdot [$b] \cdot x^[$e2] \cdot x^[$n]} + \frac{[$a] \cdot x^[$p] \cdot x^[$m]}{[$a]\cdot [$b] \cdot x^[$m]}``]

[``[$ans]``]
END_PGML_SOLUTION
##############################################

ENDDOCUMENT();