## DESCRIPTION
##   Calculate the Limit
## ENDDESCRIPTION

## Tagged by nhamblet

## DBsubject(Calculus - single variable)
## DBchapter(Limits and continuity)
## DBsection(Squeeze theorem)
## Institution(ASU)
## Level(2)
## KEYWORDS('Limit', 'Squeeze')
## MO(1)-TF-12/18/2014

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl"
);

Context("ArbitraryString");

$a = random(-4,4,1);
$b = random(-6,6,1);
$c = random(-6,6,1);
# coefficients of quadratic  x^2 + a1x + a2 passing through c at a
$a1 = $b-$a;
$a2 = -$a*$b+$c;
# coefficients of line b1x +b2  passing through c at a and tangent to quadratic above.
$b1 = $a + $b;
$b2 = -$a*($a +$b) + $c;

TEXT(beginproblem());
$showPartialCorrectAnswers=1;

BEGIN_TEXT
If \[ ${b1}x +$b2 \le f(x) \le x^2 + ${a1}x + $a2 \]  determine
\(\displaystyle \lim_{x\to $a} f(x) \) = \{ans_rule(20) \}
$PAR

What theorem did you use to arrive at your answer? $BR
\{ ans_rule(50) \}                  	                  
END_TEXT

ANS(Compute($c)->cmp);

ANS(String("The Squeeze Theorem")->cmp(
  checker => sub {
    my ($correct,$student,$ansHash) = @_;
    my $s = lc($ansHash->{original_student_ans}); # lower case of raw student input
    return ( ($s =~ /squeeze/  or $s =~ /pinch/  or $s =~/sandwich/) ? 1 : 0);
  }
));

ENDDOCUMENT();
