# DESCRIPTION
# This problem involves finding the moment generating function of the binomial
# distribution. Then, given the values of parameters, finding the first 3 moments
# of the distribution and its variance.
# WebWork problem written by Vladislav Emelyanov
# ENDDESCRIPTION

## DBsubject('Probability')
## DBchapter('Random variables')
## DBsection('Generating function')
## KEYWORDS('moments','expected value','expectation','standard deviation','variance', 'sampling distributions','binomial distribution','binomial random variable')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Vladislav Emelyanov')
## Institution('UMASS-Amherst-Math-Stat')
## Date('3/31/2017')

DOCUMENT(); 
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

Context("Numeric");
Context()->variables->are(n=>"Real",p=>"Real",t=>"Real",q=>"Real");

$p = random(0.1, .9, .1);
$n = random(10,60,1);

Context()->texStrings;
BEGIN_TEXT
If \( Y \) is \( binomial(n,p) \), find the MGF of \( Y \).
$BR
\( M(t)= \)  \{ans_rule(20) \} 
END_TEXT

$showHint = 2;
BEGIN_HINT
You may derive a formula for \( M(t) \) or find it using a reference.
END_HINT

BEGIN_TEXT
$BR $BR
If \( n = $n \) and \( p = $p \), differentiate the MGF you found above to find the first 3 moments of \( Y \) about 0.
$BR
1st Moment: \{ ans_rule(10) \}
$BR
2nd Moment: \{ ans_rule(10) \}
$BR
3rd Moment: \{ ans_rule(10) \} 
$BR
Using the moments above, calculate the variance of \( Y \).
$BR
\( var(Y)= \)  \{ans_rule(10) \} 
END_TEXT

$showHint = 2;
BEGIN_HINT
Remember: \( var(Y) = E[Y^2] - E[Y]^2 \)
END_HINT

$mgf1 = Formula("(p e^t + q)^n");
$mgf2 = Formula("(p e^t + (1-p))^n");

ANS($mgf2 ->cmp(
   checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ($ans) = $student->value;
      if ( $mgf1 == $ans || $mgf2 == $ans ) {
         return 1;
      } else {
         return 0;
      }
   }
));

$moment1 = $n * $p;
$moment2 = Formula("$n*$p - (1-$n)*$n*$p^2");
$moment3 = Formula(($n-2)*($n-1)*$n*$p*$p*$p+3*($n-1)*$n*$p*$p+$n*$p);
$variance = $n*$p*(1-$p);

ANS(num_cmp($moment1));
ANS(num_cmp($moment2));
ANS(num_cmp($moment3));
ANS(num_cmp($variance));
ENDDOCUMENT();