##DESCRIPTION
## DBsubject(Trigonometry)
## DBchapter(Analytic trigonometry)
## DBsection(Solving trigonometric equations exactly)
## Institution(NAU)
## Author(Monika Keindl)
## MLT(factorForm)
## Level(4)
## MO(1)
## KEYWORDS('identity', 'trigonometry', 'equation')
##ENDDESCRIPTION

# File Created: 7/5/05
# Location: Northern Arizona University
# Course: MAT 125

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(&beginproblem);

Context("Numeric");
Context()->variables->add(k=>"Real"); 
Context()->flags->set(
  reduceConstants => 0,
  reduceConstantFunctions => 0
);

Parser::Number::NoDecimals();

## trig functions to mix Tangent with, and their roots 
## there are two answer formulas if f = sin(x), but only 
## one if f = cos(x) because cos(x) cancels the denominator of tan(x)

$f_is_cos=random(0,1,1);  # f index

if ($f_is_cos){
  $f = Compute("cos(x)");
}
else {
  $f = Compute("sin(x)");
  $f_root = Compute("k*pi");
}

## randomize tangent formula and its additional roots 

$tan_ind=random(0,5,1);
if ($tan_ind == 0){
  $ques = Compute("3*$f*tan(x)-sqrt(3)*$f");
  $add_root = Compute("pi/6+k*pi");
}
elsif ($tan_ind == 1){
  $ques = Compute("$f*tan(x)-$f");
  $add_root = Compute("pi/4+k*pi");
}
elsif ($tan_ind == 2){
  $ques = Compute("$f*tan(x)-sqrt(3)*$f");
  $add_root = Compute("pi/3+k*pi");
}
elsif ($tan_ind == 3){
  $ques = Compute("$f*tan(x)+sqrt(3)*$f");
  $add_root = Compute("2*pi/3+k*pi");
}
elsif ($tan_ind == 4){
  $ques = Compute("$f*tan(x)+$f");
  $add_root = Compute("3*pi/4+k*pi");
}
else {
  $ques = Compute("3*$f*tan(x)+sqrt(3)*$f");
  $add_root = Compute("5*pi/6+k*pi");
}

if ($f_is_cos){ 
  $ans = $add_root;
}
else{  
  $ans = MultiAnswer($f_root, $add_root)->with(
    singleResult => 0,
    checker => sub {
      my ($correct, $student, $self) =@_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
	   ($f1 == $f2stu && $f2 == $f1stu) ) {
	  return [1,1];
      } else {
	if ($f1 == $f1stu || $f2 == $f1stu) {
	   return [1,0];
	} elsif ($f1 == $f2stu || $f2 == $f2stu) {
	      return [0,1];
	} else {
	      return [0,0];
	}
      }
    }
  );
}

Context()->texStrings;
BEGIN_TEXT
$PAR

Find all angles \(x\) (in radians) that satisfy the equation 

\[$ques=0.\]

Your answers should be formulas that look like
\[ x = \text{angle}+k*\text{period}\]
where \(k=\cdots,-2,-1,0,1,2,\cdots\) ranges over the integers, \( 0\leq \text{angle} < \text{period} \),  and period is the smallest positive "period" that makes the formula work. 
$PAR
Note: decimal answers are not allowed, so, for example, if \(\pi\) appears in your answer call it $BBOLD pi $EBOLD not 3.14159 . 

$PAR
The solutions of the equation are:
$PAR
END_TEXT
Context()->normalStrings;


if ($f_is_cos){
Context()->texStrings;
BEGIN_TEXT
\(x\) = \{$ans->ans_rule(20)\} $BR
(Warning: division by zero is not allowed!)
END_TEXT
Context()->normalStrings;
}
else {
Context()->texStrings;
BEGIN_TEXT
\(x\) = \{$ans->ans_rule(20)\} 
$BR
or
$BR
\(x\) = \{$ans->ans_rule(20)\}. 
END_TEXT
Context()->normalStrings;
}

ANS($ans->cmp());

ENDDOCUMENT(); 
