## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Inverses)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.3')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');

$a11 = random(2,5,1);
$a21 = random(2,5,1);
$a31 = 1;

$b = random(-2,2,4);
$a12 = $b * $a11;
$m = random(-1,1,2);
$a22 = $b * $a21 + $m;
$a32 = $b * $a31;

$c = random(-3,-1,1);
$d = non_zero_random(-3,3,1);
$n = random(-1,1,2);
$a13 = $c * $a11 + $d * $a12 + $n;
$a23 = $c * $a21 + $d * $a22;
$a33 = $c * $a31 + $d * $a32;

while((abs($a13)<2) or (abs($a23)<2) or (abs($a33)<2)) {
	$c = $c+1;
	$a13 = $c * $a11 + $d * $a12 + $n;
	$a23 = $c * $a21 + $d * $a22;
	$a33 = $c * $a31 + $d * $a32;
}

$det = - $a31 * $m * $n;

$ans11 = Real("($a22 * $a33 - $a23 * $a32)/$det");
$ans12 = Real("- ($a12 * $a33 - $a13 * $a32)/$det");
$ans13 = Real("($a12 * $a23 - $a13 * $a22)/$det");
$ans21 = Real("- ($a21 * $a33 - $a23 * $a31)/$det");
$ans22 = Real("($a11 * $a33 - $a13 * $a31)/$det");
$ans23 = Real("- ($a11 * $a23 - $a13 * $a21)/$det");
$ans31 = Real("($a21 * $a32 - $a22 * $a31)/$det");
$ans32 = Real("- ($a11 * $a32 - $a12 * $a31)/$det");
$ans33 = Real("($a11 * $a22 - $a12 * $a21)/$det");

$NO_SPACE = '@{}';


$A = MultiAnswer($ans11,$ans12,$ans13,$ans21,$ans22,$ans23,$ans31,$ans32,$ans33)->with(
  singleResult => 1,
  allowBlankAnswers => 0,

  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      foreach my $i (0..scalar(@c)-1) {
           if (Real($s[$i]) != $c[$i]) { return 0; };
      }
      return 1;
  }
);

Context()->texStrings;
BEGIN_TEXT

Find the inverse of the linear transformation 
\[ \begin{array}{r${NO_SPACE}r${NO_SPACE}r${NO_SPACE}r${NO_SPACE}r}
y_1 & = & $a11 x_1 & + $a12 x_2 & + $a13 x_3 \\
y_2 & = & $a21 x_1 & + $a22 x_2 & + $a23 x_3 \\
y_3 & = &      x_1 & + $a32 x_2 & + $a33 x_3 
\end{array} \]
$BR
\(x_1=\) \{$A->ans_rule(5)\}\(y_1+\) \{$A->ans_rule(5)\}\(y_2+\) \{$A->ans_rule(5)\}\(y_3\), $BR
\(x_2=\) \{$A->ans_rule(5)\}\(y_1+\) \{$A->ans_rule(5)\}\(y_2+\) \{$A->ans_rule(5)\}\(y_3\), $BR
\(x_3=\) \{$A->ans_rule(5)\}\(y_1+\) \{$A->ans_rule(5)\}\(y_2+\) \{$A->ans_rule(5)\}\(y_3\). $BR

END_TEXT
Context()->normalStrings;

ANS( $A->cmp );
;
ENDDOCUMENT();
