## DESCRIPTION
## Laplace transforms: partial fractions
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Laplace transforms)
## DBsection(Inverse transformations)
## Institution(Fort Lewis College)
## MLT(inverse_03)
## Level(4)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('5.3')
## Problem1('15')
## KEYWORDS('Laplace transform')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGunion.pl",
  "parserMultiAnswer.pl",
  "parserFunction.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
s=>"Real",
t=>"Real"
);

# Redefine the Heaviside function so the answer checker will take either 1 or h(t).
parserFunction("h(t)" => "1");

$a = random(1,4,1);
$asqr = $a*$a;
$B = random(1,5,1);
do {$C = random(1,5,1); } until ($C != $B);
$D = random(1,6,1);

$Fs = Compute("(($B+$D)*s^2 + $C*s + $D*$asqr)/(s^3 + $asqr*s)")->reduce; 
$fac1 = Compute("($B*s+$C)/(s^2+$asqr)");
$fac2 = Compute("$D/s");
$ft = Compute("$B*cos($a*t) + $C/$a*sin($a*t) + $D");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  allowBlankAnswers =>1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Consider the function \( \displaystyle F(s) = $Fs \).

\{ BeginList('OL', type=>'a') \}

$ITEM Find the partial fraction decomposition of \( F(s) \):
$BR
$BR
\( \displaystyle $Fs = \) \{ $multians->ans_rule(20) \} + \{ $multians->ans_rule(20) \}

$ITEMSEP
$ITEM  Find the inverse Laplace transform of \( \displaystyle F(s) \).
$BR
$BR
\( \displaystyle f(t) = {\mathcal L}^{-1} \left\lbrace F(s) \right\rbrace = \)
\{ ans_rule(50) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL', type=>'a') \}
END_TEXT
Context()->normalStrings;


##############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

#UNORDERED_ANS( $fac1->cmp(), $fac2->cmp() );

ANS( $ft->cmp() );

ENDDOCUMENT();


