# DESCRIPTION
# Homework problem from A First Course in Probability, Ross;
# 9th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 13/09/25

## DBsubject(Probability)
## DBchapter(Laws, theory)
## DBsection(Central limit theorem)
## Date(08/12/16)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## TitleText1('A First Course in Probability')
## AuthorText1('Sheldon Ross')
## EditionText1('9')
## Section1('8.3')
## Problem1('6')
## KEYWORDS('probability,random variables,discrete random variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGnumericalmacros.pl",
  "PGstatisticsmacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->flags->set( tolerance=>0.005, tolType=>'absolute' );

# normal_prob(a, b, mean=>0, deviation=>1);

$nrolls   = random(70, 100, 1);
$nm1rolls = $nrolls - 1;
# this keeps the probability close to the same value
$sum      = int(1.54814*sqrt(2.91967*$nrolls) + 3.5*$nrolls);

$z_score = ($sum+.5 - 3.5*$nm1rolls)/sqrt(35*$nm1rolls/12);
$prob = Compute( normal_prob(-infty, $z_score) );
# alternate values: leave off continuity correction
$z2 = ($sum - 3.5*$nm1rolls)/sqrt(35*$nm1rolls/12);
$prob_alt = Compute( normal_prob(-infty, $z2) );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Suppose we roll a fair six-sided die and sum the values obtained
on each roll, stopping once our sum exceeds $sum.  Approximate the
probability that at least $nrolls rolls are needed to get this sum.
$BR
Probability = \{ $prob->ans_rule() \}

END_TEXT
Context()->normalStrings;

ANS( $prob->cmp( checker=>sub { my ( $c,$s,$a ) = @_;
				return $c==$s || $prob_alt==$s } ) );

$sumpp5 = $sum + 0.5;
$z_score_r = sprintf("%1.5f", $z_score);

Context()->texStrings;
BEGIN_SOLUTION

Because the expected value of each die roll is \(3.5\), the
expected value of the sum of \(n\) such rolls is \(3.5n\).  Similarly,
given the reasonable expectation that the rolls are independent, we
know that the variance of the sum is the sum of the variances of each
of the rolls.  The variance of one roll is
\(\mbox{Var}(R) = E[R^2] - E[R]^2 = E[R^2] - 3.5^2\).
By direct calculation,
\(E[R^2] = \frac16(1 + 4 + \cdots + 36) = \frac{91}{6}\),
so that
\(\mbox{Var}(R) = \frac{35}{12} \approx 2.91667\).
Then the probability that at least $nrolls are needed to get $sum
is one minus the probability that $nm1rolls are not enough.  If \(X\)
is the sum of the $nm1rolls, this is
\[ P\lbrace X < 300 \rbrace =
P\lbrace \frac{X - 3.5($nm1rolls)}{\sqrt{(35/12)($nm1rolls)}} <
\frac{$sumpp5 - 3.5($nm1rolls)}{\sqrt{(35/12)($nm1rolls)}} \rbrace =
\Phi(\frac{$sumpp5 - 3.5($nm1rolls)}{\sqrt{(35/12)($nm1rolls)}} \approx
\Phi($z_score_r) = $prob. \]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
