## MO(1)
# WeBWorK problem written by Alex Jordan 
# Portland Community College
#
# Template:
# Simplify the expression. For this problem you may assume that all variables are positive. 
#
#       sqrt(x^n)n-th-root(x^2)
#
#
# Last updated: Yao 1/5/2017
# ENDDESCRIPTION


##############################################


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPowers.pl",
  "contextFraction.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);


##############################################


TEXT(beginproblem());

Context("Fraction");
$var = list_random('x','y','z','t','r','m','n','a','b','c');
Context()->variables->are($var=>'Real');
Context()->flags->set(reduceConstants=>0);
Context()->variables->set($var=>{limits=>[0,2]});
do {
   $num1 = 1;
   $den1 = random(2,5,1);
   $num2 = random(2,5,1);
   do {$den2 = random($num2+1,10,1);} until ( ($den2!=$den1) && (gcd($den2,$num2)==1) );
   $commonD = lcm($den1,$den2);
   $f1 = $commonD/$den1;
   $f2 = $commonD/$den2;
   $frac = Fraction($num1*$den2+$num2*$den1,$den1*$den2)->reduce;
   ($ansNum,$ansDen) = $frac->value;
} until ( $ansDen != $commonD );

$problem = ($den1==2) ? "\sqrt{$var} \cdot \sqrt[$den2]{$var^{$num2}}" : "\sqrt[$den1]{$var}\cdot\sqrt[$den2]{$var^{$num2}}";

Context("Numeric");
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
'/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
'//' => {class => 'bizarro::BOP::divide', isCommand => 1},
);

Context()->variables->are($var=>'Real');
Context()->flags->set(reduceConstants=>0);
Context()->operators->undefine('*');
Context()->operators->undefine('+');
Context()->functions->undefine('sqrt');
Context()->variables->set($var=>{limits=>[0,2]});

$ans = Formula("$var^($frac)");
$evaluator = $ans->cmp(
     checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is not in the expected form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0);
      return 1;
  }
);


##############################################

BEGIN_PGML

Simplify the expression, answering with rational exponents and not radicals. [@KeyboardInstructions('To enter [`x^{m/n}`], type [|x^(m/n)|]*.')@]**

[``[$problem]=``]    [_________________]{$evaluator}

END_PGML


##############################################

BEGIN_PGML_SOLUTION

We will use the following properties:

    [`` \sqrt[m]{x^n}=x^{\frac{n}{m}} ``]

    [`` x^m\cdot x^n=x^{m+n} ``]

The solution is:

    [``
\begin{aligned}
   [$problem]
   &= [$var]^{\frac{[$num1]}{[$den1]}} \cdot [$var]^{\frac{[$num2]}{[$den2]}} \\
   &= [$var]^{\frac{[$num1]}{[$den1]}+\frac{[$num2]}{[$den2]}} \\
   &= [$var]^{\frac{[$num1*$f1]}{[$den1*$f1]}+\frac{[$num2*$f2]}{[$den2*$f2]}} \\
   &= [$var]^{\frac{[$num1*$f1+$num2*$f2]}{[$den1*$f1]}} \\
   &= [$var]^{[$frac]}
\end{aligned}
    ``]

END_PGML_SOLUTION


##############################################


ENDDOCUMENT();
