## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection('Span')
## Institution('TCNJ')
## KEYWORDS('vectors' 'span')
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('2.2')
## Problem1('')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
"PG.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGgraphmacros.pl",
"PGmatrixmacros.pl", 
"PGnumericalmacros.pl",
"parserMultiAnswer.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$ans1 = non_zero_random(-5,5,1);
$ans2 = 0;
$ans3 = non_zero_random(-5,5,1);

$v1 = new Matrix(3,1);
$v1->assign(1,1, $ans1);
$v1->assign(2,1, $ans2);
$v1->assign(3,1, $ans3);

Context('Numeric');
$mp = MultiAnswer($ans1, $ans2, $ans3)->with(singleResult=>1,
  checker => sub {
    my ($correct,$student,$self) = @_;  # get the parameters
	my ($x1,$x2,$x3) = @{$student};      # extract the student answers
	my ($c1,$c2,$c3) = @{$correct};      # extract the correct answers
	if($x1 == Real(0) and $x2 == Real(0) and $x3 == Real(0)) {
	  return 0;
	}
    return 1 if( $x1*Real($c3) == $x3*Real($c1) and $x2 == Real($c2));
	return 0;
  }
);

if($ans1==1){
   $str1 = "t";}
elsif($ans1==-1){
   $str1 = "-t";}
else{
   $str1 = $str1 . "$ans1";
   $str1 = $str1 . "t";}

$str2 = "0";

if($ans3==1){
   $str3 = "t";}
elsif($ans3==-1){
   $str3 = "-t";}
else{
   $str3 = $str3 . "$ans3";
   $str3 = $str3 . "t";}

BEGIN_TEXT

\{ mbox( 'Let \(H\) be the set of all vectors of the form: ', display_matrix([[$str1],[$str2],[$str3]]), '.' ) \}
Find a vector \(v\) in \( {\mathbb R}^3\) such that \( H=\textrm{Span}(\left\{"\{"\} v\right\}) \).
$BR  
\{ mbox( ' \(v=\)', display_matrix([[$mp->ans_rule(5)],[$mp->ans_rule(5)], 
 [$mp->ans_rule(5)]]), '.' ) \}

END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();       # This should be the last executable line in the problem.
