## DESCRIPTION
## Polynomial and Rational Functions


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Functions Modeling Change')
## TitleText2('Functions Modeling Change')
## TitleText3('Functions Modeling Change')
## AuthorText1('Connally')
## AuthorText2('Connally')
## AuthorText3('Connally')
## EditionText1('3')
## EditionText2('4')
## EditionText3('5')
## Section1(9.Tools)
## Section2(11.Skills)
## Section3('11.Skills')
## Problem1(52)
## Problem2(43)
## Problem3('43')
## KEYWORDS('polynomial functions','rational functions')


####################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGunion.pl",
  "parserMultiAnswer.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

####################################
#  Setup

Context("Numeric")->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[1,3]});

$a = random(2,9,1);
$b = random(2,9,1);
$c = random(2,9,1);
$d = random(2,9,1);

$n = random(3,9,1);

$fraction = "\frac{ $a t^2 + $b t - $c}{ $d t^{$n} }";

$ans1 = Compute("($a/$d) t^(2-$n)")->reduce;
$ans2 = Compute("($b/$d) t^(1-$n)")->reduce;
$ans3 = Compute("($c/$d) t^(-$n)")->reduce;

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Split into a sum or difference of reduced fractions.
$PAR
\( \displaystyle $fraction = \)
\{ $multians->ans_rule(15) \}
+
\{ $multians->ans_rule(15) \}
-
\{ ans_rule(15) \}
END_TEXT
Context()->normalStrings;


####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );

ANS($ans3->cmp()
->withPostFilter(AnswerHints(
  -$ans3 => "There is already minus sign in front of the third answer box",
))
);

($nu,$dn) = reduce($a,$d);
$adfr = ( $dn == 1 ) ? $nu : "\frac{$nu}{$dn}";
($nu,$dn) = reduce($b,$d);
$bdfr = ( $dn == 1 ) ? $nu : "\frac{$nu}{$dn}";
($nu,$dn) = reduce($c,$d);
$cdfr = ( $dn == 1 ) ? $nu : "\frac{$nu}{$dn}";

$twomn = 2-$n;
$onemn = 1-$n;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD SOLUTION $EBOLD
$PAR
We have
\[ \frac{ $a t^2 + $b t - $c}{ $d t^{$n} } =
   \frac{$a t^2}{$d t^{$n}} + \frac{$b t}{$d t^{$n}} - \frac{$c}{$d t^{$n}} =
   $adfr\,t^{$twomn} + $bdfr\,t^{$onemn} - $cdfr\,t^{-$n}.
 \]

END_SOLUTION
Context()->normalStrings;

;
ENDDOCUMENT();
