## DESCRIPTION
## Polynomial and Rational Functions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Functions with fractional exponents and radical functions)
## DBsection(Functions with fractional exponents)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## TitleText1('Functions Modeling Change')
## AuthorText1('Connally')
## EditionText1('3')
## Section1('9')
## Problem1('46')
## TitleText2('Functions Modeling Change')
## AuthorText2('Connally')
## EditionText2('4')
## Section2('11.')
## Problem2('46')
## KEYWORDS('polynomial functions','rational functions')


####################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGunion.pl",
"parserMultiAnswer.pl",
"PGcourse.pl",
);

TEXT(beginproblem());

####################################
#  Setup

Context("Numeric")->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[1,2]});

$n = random(2,7,1);

$fraction = "\frac{ t^{-1/2} + t^{1/2} }{ t^{$n} }";

$ans1 = Compute("1/(t^($n+1/2))")->reduce;
$ans2 = Compute("1/(t^($n-1/2))")->reduce;

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Split into a sum or difference of reduced fractions.
$PAR
\( \displaystyle $fraction = \)
\{ $multians->ans_rule(15) \}
+
\{ $multians->ans_rule(15) \}
END_TEXT
Context()->normalStrings;


####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );

$npp5 = (2*$n + 1) . "/2";
$nmp5 = (2*$n - 1) . "/2";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD SOLUTION $EBOLD
$PAR
We have
\[ \frac{ t^{-1/2} + t^{1/2} }{ t^{$n} } =
   \frac{t^{-1/2}}{t^{$n}} + \frac{t^{1/2}}{t^{$n}} =
   \frac{1}{t^{$n+1/2}} + \frac{1}{t^{$n-1/2}} =
   \frac{1}{t^{$npp5}} + \frac{1}{t^{$nmp5}}.
 \]

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
