## DESCRIPTION
## Polynomial and Rational Functions


## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Functions Modeling Change')
## TitleText2('Functions Modeling Change')
## TitleText3('Functions Modeling Change')
## AuthorText1('Connally')
## AuthorText2('Connally')
## AuthorText3('Connally')
## EditionText1('3')
## EditionText2('4')
## EditionText3('5')
## Section1(9.Tools)
## Section2(11.Skills)
## Section3('11.Skills')
## Problem1(28)
## Problem2(23)
## Problem3('23')
## KEYWORDS('polynomial functions','rational functions')


####################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGunion.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

####################################
#  Setup

Context("Numeric");
Context()->variables->set(x=>{limits=>[1,2]});
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(3,9,1);
$b = random(3,9,1);
$n = random(1,4,1);

$n2p1 = 2*$n+1;

$fraction = "\frac{$a}{\sqrt{x}} - \frac{$b}{(\sqrt{x})^{$n2p1}} ";

$num = Compute("($a x^$n - $b)");
$den = Compute("(x^($n) * sqrt(x))");

$numbogus = Formula("$a-$b");
$denbogus = Formula("x^($n+1)");

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};

      if ( ($f1 == $f1stu && $f2 == $f2stu) ) {
          return [1,1];
      } elsif ($f1 == $f1stu) {
          return [1,0];
      } elsif ($f2 == $f2stu) {
          return [0,1];
      } elsif ($numbogus==$f1stu || $denbogus==$f2stu) {
          $ansHash->setMessage(1,"Find a common denominator first");
          $ansHash->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


#
#  Display the fraction and answer blanks nicely
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
  $showfraction =
  "\[ $fraction = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $showfraction =
  ColumnTable(
  "\( \displaystyle $fraction = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}
Context()->normalStrings;



####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Perform the indicated operations.
Express your answer in reduced form.
$PAR
$BCENTER
$showfraction
$ECENTER
END_TEXT
Context()->normalStrings;


####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

install_problem_grader(~~&std_problem_grader);

ANS($multians->cmp() );

$twon = 2*$n;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD SOLUTION $EBOLD
$PAR
We have
\[ \frac{$a}{\sqrt{x}} - \frac{$b}{(\sqrt{x})^{$n2p1}} =
   \frac{$a\,(\sqrt{x})^{$twon}}{(\sqrt{x})^{$n2p1}} -
      \frac{$b}{(\sqrt{x})^{$n2p1}} =
   \frac{$a\,x^{$n} - $b}{(\sqrt{x})^{$n2p1}} =
   \frac{$a\,x^{$n} - $b}{x^{$n}\,\sqrt{x}}.
 \]

END_SOLUTION
Context()->normalStrings;

;
ENDDOCUMENT();


$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ) {
          return [1,1];
      } elsif ($f1 == $f1stu) {
          return [1,0];
      } elsif ($f2 == $f2stu) {
          return [0,1];
      } elsif ( Formula("$a-$b")==$f1stu ) {
          $ansHash->setMessage(1,"Find a common denominator first");
          $ansHash->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f1/$f2==$f1stu/$f2stu ) {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];
      }
  }
);
