# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 07/26/10


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Applications of double integrals)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('19.3')
## Section2('21.3')
## Section3('21.3')
## Problem1('9')
## Problem2('11')
## Problem3('11')
## KEYWORDS('flux integral', 'surface integral', 'multivariable', 'integral', 'calculus')

## Textbook tags
## HHChapter1('Flux Integrals')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are( s=>'Real', t=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(2,8,1);
$x0 = random(5,30,5);
$y0 = random(5,30,5);

$param = MultiAnswer( Compute("s"), Compute("t"), Compute("$a*s + $b*t"),
		      Compute("0"), Compute("$x0"), Compute("0"),
		      Compute("$y0") )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ans ) = @_;

	my @res = ( 0,0,0, 0,0, 0,0 );

	if ( $s->[2] == $a*$s->[0] + $b*$s->[1] ) {
	    ## then we're ok as long as the range of values for x and y
	    ##    work
	    $res[0] = 1;
	    $res[1] = 1;
	    $res[2] = 1;
	    if ( $c->[3] == $s->[0]->eval( s=>$s->[3], t=>$s->[5] ) &&
	 	 $c->[4] == $s->[0]->eval( s=>$s->[4], t=>$s->[6] )  ) {
		$res[3] = 1;
		$res[4] = 1;
	    }
	    if ( $c->[5] == $s->[1]->eval( s=>$s->[3], t=>$s->[5] ) &&
	 	 $c->[6] == $s->[1]->eval( s=>$s->[4], t=>$s->[6] ) ) {
		$res[5] = 1;
		$res[6] = 1;
	    }
	}
	return [ @res ];
} );

$area = Compute( "sqrt($a*$a + $b*$b + 1)*$x0*$y0" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find the surface area of the
region \( S \) on the plane
\( z= $a x + $b y \)  such that \( 0\le x \le $x0 \) and
\( 0\le y\le $y0 \)
by finding a parameterization of the surface and then calculating the
surface area.
$PAR
A parameterization is
$BR
\( x(s,t) = \) \{ $param->ans_rule(35) \},
$BR
\( y(s,t) = \) \{ $param->ans_rule(35) \}, and
$BR
\( z(s,t) = \) \{ $param->ans_rule(35) \}, with
$BR
\{ $param->ans_rule(5) \} \(\le s\le\) \{ $param->ans_rule(5) \} and
$BR
\{ $param->ans_rule(5) \} \(\le t\le\) \{ $param->ans_rule(5) \}.

$PAR
Then, the surface area = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( $param->cmp() );
ANS( $area->cmp() );

$asq = $a*$a + $b*$b + 1;
$x0y0 = $x0*$y0;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

A parameterization of \(S\) is
\[ x=s, \quad y=t, \quad z=$a s + $b t, \quad
\mbox{for } 0\le s\le $x0, \quad 0\le  t \le $y0.\]
We compute
\[
\frac{\partial \vec r }{\partial s} \times 
\frac{\partial \vec r }{\partial t} =
(\vec i + $a\vec k)\times (\vec j + $b\vec k)= 
-$a\vec i - $b\vec j +\vec k,
\]
so that
\[
\left\Vert \frac{\partial \vec r }{\partial \theta} \times 
\frac{\partial \vec r }{\partial t} \right\Vert =
\sqrt{ $asq }.
\]
Then
\[
\mbox{Surface area} = \int_S dA =
\int_R \left\Vert \frac{\partial \vec r }{\partial \theta} \times
\frac{\partial \vec r }{\partial t} \right\Vert dA  =
\int_{t=0}^{$y0} \int_{s=0}^{$x0} \sqrt{$asq}\,ds\,dt  = $x0y0\sqrt{$asq}.
\]

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
