## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear independence)
## Institution(TCNJ)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('2.3')
## Problem1('')
## KEYWORDS('vectors' 'basis' 'linearly independent' 'dependent')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "PGnumericalmacros.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

$flip = 1;  # choose between dependent and independent vectors
$lambda = new Matrix(4,1);
$a=non_zero_random(1,5,1);
$b=non_zero_random(1,5,1);
$cc=non_zero_random(1,5,1);

# we'll scramble the vectors and coefficients

 
## $vec[$slice[0]] = Matrix->new_column_matrix([$v01, $v02, $v03, $v04]);


##$vec[1] = Matrix->new_column_matrix([ $v11, $v12, $v13, $v14 ]);

    $lambda->assign(1 ,1,$a);
    $lambda->assign(2 ,1,$b);
    $lambda->assign(3 ,1,$cc);
    $lambda->assign(4 ,1,-1);

$sl = new_pop_up_select_list();
$sl->qa(
"Determine whether or not the four vectors listed above are linearly independent or linearly dependent.",
'Linearly_Dependent');
$sl->choose(1);
$sl->ra_pop_up_list(['No_Answer_Entered' => 'Select an Answer', 'Linearly_Independent' => 'Linearly Independent', 'Linearly_Dependent' => 'Linearly Dependent']);
$string1=nicestring([$a, $b, $cc], ['r','u','d']);

$coef = MultiAnswer( Compute($lambda->element(1,1)),
        Compute($lambda->element(2,1)),
        Compute($lambda->element(3,1)),
        Compute($lambda->element(4,1)) )->with(singleResult=>1,
        checker=> sub {
         my ( $c, $s, $ans ) = @_;
         my @sans = @$s;
         my $zero = Compute(0);
         if ( $zero == $sans[0] && $zero == $sans[1] && $zero == $sans[2] && $zero == $sans[3] ) {
          $ans->setMessage(2, "Your responses should not be all zero.");
          return 0;
          } else {
          if ( $sans[0]*$b - $sans[1]*$a == 0 &&
             $sans[0]*$cc - $sans[2]*$a == 0 &&
             $sans[0]*(-1) - $sans[3]*$a  == 0 &&
             $sans[1]*$cc - $sans[2]*$b  == 0 &&
             $sans[1]*(-1) - $sans[3]*$b == 0  &&
             $sans[2]*(-1) - $sans[3]*$cc == 0) {
            return 1;
          } else {
            return 0;
          }
}});
    

BEGIN_TEXT
Let \(  S= \left\{"\{"\}r, u, d, x \right\} \) be a set of vectors.  
$BR $BR
If \(x = $string1 \),  determine whether or not \(S\) is linearly independent.

\{$sl->print_q()\} 

$BR
If \(S\) is dependent, enter a non-trivial linear relation below.  Otherwise, enter 0's for the coefficients.
$BR
$BR
\{ $coef->ans_rule(10) \}\(r+\)
\{ $coef->ans_rule(10) \}\(u+\)
\{ $coef->ans_rule(10) \}\(d+\)
\{ $coef->ans_rule(10) \}\(x =0\).
$BR

END_TEXT
ANS str_cmp($sl->ra_correct_ans);
ANS($coef->cmp());


ENDDOCUMENT() ;        
