## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear independence)
## Institution(TCNJ)
## MLT(indep2Q)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('2.3')
## Problem1('')
## KEYWORDS('vectors' 'basis' 'linearly independent' 'dependent')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserPopUp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector");

$showPartialCorrectAnswers = 1;

$flip = random(0,1,1);  # choose between dependent and independent vectors

$vec[0] = Vector( non_zero_random(-5,5,1), random(-5,5,1), non_zero_random(-5,5,1));

$vec[1] = 
  Real(non_zero_random(-2,2,1)) * $vec[$slice[0]] + 
     Vector( 0, non_zero_random(-3,3,1), non_zero_random(-3,3,1));

if ($flip==0) {
  $vec[2] = Real(non_zero_random(-2,2,1)) * $vec[0]  
           + Real(random(-2,2,1)) * $vec[1]  
           + Vector( 0, 0, non_zero_random(-5,5,1) );
  @lambda = (0,0,0);
} else {
  @lambda = ( Real(non_zero_random(-5,5,1)), Real(non_zero_random(-5,5,1)), Real(-1));
  $vec[2] = $lambda[0]*$vec[0] + $lambda[1]*$vec[1];
}

# we'll scramble the vectors and coefficients
@slice = NchooseK(3,3);
 

($vec0,$vec1,$vec2)=@vec[@slice];  
($lam0,$lam1,$lam2)=@lambda[@slice];

$multians = MultiAnswer($lam0,$lam1,$lam2)->with(
  singleResult => 1,
  checker => sub {
    my ($correct,$student,$self)=@_;
    my @cor; my @stu;
    @cor[@slice] = @{$correct};  # unscramble correct lambdas
    @stu[@slice] = @{$student};  # unscramble student's lambdas
    if ($flip == 0){
      if ( $stu[0] == 0 && $stu[1] == 0 && $stu[2] == 0 ){
        return 1;
      } 
      else {
        return 0;
      }
    }
    else {
      if ( $stu[2] !=0 && $stu[0]*$cor[2] == $cor[0]*$stu[2] && $stu[1]*$cor[2] == $cor[1]*$stu[2] ){
        return 1;
      }
      else { 
        return 0;
      } 
    }
  }
);

$selection = ($flip==0) ? "Linearly Independent" : "Linearly Dependent";
$popUp = PopUp(["Select an answer:", "Linearly Independent", "Linearly Dependent"], $selection);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Let A = \($vec0\), B = \($vec1\), and C = \($vec2 \) 

$PAR
Determine whether or not the three vectors listed above are linearly independent or linearly dependent.
$BR \{$popUp->menu()\}

$PAR
If they are linearly dependent, determine a non-trivial linear relation.
Otherwise, if the vectors are linearly independent enter 0's for 
the coefficients, since that relationship $BBOLD always $EBOLD holds.
$BR
\(0\) = \{$multians->ans_rule(5)\}\(A+\) \{$multians->ans_rule(5)\}\(B+\) \{$multians->ans_rule(5)\} \(C\).

END_TEXT
Context()->normalStrings;

ANS($popUp->cmp());
ANS($multians->cmp());

ENDDOCUMENT() ;        
