## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear independence)
## Institution(TCNJ)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('2.3')
## Problem1('')
## KEYWORDS('vectors' 'basis' 'linearly independent' 'dependent')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "PGnumericalmacros.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

$lambda = new Matrix(3,1);
$a=non_zero_random(1,5,1);
$b=non_zero_random(1,5,1);
$c=non_zero_random(1,5,1);

# we'll scramble the vectors and coefficients

## $vec[$slice[0]] = Matrix->new_column_matrix([$v01, $v02, $v03, $v04]);


##$vec[1] = Matrix->new_column_matrix([ $v11, $v12, $v13, $v14 ]);

$lambda->assign(1 ,1,-$a);
$lambda->assign(2 ,1,-$b);
$lambda->assign(3 ,1,1);

$sl = new_pop_up_select_list();
$sl->qa(
"Is \(T\) linearly independent or dependent?",
'Linearly_Independent');
$sl->choose(1);
$sl->ra_pop_up_list(['No_Answer_Entered' => 'Select an Answer', 'Linearly_Independent' => 'Linearly Independent', 'Linearly_Dependent' => 'Linearly Dependent']);

$string1=nicestring([$a, $b, $c], ['r','u','d']);

$coef = MultiAnswer( Compute(0), Compute(0), Compute(0) )->with(
        singleResult=>1,
        checker=> sub {
	    my ( $c, $s, $ans ) = @_;
	    return $c->[0]==$s->[0] && $c->[1]==$s->[1] && $c->[2]==$s->[2];
	} );

BEGIN_TEXT
Suppose \(  S= \left\{"\{"\} r, u, d\right\} \) is a set of linearly independent vectors.  
$BR $BR
If \(x = $string1 \),  determine whether \(T=
\left\{"\{"\} r,u,x\right\} \) is a linearly independent set.

\{$sl->print_q()\} 
$BR
$BR
If \(T\) is dependent, enter a non-trivial linear relation below.  Otherwise, enter 0's for the coefficients.
$BR
$BR

\{ $coef->ans_rule(10) \} \( r + \)
\{ $coef->ans_rule(10) \} \( u + \)
\{ $coef->ans_rule(10) \} \( x = 0\).

END_TEXT

ANS( str_cmp($sl->ra_correct_ans) );
ANS( $coef->cmp() );

ENDDOCUMENT();
