## DESCRIPTION
## Rational Expressions: Addition and Subtraction
## ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(Spring 2012)
## Institution(University of Missouri)
## MLT(SharedFactor)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.6')
## Problem1('')
## KEYWORDS('Algebraic fractions')
## Adapted from a problem written by Paul Pearson (Fort Lewis)


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

###################################
# Setup

Context("Numeric");

Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real");

$var = list_random('x','y','z','w','p','q','r','s','t','u','v');

$a = 2*list_random(-5,-4,-3,-2,2,3,4,5);
$b = random(-1,1,2)*list_random(2,3,4,5,6,7);
$c = $a/2;

if ($a + $b == 0) {
  $den1 = "$var^2 +".($a*$b); #($var + $a)($var + $b)
} elsif ($a + $b == -1) {
  $den1 = "$var^2 - $var +".($a*$b);
} elsif ($a + $b == 1) {
  $den1 = "$var^2 + $var +".($a*$b);
} else {
  $den1 = "$var^2 +".($a+$b)."$var +".($a*$b); #($var + $a)($var + $b)
}

if ($c + $b == 0) {
  $den2 = "$var^2 +".($c*$b); #($var + $a)($var + $b)
} elsif ($c + $b == -1) {
  $den2 = "$var^2 - $var +".($c*$b);
} elsif ($c + $b == 1) {
  $den2 = "$var^2 + $var +".($c*$b);
} else {
  $den2 = "$var^2 +".($b+$c)."$var +".($c*$b); #($var + $b)($var + $c)
}

$frac = "\frac{$var + $c}{$den1} - \frac{$var}{$den2}";

$num = Formula(" ($c)**2")->reduce;
$den = Formula("($var + $a)($var + $b)($var +$c)   ")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}



###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Perform the indicated operation, and simplify if possible.  Assume any factors you cancel are not zero.
$BR
$BR
$BCENTER
$displayfrac
$ECENTER

END_TEXT


Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 0;

ANS($multians->cmp());

ENDDOCUMENT();
