## DESCRIPTION
## Find a value of x for which a 3x3 matrix, depending on x, has a given determinant.
##   The matrix will have small integer entries.
##   There will usually (always?) only be one possible solution. If multiple solutions work, any of them will be accepted.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Determinants')
## DBsection('Computing determinants')
## Date('March 25, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('3.2: Determinants')
## Problem1('')
## KEYWORDS('linear algebra','matrices','determinant')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "contextFraction.pl",
    "PGchoicemacros.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Matrix");

do {
    $A = Matrix([[random(-2,3),random(-2,3),random(-2,3)],[random(-2,3),random(-2,3),random(-2,3)],[random(-2,3),random(-2,3),random(-2,3)]]);
    
    $r = random(1,3);
    $c = random(1,3);
    
    $ans = $A->element($r,$c);
    
    foreach $i  (0..2) {
        foreach $j (0..2) {
            $em[$i][$j] = 0;
        }
    }
    $em[$r-1][$c-1] = 1;

    $Em = Matrix(@em);

    $D = $A + $Em;
    $F = $A - $Em;

    $dt = $A->det;
} until (abs($D->det - $dt) > 0.1 && abs($F->det - $dt) > 0.1);

foreach $i  (0..2) {
    foreach $j (0..2) {
        if($i == $r-1 && $j == $c-1){
            $ad[$i][$j] = "x";
        }else{
            $ad[$i][$j] = $A->element($i+1,$j+1);
        }
    }
}

   
##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a real number \( x \) for which the matrix \( \displaystyle \begin{bmatrix} $ad[0][0] & $ad[0][1] & $ad[0][2] \\ $ad[1][0] & $ad[1][1] & $ad[1][2] \\ $ad[2][0] & $ad[2][1] & $ad[2][2]\end{bmatrix} \) has determinant equal to \( $dt \).
$BR
$BR
\( x \): \{ ans_rule(3) \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $ans->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    
    foreach $i  (0..2) {
        foreach $j (0..2) {
            if($i == $r-1 && $j == $c-1){
                $ansm[$i][$j] = $student;
            }else{
                $ansm[$i][$j] = $A->element($i+1,$j+1);
            }
        }
    }
    $AnsMat = Matrix(@ansm);
    
    if( $AnsMat->det == $dt )
        { return 1; } else { return 0; }
}));

ENDDOCUMENT();
