# DESCRIPTION
# Problem from Holt Linear Algebra, 1st ed.
# WeBWorK problem written by Gavin LaRose,
# <glarose(at)umich(dot)edu>
# ENDDESCRIPTION

## DBsubject('Linear Algebra')
## DBchapter('Matrices')
## DBsection('Linear Transformations')
## KEYWORDS('linear transformations, range')
## TitleText1('Linear Algebra')
## EditionText1('1')
## AuthorText1('Holt')
## Section1('3.1')
## Problem1('40')
## Author('Gavin LaRose')
## Institution('UMich')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserPopUp.pl",
# "compoundProblem.pl",
);

# make sure we're in the context we want
Context("Matrix");
$showPartialCorrectAnswers = 0;

$mat = Matrix( [[1,1]] );
$tfPop = PopUp( [ '?', 'True', 'False' ], 'False' );

TEXT(beginproblem());

Context()->texStrings;
BEGIN_TEXT
${BITALIC}(Note that this is a compound problem, for which the
second part of the problem is shown after the first is answered
correctly.)$EITALIC
$PAR

${BBOLD}(a)$EBOLD
Determine if the following is true or false:
$PAR
The image of a linear transformation \(T\) must be a subset of the
domain.
$PAR
This statement is \{ $tfPop->menu() \}
$PAR

END_TEXT
Context()->normalStrings;

ANS( $tfPop->cmp() );

$ans_hash1 = $tfPop->cmp()->evaluate( $inputs_ref->{ANS_NUM_TO_NAME(1)} );

if ( $ans_hash1->{score} == 1 ) {

Context()->texStrings;
BEGIN_TEXT

${BBOLD}(b)$EBOLD
Let \( T(\mathbf{x}) = A\mathbf{x} \).
Give a matrix \(A\) for which the range of
the linear transformation \(T\) is not a subset of the domain.
$BR
\( A = \) \{ $mat->ans_rule(25) \}
$BR
${BITALIC}(Enter your answer delimiting the matrix and
rows with square brackets: thus, to enter
\(\begin{bmatrix} 1 & 2\\ 3 & 4\end{bmatrix}\) you would enter
${BBOLD}[[1,2],[3,4]]$EBOLD.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $mat->cmp( checker=>sub {
		   my ($cor, $stu, $ans) = @_;
		   my @dim = $stu->dimensions;
		   return $dim[0] != $dim[1]; } ) );
}

Context()->texStrings;
BEGIN_SOLUTION

This is False.  For instance, let \(A = $mat\); then
\(T\left(\begin{bmatrix} x_1\\ x_2\end{bmatrix}\right) = x_1 + x_2\)
has \(\mbox{range}(T) = \mathbf{R}\), which is not a subset of
\(\mbox{domain}(T) = \mathbf{R}^2\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
