##DESCRIPTION
##  College Algebra Review: Factoring quadratics with leading term 1
##ENDDESCRIPTION

##KEYWORDS('algebra', 'factoring', 'quadratic')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Expanding and Factoring')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.4')
## Problem1('')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PeriodicRerandomization.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(x=>"Real",y=>"Real",t=>"Real");

$a = list_random(-12,-11,-10,-8,-7-6,-5,-3,-2,2,3,5,6,7,8,10,11,12);
do {$c = list_random(2,3,4,5,6,7,8,9,10,11,12)} 
until (gcd($a, $c ) == 1);
$e1 = $c**2 ;
$e3 = -$a**2;
$texanswer = "\( \left($c x + $a t \right) \left($c y - $a t \right)\)";
$fac1 = Formula(" $c x^2 + $a t ");
$fac2 = Formula(" $c x^2 - $a t ")->reduce;

$expression = "\( $e1 x^4  + $e3 t^2 \)";

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
#  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ||
           ($f1 == -$f1stu && $f2 == - $f2stu) ||
           ($f1 == -$f2stu && $f2 == -$f1stu)  ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text

BEGIN_TEXT

Factor the expression $expression.  Simplify your 
answer as much as possible.
$BR
$BR
\( \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \)
END_TEXT


if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    $texanswer
    $PAR
    END_TEXT
} else {
  BEGIN_TEXT
  $BR
  END_TEXT
}

Context()->normalStrings;

PeriodicStatus(); 

################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

COMMENT('MathObject version', 'Features PeriodicRerandomization');
ENDDOCUMENT();
