# ************************************
#   Comments + Keyword Clarification
#
#    Anton, Calculus, 10.6.6a & b & c
# ************************************
# DESCRIPTION
# Find the polar equation of a conic given information about vertices, eccentricity, etc.
#
# ENDDESCRIPTION
## DBsubject('Calculus')
## DBchapter('Parametric and Polar Curves; Conic Sections')
## DBsection('Conic Sections in Polar Coordinates')
## KEYWORDS('Conics', 'Polar coordinates')
## TitleText1('Calculus')
## EditionText1('9')
## AuthorText1('Anton')
## Section1('10.6')
## Problem1('6abc')
## Author('')
## Institution('John Wiley & sons')
#***************************************************************
#   Initialization - Load standard PG and MathObjects macros   
#***************************************************************
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "littleneck.pl",
  "PGcourse.pl"
);
# allow the student to change the seed for this problem.

rand_button();

#
# Set Up - Set up variable context and formulation of problem
#
Context("Numeric");
Context()->variables->add(t=>'Real'); 

#****************************************************
#  Create polar equation of an ellipse with axis parallel to either 
#one of the coordinate axes.
#****************************************************




#Create an array for the four choices


@fn = ( 
  ['sin', '+' , '\pi/2', '3\pi/2', 'above', 'less than'],
  ['sin', '-' , '\pi/2', '3\pi/2', 'below', 'greater than'],
  ['cos', '+' , '0', '\pi', 'right of', 'less than'],
  ['cos', '-' , '0', '\pi', 'left of', 'greater than'],
);

## choose one of the orientations

# 

($func,$sign,$t1,$t2,$orient,$r1r2)=@{$fn[random(0,$#fn,1)]};

# create some random numbers

$k = random(2,3,1);
if($k==2) {$small = random(2,5,1);} else {$small=random(2,6,2);};
$big = $small*($k+1)/($k-1);
$s2 = $s1*($k+1)/($k-1);
if($r1r2 eq 'less than') {$r1 = $small; $r2 = $big;} else {$r1=$big; $r2 = $small;};

$e = 1/$k;
$a = ($big +$small)/2;
$top = 2*$r1*$r2/($big - $small);

#data for part b
#

if ($func eq 'sin') {$func_b = 'cos'; $t_b = '\pi'} else {$func_b = 'sin'; $t_b = '\pi/2';};
$r_b = random(2,6,1);
$top_b = 2*$r_b;

#data for part c
#

$e_c = random(2,5,1);
$r_c = random(2,8,1);
while ($e_c == $r_c) {$r_c = random(2,8,1);}
$top_c = $r_c*(1+$e_c);
$top_c2 = $r_c*(1-$e_c);


#**************
#   Answers
#**************

Context()->texStrings;

$ans_a = Formula("$top/($k $sign $func(t))")->reduce;
$ans_b = Formula("$top_b/(1 + $func_b(t))")->reduce;
$ans_c = Formula("$top_c/(1 + $e_c cos(t))")->reduce;
$ans_c2 = Formula("$top_c2/(1 - $e_c cos(t))")->reduce;

# *********************
#   Problem Statement
# *********************
TEXT(beginproblem());
#
# Strings in Tex form (formula)
#
#Context()->texStrings;
BEGIN_TEXT

Find a polar equation for the conic described below. Points are in polar coordinates and directrices in rectangular coordinates for simplicity. 
$BR (In some cases there may be more than one conic that satisfies the conditions. Enter the equation of one of the conics as your answer in this case).
$BR Enter t for the variable \( \theta \).
$PAR

(a) Ellipse; ends of major axis \( ($r1, $t1) \) and \( ($r2, $t2) \).
$BR
Equation of the ellipse is  \(r= \) \{ ans_rule(35) \} 
$PAR

(b) Parabola; vertex \( ($r_b, $t_b) \).
$BR
Equation of the parabola is \(r= \) \{ ans_rule(35) \}
$PAR

(c) Hyperbola; \( e = $e_c \); vertex \( ($r_c, 0) \).
$BR
Equation of the hyperbola is \(r= \) \{ ans_rule(35) \}

 
END_TEXT
#
# Expect strings in standard form
#
Context()->normalStrings;
# *******************************
#   Answer and Solution section
# *******************************
#
# Check student solution
#
ANS($ans_a->cmp());
ANS($ans_b->cmp());
ANS($ans_c->cmp(checker => sub {
       my ($correct,$student) = @_;
       return $ans_c == $student || $ans_c2 == $student;
     }));
#ANS($ans_c->cmp());
#***************************
#  Explanation of solution  
#***************************
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

(a) Since the vertices are at \( ($r1, $t1) \) and \( ($r2, $t2) \), we know that the directrix is $orient the pole. Hence, the equation is of the form


\[ r = \frac {a(1-e^2)}{1 $sign e \$func(\theta)} \]

Since, \( r_0 = \) $small and \( r_1 = \) $big, we have \( e = \frac{1}{$k} \) and \( a = \) $a. Hence, we can write the equation as

\( r = \frac {\frac{$top}{$k}}{1 $sign \frac{1}{$k} \$func(\theta)} = \frac {$top}{$k $sign \$func(\theta)} \)

$PAR

(b) With the vertex at \( ($r_b, $t_b) \), the equation must be of the form \( r = \frac {d}{1 + \$func_b(\theta)} \).
$BR Since \( ($r_b, $t_b) \) satisfies the equation, \( d= $top_b \) and equation is \( r = \frac {$top_b}{1 + \$func_b(\theta)} \).
$PAR

(c) With a vertex at \( ($r_c, 0) \), the conic must be of the form \( r = \frac {ed}{1 + e\cos(\theta)} \) or \( r = \frac {ed}{1 - e\cos(\theta)} \).
$BR
If we substitute $ed for \( e \) and the coordinates of the vertex into the equation, we get

\( r = \frac {$top_c}{1 + $e_c \cos(\theta)} \)  or \( r = \frac {$top_c2}{1 - $e_c \cos(\theta)} \)
$BR
Note: The second equation is valid when the given vertex of the hyperbola is associated with the focus not at the pole.



END_SOLUTION
Context()->normalStrings;
ENDDOCUMENT();
