#Use laws of exponents with rational exponents.
# ENDDESCRIPTION
## DBsubject()
## DBchapter()
## DBsection()
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## Date(6/2019)
## KEYWORDS()

 
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
  "contextFraction.pl",
);

######################################
TEXT(beginproblem());
Context("Fraction");
$var = list_random('a','b','c','m','n','p','y','z');
Context()->variables->add($var=>'Real');
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);
#Context()->flags->set(bizarroPow=>1);
$n1 = random(1, 9, 2);
do{
$n2 = random(3, 11, 2);
}until($n2 > $n1);
$n3 = random(1, 10);

$d1 = 2;
$d2 = 2;
$d3 = 2;

$e1 = Fraction($n1, $d1)->reduce;
$e2 = Fraction($n2, $d2)->reduce;
$e3 = Fraction($n3, $d3)->reduce;

$diff = ($n1-$n2)/2;

$e4 = Fraction($n3 - $n1 + $n2, $d1)->reduce;

#Answer
$ga = Formula("1/$var^($e4)" );

#####################################################

BEGIN_PGML
Simplify by applying the laws of exponents. Write your answers with positive rational exponents only.


[``\frac{{[$var]}^{[$e1]}{[$var]}^{-[$e2]}}{{[$var]}^{[$e3]}} =``][__________]


END_PGML
####################################################

ANS($ga -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

#################################
BEGIN_PGML_SOLUTION
[``\frac{{[$var]}^{[$e1]}{[$var]}^{-[$e2]}}{{[$var]}^{[$e3]}}``]

[``\frac{[$var]^{[$e1]-[$e2]}}{[$var]^{[$e3]}}``]

[``\frac{[$var]^{[$diff]}}{[$var]^{[$e3]}}``]

[``[$var]^{[$diff]-[$e3]}``]

[``[$var]^{-[$e4]}``]

[``[$ga]``]


END_PGML_SOLUTION

ENDDOCUMENT();