#Use laws of exponents with rational exponents.
# ENDDESCRIPTION
## DBsubject()
## DBchapter()
## DBsection()
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## Date(6/2019)
## KEYWORDS()

 
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
  "contextFraction.pl",
);

######################################
TEXT(beginproblem());
Context("Fraction");
$var = list_random('a','b','c','m','n','p','y','z');
Context()->variables->add($var=>'Real');
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);
#Context()->flags->set(bizarroPow=>1);

$d1 = list_random(3, 4);
$n1 = list_random(1, 2, 5);
$d2 = random(2, 8);
do{
$n2 = random(1, 8);
}until($n2 != $d2);

$e1 = Fraction($n1, $d1)->reduce;
$e2 = Fraction($n2, $d2)->reduce;

$e3 = Fraction($n1*$n2, $d1*$d2)->reduce;

#for solution
$newn = -$n1*$n2;
$newd = $d1*$d2;

#Answer
$ga = Formula("1/$var^($e3)" );

#####################################################

BEGIN_PGML
Simplify by applying the laws of exponents. Write your answers with positive rational exponents only.


[``({{[$var]}^{-[$e1]}})^{[$e2]} =``][__________]


END_PGML
####################################################

ANS($ga -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

########################################################
BEGIN_PGML_SOLUTION
Use the power rule and multiply the exponents:

[`[$e1]\cdot[$e2] = \frac{[$newn]}{[$newd]} = -[$e3]`]


[`[$var]^{-[$e3]} = \frac{1}{[$var]^[$e3]}`]


END_PGML_SOLUTION



ENDDOCUMENT();