## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.2')
## Problem1('4')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "PGunion.pl",
  "parserPopUp.pl",
  "contextInequalities.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

install_weighted_grader();

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");

$a = random(2,6,1);
do { $b = random(2,6,1); } until ($b != $a);

do { $a11 = random(2,9,1); } until ($a11 != $a && $a11 != $b);
do { $a12 = random(2,9,1); } until ($a12 != $a && $a12 != $b);
do { $a21 = random(2,9,1); } until ($a21 != $a && $a21 != $b);
do { $a22 = random(2,9,1); } until ($a22 != $a && $a22 != $b);

# In order to avoid getting errors, make sure all of these
# are formulas, e.g., $P12 is a Formula instead of Real or Compute.
#
$Q11 = Formula("$a11 t");  $Q12 = Formula("$a12");
$Q21 = Formula("$a21");    $Q22 = Formula("$a22 t");

$P11 = Formula("$a11 t / (t+$a)");  $P12 = Formula("$a12 / (t+$a)");
$P21 = Formula("$a21 / (t-$b)");    $P22 = Formula("$a22 t / (t-$b)");

$g1 = Formula("0");
$g2 = Formula("0");

$multians1 = MultiAnswer($P11, $P12, $P21, $P22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);

$multians2 = MultiAnswer($g1, $g2)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  format =>  "<table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c1, $c2 ) = @{$correct};
      my ( $s1, $s2 ) = @{$student};
      my @score = (0,0);
      if ($c1->typeMatch($s1) && $c1 == $s1) { $score[0] = 1; }
      if ($c2->typeMatch($s2) && $c2 == $s2) { $score[1] = 1; }
      return [ @score ];
  }

);

$popup = PopUp(["Choose","homogeneous","nonhomogeneous"],"homogeneous");


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose
$PAR
\[
\begin{array}{rclcl}
(t + $a) y_1^{\,\prime} & = & $Q11 y_1 + $Q12 y_2, && y_1(1) = 0, \\
(t - $b) y_2^{\,\prime} & = & $Q21 y_1 + $Q22 y_2, && y_2(1) = 2. 
\end{array}
\]
$BR

\{ BeginList("OL",type=>'a') \}
$ITEM
This system of linear differential equations can
be put in the form 
\( \boldsymbol{\vec{y}^{\,\prime}} = P(t) \boldsymbol{\vec{y}} + \boldsymbol{\vec{g}}(t) \).
Determine \( P(t) \) and \( \boldsymbol{\vec{g}}(t) \).

$PAR
$BCENTER
\{ 
mbox(
  "\( P(t) = \)",
  display_matrix([
  [$multians1->ans_rule(15),$multians1->ans_rule(15)], 
  [$multians1->ans_rule(15),$multians1->ans_rule(15)]
  ], align=>'cc')
);
\}
$ECENTER

$PAR
$BCENTER
\{ 
mbox(
  "\( \boldsymbol{\vec{g}}(t) = \)",
  display_matrix([
  [$multians2->ans_rule(15)], 
  [$multians2->ans_rule(15)]
  ], align=>'cc')
);
\}
$ECENTER

$ITEMSEP
$ITEM
Is the system homogeneous or nonhomogeneous? 
\{ $popup->menu() \}

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 50 );
WEIGHTED_ANS( $multians2->cmp(), 10 );

WEIGHTED_ANS( $popup->cmp(), 5 );


Context("Inequalities")->variables->are(t=>"Real");

$interval = Compute("-$a < t < $b");

Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM
Find the largest interval  \( a < t < b \) 
such that a unique solution of the initial 
value problem is guaranteed to exist.
$BR
$BR
Interval: 
\{ ans_rule(30) \}
\{ AnswerFormatHelp("inequalities") \}

\{ EndList("OL") \}

END_TEXT
Context()->normalStrings;



WEIGHTED_ANS( $interval->cmp(), 35 );

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
