## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.2')
## Problem1('19')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

install_weighted_grader();


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real",
);

$a = random(2,8,1);
do { $b = random(2,5,1); } until ($b != $a);
$c = random(-8,-2,1);

foreach my $i (0..3) {
  foreach my $j (0..3) {
  $P[$i][$j] = Formula("0");
  }
}

$P[0][1] = Formula("1");
$P[1][0] = Formula("$a");
$P[1][1] = Formula("1/t");
$P[1][2] = Formula("-t");
$P[1][3] = Formula("sin(t)");
$P[2][3] = Formula("1");
$P[3][0] = Formula("1");
$P[3][3] = Formula("$c");

$multians1 = MultiAnswer(
$P[0][0], $P[0][1], $P[0][2], $P[0][3],
$P[1][0], $P[1][1], $P[1][2], $P[1][3],
$P[2][0], $P[2][1], $P[2][2], $P[2][3],
$P[3][0], $P[3][1], $P[3][2], $P[3][3]
)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rrrr} 
  %s & %s & %s & %s \\ 
  %s & %s & %s & %s \\ 
  %s & %s & %s & %s \\ 
  %s & %s & %s & %s \\ 
  \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..15) {
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j]=1; } else { $score[$j]=0; }
      }
      return [ @score ];
  }

);


foreach my $i (0..3) {
  $g[$i] = Formula("0");
}
$g[1] = Formula("e^($b t)");

$multians2 = MultiAnswer($g[0], $g[1], $g[2], $g[3])->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{r} %s \\ %s \\ %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c1, $c2, $c3, $c4 ) = @{$correct};
      my ( $s1, $s2, $s3, $s4 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c1->typeMatch($s1) && $c1 == $s1) { $score[0] = 1; }
      if ($c2->typeMatch($s2) && $c2 == $s2) { $score[1] = 1; }
      if ($c3->typeMatch($s3) && $c3 == $s3) { $score[2] = 1; }
      if ($c4->typeMatch($s4) && $c4 == $s4) { $score[3] = 1; }
      return [ @score ];
  }

);



######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Consider the system of higher order differential equations
$PAR
\[
\begin{array}{rcl}
y^{\,\prime\prime} & = & t^{-1} y^{\,\prime} + $a y - t z + (\sin t) z^{\,\prime} + e^{$b t}, \\
z^{\,\prime\prime} & = & y + $c z^{\,\prime}.
\end{array}
\]
$BR
Rewrite the given system of two second order differential
equations as a system of four first order linear differential
equations of the form \( \boldsymbol{\vec{y}^{\,\prime}} = P(t) \boldsymbol{\vec{y}} +
\boldsymbol{\vec{g}}(t) \).  Use the following change of variables
$PAR
\[
\boldsymbol{\vec{y}}(t) 
= 
\left\lbrack
\begin{array}{r}
y_1(t) \\
y_2(t) \\
y_3(t) \\
y_4(t)
\end{array}
\right\rbrack
=
\left\lbrack
\begin{array}{r}
y(t) \\
y^{\,\prime}(t) \\
z(t) \\
z^{\,\prime}(t)
\end{array}
\right\rbrack.
\]
$BR
$PAR
$BCENTER
\{ 
mbox(
  display_matrix([["\(y_1^{\,\prime}\)"],["\(y_2^{\,\prime}\)"],["\(y_3^{\,\prime}\)"],["\(y_4^{\,\prime}\)"]], align=>'c'),
  ' =',
  display_matrix([
  [$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6)],
  [$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6)],
  [$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6)],
  [$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6),$multians1->ans_rule(6)]
  ], align=>'cccc'),
  display_matrix([["\(y_1\)"],["\(y_2\)"],["\(y_3\)"],["\(y_4\)"]], align=>'c'),
  ' + ',
  display_matrix([
  [$multians2->ans_rule(6)],[$multians2->ans_rule(6)],[$multians2->ans_rule(6)],[$multians2->ans_rule(6)]
  ], align=>'c')
);
\}
$ECENTER

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 80 );
WEIGHTED_ANS( $multians2->cmp(), 20 );

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
