## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Distinct real eigenvalues)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## Static(1)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.2')
## Problem1('7')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real",
'c1'=>"Real",
'c2'=>"Real",
);
Context()->reduction->set('(-x)+y'=>0);
Context()->reduction->set('(-x)-y'=>0);

$i = random(0,5,1);

#$i = ($i + 2) % 6;

if ($i == 0) {

  @matrix = (1,3,3,1);
  @evect1 = (1,1);
  $eval1  = 4;
  @evect2 = (1,-1);
  $eval2  = -2;

} elsif ($i == 1) {

  @matrix = (7,2,-4,1);
  @evect1 = (1,-1);
  $eval1  = 5;
  @evect2 = (1,-2);
  $eval2  = 3;

} elsif ($i == 2) {

  @matrix = (4,1,1,4);
  @evect1 = (1,1);
  $eval1  = 5;
  @evect2 = (1,-1);
  $eval2  = 3;

} elsif ($i == 3) {

  @matrix = (1,-2,1,4);
  @evect1 = (2,-1);
  $eval1  = 2;
  @evect2 = (1,-1);
  $eval2  = 3;

} elsif ($i == 4) {

  @matrix = (3,2,4,1);
  @evect1 = (1,1);
  $eval1  = 5;
  @evect2 = (-1,2);
  $eval2  = -1;

} else {

  @matrix = (5,3,3,5);
  @evect1 = (1,-1);
  $eval1  = 2;
  @evect2 = (1,1);
  $eval2  = 8;
}

$y1 = Formula("$evect1[0] c1 e^($eval1 t) + $evect2[0] c2 e^($eval2 t)")->reduce();
$y2 = Formula("$evect1[1] c1 e^($eval1 t) + $evect2[1] c2 e^($eval2 t)")->reduce();

$P11 = Formula("$evect1[0] e^($eval1 t)");
$P21 = Formula("$evect1[1] e^($eval1 t)");

$P12 = Formula("$evect2[0] e^($eval2 t)");
$P22 = Formula("$evect2[1] e^($eval2 t)");

$multians1 = MultiAnswer($P11, $P21, $P12, $P22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0'><tr>
    <td><table border='0' cellspacing='10'><tr><td> %s </td></tr><tr><td> %s</td></tr></table></td>
    <td><table border='0' cellspacing='10'><tr><td> %s </td></tr><tr><td> %s</td></tr></table></td>
  </tr></table>",
  tex_format => "\left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack, 
  \left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $i (0..3) {
        if ($c[$i]->typeMatch($s[$i]) && $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; }
      }
      return [ @score ];
  }

);

foreach my $j (0..3) {
  if ($matrix[$j] == 1) { $m[$j] = ""; } else { $m[$j] = $matrix[$j]; }
}


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
$BITALIC This is the third part of a three-part problem. $EITALIC
$BR
$BR
Consider the system of differential equations
$PAR
\[
\begin{array}{rcl}
y_1^{\,\prime} & = & $m[0] y_1 + $m[1] y_2, \\
y_2^{\,\prime} & = & $m[2] y_1 + $m[3] y_2, 
\end{array}
\]
$BR
with solutions
$PAR
\[
\begin{array}{rcl}
y_1(t) & = & $y1, \\
y_2(t) & = & $y2, 
\end{array}
\]
$BR
for any constants \( c_1 \) and \( c_2 \).

Rewrite the solution of the equations in vector 
form as  \( \boldsymbol{\vec{y}}(t) = c_1 \boldsymbol{\vec{y}_1}(t) + 
c_2 \boldsymbol{\vec{y}_2}(t) \).
$PAR
$BCENTER
\{ 
mbox(
  "\( \boldsymbol{\vec{y}}(t) = c_1 \)",
  display_matrix([
  [$multians1->ans_rule(15)],[$multians1->ans_rule(15)] 
  ], align=>'c'),
  "\( + \ c_2 \)",
  display_matrix([
  [$multians1->ans_rule(15)],[$multians1->ans_rule(15)] 
  ], align=>'c')
);
\}
$ECENTER

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );


COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
