## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Complex eigenvalues)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.3')
## Problem1('10')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserPopUp.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);

install_weighted_grader();


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real",
);

$k = random(2,9,1);

@P = (
[0,$k],
[-($k),0]
);

@y1 = ( Formula("cos($k t)"), Formula("-sin($k t)") ); 
@y2 = ( Formula("-$k sin($k t)"), Formula("-$k cos($k t)") ); 

$wronskian = Formula("-$k");


$pop0 = PopUp(["Choose","dependent","independent"],"independent");
$pop1 = PopUp(["Choose","zero","nonzero"],"nonzero");
$pop2 = PopUp(["Choose","do","do not"],"do");


$multians1 = MultiAnswer($y1[0], $y2[0], $y1[1], $y2[1])->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
$BITALIC This is the second part of a two-part problem. $EITALIC
$BR
$BR
Let
$PAR
\[
\boldsymbol{\vec{y}_1}(t) =
\left\lbrack
\begin{array}{c}
$y1[0] \\
$y1[1]
\end{array}
\right\rbrack,
\ \ \
\boldsymbol{\vec{y}_2}(t) =
\left\lbrack
\begin{array}{c}
$y2[0] \\
$y2[1]
\end{array}
\right\rbrack.
\]
$BR

Compute the Wronskian to determine whether the functions
\( \boldsymbol{\vec{y}_1}(t) \) and \( \boldsymbol{\vec{y}_2}(t) \)
are linearly independent.
$PAR
$BCENTER
\{ 
mbox(
  "Wronskian = \( \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  " = ".$SPACE.ans_rule(20)
);
\}
$ECENTER
$PAR
These functions are linearly \{ $pop0->menu() \} because
the Wronskian is \{ $pop1->menu() \} for all \(t\).
Therefore, the solutions \( \boldsymbol{\vec{y}_1}(t) \) and 
\( \boldsymbol{\vec{y}_2}(t) \) to the system
$PAR
\[
\boldsymbol{\vec{y}^{\,\prime}} =
\left\lbrack
\begin{array}{rr}
$P[0][0] & $P[0][1] \\
$P[1][0] & $P[1][1]
\end{array}
\right\rbrack
\boldsymbol{\vec{y}}
\]
$BR
\{ $pop2->menu() \}
form a fundamental set (i.e., linearly independent set) of solutions.

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 50 );

WEIGHTED_ANS( $wronskian->cmp(), 40 );

WEIGHTED_ANS( $pop0->cmp(), 4 );
WEIGHTED_ANS( $pop1->cmp(), 3 );
WEIGHTED_ANS( $pop2->cmp(), 3 );

COMMENT("Uses parserMultiAnswer.pl to format the matrix nicely.");

ENDDOCUMENT();
