## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-MN-setup_only)
## Level(2)
## Static(1)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.2')
## Problem1('7')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real",
'c1'=>"Real",
'c2'=>"Real",
);
Context()->reduction->set('(-x)+y'=>0);
Context()->reduction->set('(-x)-y'=>0);

$i = random(0,5,1);

#$i = ($i + 1) % 6;

if ($i == 0) {

  @matrix = (1,3,3,1);
  @evect1 = (1,1);
  $eval1  = 4;
  @evect2 = (1,-1);
  $eval2  = -2;

} elsif ($i == 1) {

  @matrix = (7,2,-4,1);
  @evect1 = (1,-1);
  $eval1  = 5;
  @evect2 = (1,-2);
  $eval2  = 3;

} elsif ($i == 2) {

  @matrix = (4,1,1,4);
  @evect1 = (1,1);
  $eval1  = 5;
  @evect2 = (1,-1);
  $eval2  = 3;

} elsif ($i == 3) {

  @matrix = (1,-2,1,4);
  @evect1 = (2,-1);
  $eval1  = 2;
  @evect2 = (1,-1);
  $eval2  = 3;

} elsif ($i == 4) {

  @matrix = (3,2,4,1);
  @evect1 = (1,1);
  $eval1  = 5;
  @evect2 = (-1,2);
  $eval2  = -1;

} else {

  @matrix = (5,3,3,5);
  @evect1 = (1,-1);
  $eval1  = 2;
  @evect2 = (1,1);
  $eval2  = 8;
}

$y1 = Formula("$evect1[0] c1 e^($eval1 t) + $evect2[0] c2 e^($eval2 t)")->reduce();
$y2 = Formula("$evect1[1] c1 e^($eval1 t) + $evect2[1] c2 e^($eval2 t)")->reduce();

($P11,$P12,$P21,$P22) = @matrix;

$multians1 = MultiAnswer($P11, $P12, $P21, $P22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td></tr><tr><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);


foreach my $j (0..3) {
  if ($matrix[$j] == 1) { $m[$j] = ""; } else { $m[$j] = $matrix[$j]; }
}


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
$BITALIC This is the first part of a three-part problem. $EITALIC
$BR
$BR
Consider the system of differential equations
$PAR
\[
\begin{array}{rcl}
y_1^{\,\prime} & = & $m[0] y_1 + $m[1] y_2, \\
y_2^{\,\prime} & = & $m[2] y_1 + $m[3] y_2. 
\end{array}
\]
$BR
Rewrite the equations in vector form as 
\( \boldsymbol{\vec{y}^{\,\prime}}(t) = A \boldsymbol{\vec{y}}(t) \).
$PAR
$BCENTER
\{ 
mbox(
  "\( \boldsymbol{\vec{y}^{\,\prime}}(t) = \)",
  display_matrix([
  [$multians1->ans_rule(15),$multians1->ans_rule(15)], 
  [$multians1->ans_rule(15),$multians1->ans_rule(15)]
  ], align=>'cc'),
  "\( \boldsymbol{\vec{y}}(t) \)"
);
\}
$ECENTER

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
