## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Verification of solutions)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.3')
## Problem1('5')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real",
);

@P = (
[2,1,1],
[1,1,2],
[1,2,1]
);

$row1 = "2y_1 +  y_2 +  y_3";
$row2 = " y_1 +  y_2 + 2y_3";
$row3 = " y_1 + 2y_2 +  y_3";

$i = random(0,2,1);

if ($i == 0) {

  @evect = (2,-1,-1);
  $eval  = 1;

} elsif ($i == 1) {

  @evect = (0,1,-1);
  $eval  = -1;

} else {

  @evect = (1,1,1);
  $eval  = 4;

}

$y1 = Formula("$evect[0] e^($eval t)")->reduce();
$y2 = Formula("$evect[1] e^($eval t)")->reduce();
$y3 = Formula("$evect[2] e^($eval t)")->reduce();

$y1p = $y1->D('t');
$y2p = $y2->D('t');
$y3p = $y3->D('t');



$multians1 = MultiAnswer(
$y1p, 2*$y1, $y2, $y3
)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "%s = (%s) + (%s) + (%s)",
  tex_format =>  "%s = (%s) + (%s) + (%s)",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = (0,0,0,0);
      foreach my $j (0..3) {
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians2 = MultiAnswer(
$y2p, $y1, $y2, 2*$y3
)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "%s = (%s) + (%s) + (%s)",
  tex_format =>  "%s = (%s) + (%s) + (%s)",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = (0,0,0,0);
      foreach my $j (0..3) {
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians3 = MultiAnswer(
$y3p, $y1, 2*$y2, $y3
)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "%s = (%s) + (%s) + (%s)",
  tex_format =>  "%s = (%s) + (%s) + (%s)",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = (0,0,0,0);
      foreach my $j (0..3) {
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);



######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
$BITALIC This is the second part of a two-part problem. $EITALIC
$BR
$BR
Show that 
$PAR
\[
\boldsymbol{\vec{y}}(t) =
\left\lbrack
\begin{array}{r}
$y1 \\
$y2 \\
$y3
\end{array}
\right\rbrack
\]
$BR
is a solution to the system of linear homogeneous 
differential equations 
$PAR
\[
\begin{array}{rcl}
y_1^{\,\prime} & = & $row1, \\
y_2^{\,\prime} & = & $row2, \\
y_3^{\,\prime} & = & $row3.
\end{array}
\]
$BR

\{ BeginList("OL",type=>'a') \}

$ITEM
Find the value of each term in the equation
\( y_1^{\,\prime} = $row1 \) in terms of the variable \( t \).
(Enter the terms in the order given.)
$BR
$BR
\{ $multians1->ans_rule(10) \} = 
\{ $multians1->ans_rule(10) \} + 
\{ $multians1->ans_rule(10) \} + 
\{ $multians1->ans_rule(10) \}. 

$ITEMSEP
$ITEM
Find the value of each term in the equation
\( y_2^{\,\prime} = $row2 \) in terms of the variable \( t \).
(Enter the terms in the order given.)
$BR
$BR
\{ $multians2->ans_rule(10) \} = 
\{ $multians2->ans_rule(10) \} + 
\{ $multians2->ans_rule(10) \} + 
\{ $multians2->ans_rule(10) \}. 

$ITEMSEP
$ITEM
Find the value of each term in the equation
\( y_3^{\,\prime} = $row3 \) in terms of the variable \( t \).
(Enter the terms in the order given.)
$BR
$BR
\{ $multians3->ans_rule(10) \} = 
\{ $multians3->ans_rule(10) \} + 
\{ $multians3->ans_rule(10) \} + 
\{ $multians3->ans_rule(10) \}. 


\{ EndList("OL") \}

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );
ANS( $multians3->cmp() );


COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
