## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Repeated eigenvalues)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-RR-LI)
## Level(3)
## Static(1)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.3')
## Problem1('21')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real",
);
Context()->reduction->set('(-x)-y'=>0);
Context()->reduction->set('(-x)+y'=>0);
Context()->variables->set(t=>{limits=>[1,2]});

$tinit = random(2,5,1);

@P = (
[Formula("2 t^(-2)"),Formula("1 - 2 t^(-1) + 2 t^(-2)")],
[Formula("-2 t^(-2)"),Formula("2 t^(-1) - 2 t^(-2)")]
);

@y1 = ( 
  Formula("t^2 - 2 t"), 
  Formula("2 t") 
); 
@y2 = ( 
  Formula("t - 1"), 
  Formula("1") 
); 


$multians1 = MultiAnswer($y1[0], $y2[0], $y1[1], $y2[1])->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
$BITALIC This is the third part of a four-part problem. $EITALIC
$BR
$BR
If the given solutions
$PAR
\[
\boldsymbol{\vec{y}_1}(t) =
\left\lbrack
\begin{array}{c}
$y1[0] \\
$y1[1]
\end{array}
\right\rbrack,
\ \ \
\boldsymbol{\vec{y}_2}(t) =
\left\lbrack
\begin{array}{c}
$y2[0] \\
$y2[1]
\end{array}
\right\rbrack.
\]
$BR
form a fundamental set (i.e., linearly independent set) of solutions for the system
$PAR
\[
\boldsymbol{\vec{y}^{\,\prime}}
=
\left\lbrack
\begin{array}{cc}
$P[0][0] & $P[0][1] \\
$P[1][0] & $P[1][1]
\end{array}
\right\rbrack
\boldsymbol{\vec{y}},
\]
$BR
state the general solution of the linear homogeneous
system, and if they do not, enter ${BITALIC}NONE${EITALIC}
in all of the answer blanks.  Express the general 
solution as the product 
\( \boldsymbol{\vec{y}}(t) = \Psi(t) \boldsymbol{\vec{c}} \),
where \( \Psi(t) \) is a square matrix whose columns
are the solutions forming the fundamental set and 
\( \boldsymbol{\vec{c}} \) is a column vector of arbitrary 
constants.
$PAR
$BCENTER
\{ 
mbox(
  display_matrix([["\( y_1(t) \)"],["\( y_2(t) \)"]], align=>'cc'),
  " =",
  display_matrix([
  [$multians1->ans_rule(20),$multians1->ans_rule(20)], 
  [$multians1->ans_rule(20),$multians1->ans_rule(20)]], align=>'cc'),
  display_matrix([["\( c_1 \)"],["\( c_2 \)"]], align=>'cc')
);
\}
$ECENTER

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

COMMENT("Uses parserMultiAnswer.pl to format the matrix nicely.");

ENDDOCUMENT();
