## DESCRIPTION
## Linear systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Verification of solutions)
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-VS-find_parameters)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('3.3')
## Problem1('2')
## KEYWORDS('differential equations','linear systems')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "unionLists.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->add(
t=>"Real",
a=>"Real",b=>"Real",k=>"Real",
A=>"Real",B=>"Real"
);

$answer1 = Compute("a k e^(k t)");
$answer2 = Compute("b k e^(k t)");

$i = random(0,2,1);

if ($i == 0) {

  @matrix = (1,3,3,1);
  @evect1 = (1,1);
  $eval1  = 4;
  @evect2 = (1,-1);
  $eval2  = -2;

} elsif ($i == 1) {

  @matrix = (7,2,-4,1);
  @evect1 = (1,-1);
  $eval1  = 5;
  @evect2 = (1,-2);
  $eval2  = 3;

} else {

  @matrix = (5,3,3,5);
  @evect1 = (1,-1);
  $eval1  = 2;
  @evect2 = (1,1);
  $eval2  = 8;
}

$x1a = Compute("$evect1[0] A e^($eval1 t) + $evect2[0] B e^($eval2 t)");
$x2a = Compute("$evect1[1] A e^($eval1 t) + $evect2[1] B e^($eval2 t)");

$x1b = Compute("$evect1[0] B e^($eval1 t) + $evect2[0] A e^($eval2 t)");
$x2b = Compute("$evect1[1] B e^($eval1 t) + $evect2[1] A e^($eval2 t)");

$multians1 = MultiAnswer($answer1, $answer2)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td></tr><tr><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack \begin{array}{r} %s \\ %s \end{array} \right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians2 = MultiAnswer($x1a, $x2a)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td></tr><tr><td> %s </td></tr></table>",
  tex_format =>  "\begin{array}{r} x_1(t) = %s \\ x_2(t) = %s \end{array}",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);



######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

  

$ITEM Assuming \( a, b \) and \( k \) are constants, 
calculate the following derivative.
\{ 
mbox([
  "\( \displaystyle \frac{d}{dt} \left( \left\lbrack \begin{array}{c} a \\ b \end{array} \right\rbrack e^{k t} \right) = \)",
  display_matrix([[$multians1->ans_rule(10)], [$multians1->ans_rule(10)]], align=>'c')], 
  valign=>'middle'
);
\}

$ITEMSEP
$ITEM Find a value of \( k \) so that 
\( \displaystyle 
\left\lbrack 
\begin{array}{r} $evect1[0] \\ $evect1[1] \end{array} 
\right\rbrack 
e^{k t}
\) 
is a solution to
\( \displaystyle
\boldsymbol{\vec{x}^{\,\prime}}
= 
\left\lbrack
\begin{array}{rr}
$matrix[0] & $matrix[1] \\
$matrix[2] & $matrix[3]
\end{array}
\right\rbrack 
\boldsymbol{\vec{x}}.
\)
$BR
\( k = \) 
\{ ans_rule(10) \}

$ITEMSEP
$ITEM Find a value of \( k \) so that 
\( \displaystyle 
\left\lbrack 
\begin{array}{r} $evect2[0] \\ $evect2[1] \end{array} 
\right\rbrack 
e^{k t}
\) 
is a solution to
\( \displaystyle
\boldsymbol{\vec{x}^{\,\prime}}
= 
\left\lbrack
\begin{array}{rr}
$matrix[0] & $matrix[1] \\
$matrix[2] & $matrix[3]
\end{array}
\right\rbrack 
\boldsymbol{\vec{x}}.
\)
$BR
\( k = \) 
\{ ans_rule(10) \}

$ITEMSEP
$ITEM Write down the general solution in 
the form \( x_1(t) = ? \) and \( x_2(t) = ? \), 
i.e., write down a formula for each 
component of the solution.  Use \( A \)
and \( B \) to denote arbitrary constants.
$BR
\( x_1(t) = \) 
\{ $multians2->ans_rule(40) \}
$BR
\( x_2(t) = \) 
\{ $multians2->ans_rule(40) \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

#ANS( $answer1->cmp() );
#ANS( $answer2->cmp() );

ANS( Compute("$eval1")->cmp() );
ANS( Compute("$eval2")->cmp() );

ANS( $multians2->cmp() );

;

ENDDOCUMENT();
