## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-MN-matrix_derivatives)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.1')
## Problem1('22')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");

$b11 = non_zero_random(-5,5,1);
$b12 = non_zero_random(-5,5,1);
$b21 = non_zero_random(-5,5,1);
$b22 = non_zero_random(-5,5,1);

$d11 = non_zero_random(-5,5,1);
$d12 = non_zero_random(-5,5,1);
$d21 = non_zero_random(-5,5,1);
$d22 = non_zero_random(-5,5,1);

$c11 = $b11 - 1/2 - $d11;
$c12 = $b12 - 1/6 - $d12;
$c21 = $b21 - $d21;
$c22 = $b22 - $d22;

$A11 = Formula("t^2 / 2 + $c11 t + $d11")->reduce();
$A12 = Formula("t^3 / 6 + $c12 t + $d12")->reduce();
$A21 = Formula("$c21 t + $d21")->reduce();
$A22 = Formula("$c22 t + $d22")->reduce();

$multians = MultiAnswer($A11, $A12, $A21, $A22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose
$PAR
\[
A''(t) =
\left\lbrack
\begin{array}{rr}
1 & t \\
0 & 0
\end{array}
\right\rbrack,
\ \ \
A(0) = 
\left\lbrack
\begin{array}{rr}
$d11 & $d12 \\
$d21 & $d22
\end{array}
\right\rbrack,
\ \ \
A(1) = 
\left\lbrack
\begin{array}{rr}
$b11 & $b12 \\
$b21 & $b22
\end{array}
\right\rbrack.
\]
$BR

Determine the general form of \( A(t) \) by constructing
antiderivatives as needed and imposing any given
constraints.

$PAR
$BCENTER
\{ 
mbox(
  "\( A(t) = \)",
  display_matrix([
  [$multians->ans_rule(20),$multians->ans_rule(20)], 
  [$multians->ans_rule(20),$multians->ans_rule(20)]
  ], align=>'cc')
);
\}
$ECENTER

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
