































































## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.1')
## Problem1('8')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "unionLists.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

install_weighted_grader();

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real",k=>"Real");
Context()->variables->set(t=>{limits=>[0.1,1]});

$a = random(2,6,1);
do { $b = random(2,6,1); } until ($b != $a);

$a11 = Formula(" $a sin(t)");  $a12 = Formula("-$b cos(t)");
$a21 = Formula(" $a sin(t)");  $a22 = Formula(" $b cos(t)");

$answer = Formula("k*pi/2")->with(period=>pi/2);

$a2 = 2 * $a;
$b2 = 2 * $b;

$ainv11 = Formula(" 1/($a2 sin(t))");  $ainv12 = Formula("1/($a2 sin(t))");
$ainv21 = Formula("-1/($b2 cos(t))");  $ainv22 = Formula("1/($b2 cos(t))");


$multians = MultiAnswer($ainv11, $ainv12, $ainv21, $ainv22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\\\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);



######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Let
$PAR 
\[
A(t) = 
\left\lbrack
\begin{array}{rr}
$a11 & $a12 \\
$a21 & $a22
\end{array}
\right\rbrack.
\]
$BR

\{ BeginList("OL",type=>'a') \}
$ITEM
Find the values of \( t \) such that \( A(t) \) is 
${BBOLD}not${EBOLD} invertible.  You may use \( k \) 
to denote any possible integer in your answer (e.g.,
if the answer is all integer multiples of \( 5 \), 
you would enter \( 5k \), where \( k \) is any integer).    
$BR
$BR
\( A(t) \) is ${BBOLD}not${EBOLD} invertible when 
\( t = \) \{ ans_rule(20) \}, 
where \( k \) is any integer.

$ITEMSEP
$ITEM
Find a formula for \( A^{-1}(t) \) for the values of 
\( t \) for which \( A(t) \) is invertible.
$BCENTER
\{ 
mbox(
  "\( A^{-1}(t) = \)",
  display_matrix([
  [$multians->ans_rule(15),$multians->ans_rule(15)], 
  [$multians->ans_rule(15),$multians->ans_rule(15)]
  ], align=>'cc')
);
\}
$ECENTER

\{ EndList("OL") \}

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $answer->cmp(), 20 );

WEIGHTED_ANS( $multians->cmp(), 80 );

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
