## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-MN-matrix_derivatives)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.1')
## Problem1('13')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "unionLists.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "contextInequalities.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

install_weighted_grader();

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[0.1,1]});

$a = random(1,9,1);
do { $b = random(2,6,1); } until ($b != $a);
do { $c = random(2,6,1); } until ($c != $b && $c != $a);

$a11 = Formula("$a");            $a12 = Formula("ln|t|");
$a21 = Formula("sqrt($b - t)");  $a22 = Formula("e^($c t)");

$ap11 = Formula("0");  $ap12 = Formula("1/t");
$ap21 = $a21->D('t');  $ap22 = $a22->D('t');

$app11 = $ap11->D('t');  $app12 = $ap12->D('t');
$app21 = $ap21->D('t');  $app22 = $ap22->D('t');


$multians1 = MultiAnswer($ap11, $ap12, $ap21, $ap22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\\\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);


$multians2 = MultiAnswer($app11, $app12, $app21, $app22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\\\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);



######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Let
$PAR 
\[
A(t) = 
\left\lbrack
\begin{array}{cc}
$a11 & $a12 \\
$a21 & $a22
\end{array}
\right\rbrack.
\]
$BR

\{ BeginList("OL",type=>'a') \}

$ITEM
Find \( A'(t) \).
$BCENTER
\{ 
mbox(
  "\( A'(t) = \)",
  display_matrix([
  [$multians1->ans_rule(25),$multians1->ans_rule(25)], 
  [$multians1->ans_rule(25),$multians1->ans_rule(25)]
  ], align=>'cc')
);
\}
$ECENTER

$ITEMSEP
$ITEM
Find \( A''(t) \).
$BCENTER
\{ 
mbox(
  "\( A''(t) = \)",
  display_matrix([
  [$multians2->ans_rule(25),$multians2->ans_rule(25)], 
  [$multians2->ans_rule(25),$multians2->ans_rule(25)]
  ], align=>'cc')
);
\}
$ECENTER
END_TEXT
Context()->normalStrings;

Context("Inequalities")->variables->are(t=>"Real");

@interval = ();

$interval[0] = Compute("t < 0 or 0 < t <= $b");
$interval[1] = Compute("t < 0 or 0 < t < $b");
$interval[2] = Compute("t < 0 or 0 < t < $b");

Context()->texStrings;
BEGIN_TEXT

$ITEMSEP
$ITEM
\( A(t) \) is defined for all \( t \) in the interval(s) 
\{ ans_rule(30) \}
\{ AnswerFormatHelp("intervals") \}
$BR
$BR
\( A'(t) \) is defined for all \( t \) in the interval(s) 
\{ ans_rule(30) \}
\{ AnswerFormatHelp("intervals") \}
$BR
$BR
\( A''(t) \) is defined for all \( t \) in the interval(s) 
\{ ans_rule(30) \}
\{ AnswerFormatHelp("intervals") \}

\{ EndList("OL") \}

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 35 );
WEIGHTED_ANS( $multians2->cmp(), 35 );

foreach my $i (0..2) {
  WEIGHTED_ANS( $interval[$i]->cmp(), 10 );
}

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
