# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 03/25/10


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Double integrals over general regions)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('16.2')
## Section2('16.2')
## Section3('16.2')
## Problem1('33')
## Problem2('42')
## Problem3('56')
## KEYWORDS('iterated integral', 'double integral', 'definite integrals', 'functions', 'multivariable')

## Textbook tags
## HHChapter1('Integrating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( y=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,5,1);
$m = list_random(2,3);

$a1 = Formula( 0 );    $b1 = Formula( $a );
$c1 = Formula( "x" );  $d1 = Formula( "$m*x" );
$dx1 = Compute( "y" ); $dy1 = Compute( "x" );

$a2 = Formula( 0 );    $b2 = Formula( $a );
$c2 = Formula( "y/$m" );  $d2 = Formula( "y" );
$a3 = Formula( $a );   $b3 = Formula( $m*$a );
$c3 = Formula( "y/$m" );  $d3 = Formula( "$a" );

$dx2 = Compute( "x" ); $dy2 = Compute( "y" );

$mp = MultiAnswer( $dx2, $dy2, $dx2, $dy2, $a2, $b2, $c2, $d2,
		   $a3, $b3, $c3, $d3 )->with(
    singleAnswer=>0,
    checker=>sub {
	my ( $c, $s, $ansHash ) = @_;
	@res = ( 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0 );
	my $done1 = 0;
	my $done2 = 0;

	## check the two integrals separately, and require that the
	##    order of the differentials be correct before giving
	##    credit for the limits.
	if ( $c->[0] == $s->[0] && $c->[1] == $s->[1] ) {
	    ## then the dy dx matches correctly for the first integral
	    $res[0] = 1;
	    $res[1] = 1;
	    if ( $c->[4] == $s->[4] && $c->[5] == $s->[5] ) {
		## then outer limits are correct
		$res[4] = 1;
		$res[5] = 1;
		$done1 = 1;
		$res[6] = ( $c->[6] == $s->[6] ) ? 1 : 0;
		$res[7] = ( $c->[7] == $s->[7] ) ? 1 : 0;
	    } elsif ( $c->[8] == $s->[4] && $c->[9] == $s->[5] ) {
		$res[4] = 1;
		$res[5] = 1;
		$done2 = 1;
		$res[6] = ( $c->[10] == $s->[6] ) ? 1 : 0;
		$res[7] = ( $c->[11] == $s->[7] ) ? 1 : 0;
	    }
	}
	## the second integral
	if ( $c->[2] == $s->[2] && $c->[3] == $s->[3] ) {
	    ## correct dy dx
	    $res[2] = 1;
	    $res[3] = 1;
	    if ( $done1 ) {
		$res[8] = ( $c->[8] == $s->[8] ) ? 1 : 0;
		$res[9] = ( $c->[9] == $s->[9] ) ? 1 : 0;
		$res[10] = ( $c->[10] == $s->[10] ) ? 1 : 0;
		$res[11] = ( $c->[11] == $s->[11] ) ? 1 : 0;
	    } elsif ( $done2 ) {
		$res[8] = ( $c->[4] == $s->[8] ) ? 1 : 0;
		$res[9] = ( $c->[5] == $s->[9] ) ? 1 : 0;
		$res[10] = ( $c->[6] == $s->[10] ) ? 1 : 0;
		$res[11] = ( $c->[7] == $s->[11] ) ? 1 : 0;
	    } else {
		$res[8] = ( $c->[4] == $s->[8] || $c->[8] == $s->[8] ) ? 1 : 0;
		$res[9] = ( $c->[5] == $s->[9] || $c->[9] == $s->[9] ) ? 1 : 0;
		$res[10] = ( $c->[6] == $s->[10] || $c->[10] == $s->[10] ) ?
		    1 : 0;
		$res[11] = ( $c->[7] == $s->[11] || $c->[11] == $s->[11] ) ?
		    1 : 0;
	    }
	}
	return [ @res ];
    } );

$intVal = Compute( "(($a*$a - 1)*e^($a*$a)/2 + (1/2))*($m - 1)" )->with(
    tolType=>'absolute', tolerance=>.0001 );


Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Let \( f(x,y)=x^2e^{x^2} \) and let \( R \) be the triangle bounded by
the lines \( x=$a \), \( x=y/$m \), and \( y=x \) in the \( xy \)-plane.
$PAR
${BBOLD}(a)$EBOLD
Express \( \int_R{f}\,dA \) as a double integral in two different ways
by filling in the values for the integrals below.
(For one of these it will be necessary to write the double integral as
a sum of two integrals, as indicated; for the other, it can be written
as a single integral.)
$PAR
\( \int_R{f}\,dA = \int_a^b\int_c^d f(x,y)\, d\)\{ans_rule(3)\}
\(d\)\{ans_rule(3)\}
$BR
where \(a = \) \{ans_rule(15)\},
\(b = \) \{ans_rule(15)\}, \(c = \) \{ans_rule(15)\}, and
\(d = \) \{ans_rule(15)\}.
$PAR
And
\( \int_R{f}\,dA = \int_a^b\int_c^d f(x,y)\, d\)\{$mp->ans_rule(3)\}
\(d\)\{$mp->ans_rule(3)\} 
\( + \int_m^n\int_p^q f(x,y)\, d\)\{$mp->ans_rule(3)\}
\(d\)\{$mp->ans_rule(3)\}
$BR
where \( a = \) \{ $mp->ans_rule(15) \}, \( b = \) \{ $mp->ans_rule(15) \},
\( c = \) \{ $mp->ans_rule(15) \}, \( d = \) \{ $mp->ans_rule(15) \},
\( m = \) \{ $mp->ans_rule(15) \}, \( n = \) \{ $mp->ans_rule(15) \},
\( p = \) \{ $mp->ans_rule(15) \}, and \( q = \) \{ $mp->ans_rule(15) \}.

$PAR
${BBOLD}(b)$EBOLD
Evaluate one of your integrals to find the value of \( \int_R{f}\,dA \).
$BR
\( \int_R{f}\,dA = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($dx1->cmp() ); ANS($dy1->cmp() );
ANS($a1->cmp() );  ANS($b1->cmp() );  ANS($c1->cmp() );  ANS($d1->cmp() );
ANS($mp->cmp() );
ANS($intVal->cmp() );

$gr = init_graph(-1,-1,($a+1),($m*$a+1), axes=>[0,0], size=>[250,250],
		 ticks=>[($a+2),($m*$a+2)]);
$gr->moveTo(0,0);
$gr->lineTo($a,$a,'blue',2);
$gr->lineTo($a,$m*$a,'blue',2);
$gr->lineTo(0,0,'blue',2);
$gr->new_color( 'ltblue', 214, 230, 244 );
$gr->fillRegion( [$a-0.1, $a, 'ltblue'] );

$asq = $a*$a;
$ma = $m*$a;
$mm1 = ( $m == 2 ) ? '' : 2;
$mm1o2 = ( $m == 2 ) ? "\frac{1}{2}" : '';
($c1n, $c1d) = reduce( ($m-1)*$asq, 2 );
$c1 = ( $c1d == 1 ) ? $c1n : "\frac{$c1n}{$c1d}";
($c2n, $c2d) = reduce( ($asq-1)*($m-1), 2 );
$c2 = ( $c2d == 1 ) ? $c2n : "\frac{$c2n}{$c2d}";
$mm1o2f = ( $m == 2 ) ? "\frac12" : 1;
$val = "$c2 e^{$asq} + $mm1o2f";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The line \(x=y/$m\) is the line \(y=$m x\), and \(y=x\) and \(y=$m x\)
intersect at \(x=0\).  Thus, \(R\) is the shaded region in the figure
below.
$PAR
$BCENTER
\{ image( insertGraph( $gr ), tex_size=>200, height=>250, width=>250,
          extra_html_tags=>'alt="Graph of the triangular region with ' .
	  'vertices (0,0), (' . $a . ',' . $a . ') and (' . $a . ',' .
	  $ma . ')."' ) \}
$ECENTER
$PAR
One expression for the integral is
\[
  \int_R f\,dA =\int_0^{$a}\int_x^{$m x} x^2e^{x^2}\,dy\,dx.
\]
Another expression is obtained
by reversing the order of integration.  When we do this, it is
necessary to split \(R\) into two regions on a line parallel to the
\(x\)-axis along the point of intersection
of \(y=x\) and \(x=$a\); this line is \(y=$a\).  Then we obtain
\[
\int_Rf\,dA =
\int_0^{$a}\int_{y/$m}^{y} x^2e^{x^2}\,dx\,dy+
      \int_{$a}^{$ma} \int_{y/2}^{$a} x^2e^{x^2}\,dx\,dy.
\]

$PAR
${BBOLD}(b)$EBOLD
We evaluate the first integral.  Integrating with respect to \(y\) first:
\[
  \int_0^{$a}\int_x^{$m x} x^2 e^{x^2}\,dy\,dx =
    \int_0^{$a} (x^2e^{x^2}y)\bigg|_x^{$m x}\,dx =
    \int_0^{$a} $mm1 x^3 e^{x^2}\,dx.
\]
To do the second integral, we use integration by parts with
\(u=x^2\), \(v'=xe^{x^2}\).  Then \(u'=2x\) and
\(v=\frac{1}{2}e^{x^2}\), so
\[
\int_0^{$a} $mm1 x^3e^{x^2}\,dx = $mm1o2 x^2e^{x^2}\bigg|_0^{$a}
    - \int_0^{$a} $mm1 x e^{x^2}\,dx.
\]
The second integral is then easily obtained, giving
\[
\int_0^{$a} $mm1 x^3e^{x^2}\,dx = \left($mm1o2 x^2e^{x^2}
   - $mm1o2 e^{x^2}\right)\bigg|_0^{$a} =
   $c1 e^{$asq} - $mm1o2 e^{$asq} + $mm1o2 = $val.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
