## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Complex eigenvalues and eigenvectors)
## Date(November 2018)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

# Randomization taken from:
# Library/TCNJ/TCNJ_CharacteristicPolynomial/problem6.pg

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGcourse.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Complex-Matrix");

$A11 = random(-1,1,2);
$A12 = non_zero_random(-3,3,1);
$mult = non_zero_random(-2,2,1);
$A21 = $mult * $a11;
$b1 = random(-1,1,2);
$A22 = $mult * $a12 + $b1;

$A = Matrix([[$A11, $A12],[$A21, $A22]]);

$B = $A->inverse;

$real = non_zero_random(-8,8,1);
do { $imag = random(1,8,1); } until ($imag != $real);

$E = Matrix([
[$real, -1*$imag],
[$imag, $real]
]);

$M = $A * $E *$B;

$e1 = $real + $imag * i;
$e2 = $real - $imag * i;

$v1 = Matrix([[$A11 - $A12 * i],[$A21 - $A22 * i]]);
$v2 = Matrix([[$A11 + $A12 * i],[$A21 + $A22 * i]]);

Context()->texStrings;
BEGIN_TEXT
Find the eigenvalues and eigenvectors for \( A = $M \).
$BR
$BR
The eigenvalue \( a + b i = \) \{ $e1->ans_rule(5) \} has an eigenvector \{ $v1->ans_array(5) \}.
$BR
$BR
The eigenvalue \( a - b i = \) \{ $e2->ans_rule(5) \} has an eigenvector \{ $v2->ans_array(5) \}.
END_TEXT
Context()->normalStrings;

ANS( $e1->cmp );
ANS( $v1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0]]);
return $M * $student == $e1 * $student;
}) );

ANS( $e2->cmp );
ANS( $v2->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0]]);
return $M * $student == $e2 * $student;
}) );

ENDDOCUMENT();
