# DESCRIPTION
# Solve an equation for a certain variable in terms of other variables
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Equations and Inequalities')
## DBsection('Other Equations')
## KEYWORDS()
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","z","p","q","k","u","v","h","s");
@caps = ("C","N","M","R","L","J","F");
$c1 = random(0,scalar(@caps)-1,1);
$c2 = random(0,scalar(@vars)-1,1);
do {$c3 = random(0,scalar(@vars)-1,1)} until ($c3 != $c2);
do {$c4 = random(0,scalar(@vars)-1,1)} until ($c4 != $c2 && $c4 != $c3);
$var1 = $caps[$c1];
$var2 = $vars[$c2];
$var3 = $vars[$c3];
$var4 = $vars[$c4];
Context()->variables->are(lc($var1)=>"Real", lc($var2)=>"Real", lc($var3)=>"Real", lc($var4)=>"Real", uc($var1)=>"Real", uc($var2)=>"Real", uc($var3)=>"Real", uc($var4)=>"Real", );


for ($i = 0; $i < random(10,50); $i++) {
  $case = random(0,1,1);
}

# $a = random(2,15)*random(-1,1,2);
$a = 1;

if ($case=0) {
   $b = 1; #random(2,6);
   $c = random(2,15);
   do {$d = random(2,15)*random(-1,1,2)} until (gcd($c,$d) == 1);

   $dispeqn = "\displaystyle $var2 = $var3 \left( $c $var4 + $d $var1\right)";
}
else {
   $b = random(2,6);
   $c = 1; #random(2,15);
   do {$d = random(2,15)*random(-1,1,2)} until (gcd($c,$d) == 1);

   $dispeqn = "\displaystyle $var2 = \frac{$var3}{$b} \left( $var4 + $d $var1\right)";
}


$var = ($case == 0) ? $var4 : $var1;
$e1 = $b;
$num = ($case == 0) ? Formula("$e1*$var2 - $d $var1 $var3")->reduce : Formula("$e1*$var2 - $c $var4 $var3")->reduce;
$den = Formula("$c $var3")->reduce;
Context()->texStrings;
$dispans = "\displaystyle $var = \frac{$num}{$den}";

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
 #         $self->setMessage(1,"Find a common denominator first");
  #        $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $var = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $var = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT

Solve the equation for \($var\):
$PAR
\[ 
$dispeqn
\]
$PAR
$displayfrac
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp(studentsMustReduceFractions=>1));


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \($dispans\)
  $PAR
  END_TEXT
}
else  {
    BEGIN_TEXT
    $PAR
    $BBOLD Note: $EBOLD
    Write your answer as a single fraction, and type the numerator and denominator separately in the answer blanks provided.
    $PAR
    END_TEXT
} 
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
