## DESCRIPTION
## Systems of ODEs: Nonhomogeneous Systems
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Nonhomogeneous systems)
## Date(24/11/2015)
## Institution(METU-NCC)
## Author(Benjamin Walter)
## Level(4)
## Static(0)
## MO(1)


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "answerHints.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

#####################################
#  Setup

Context("Numeric");
Context()->variables->are(
   t=>"Real",
);
Context()->flags->set(
   zeroLevel=>0.001,
   limits=>[0,2],
   zeroLevelTol=>0.001
);



# 1) choose matrix of form
#    [ a  -b ]
#    [ c  -a ]   (eigenvalues = 0,0)
#
#  where a^2 - bc = 0  possibilities are in @CNum and @OffD
#
# 2) phase shift to get
#    [ p+a  -b ]
#    [  c  p-a ]
#
#  eigenvalues p, p
#  eigenvector [ b a ] generalized [ 0 -1 ]

my $a, $b, $c, $p, $i, $k, $sign1, $sign2, $l, $v1, $v2, $w1, $w2, $e;

my @CNum =  (1,2,2,3,3,4,4,5,6,6,0);              # a
my @OffD = ([1,2,1,3,1,4,2,5,6,4,random(1,6,1)],  # b
            [1,2,4,3,9,4,8,5,6,9,0]);             # c
$i = random(0,10,1);
$k = 0;
$sign1 = non_zero_random(-1,1,1);
$sign2 = non_zero_random(-1,1,1);

$a = $sign1*$CNum[$i];

$b = $sign2*$OffD[$k][$i];
$c = $sign2*$OffD[($k+1)%2][$i];

#
# choose phase so that matrix entries stay <= 10
#   --> Keep eigenvalues between -4 and 4
#
$p = random( max(-10+abs($a),-4),
             min(10-abs($a),  4),   1);

$A11 = $p+$a;  $A12 =  -$b;
$A21 = $c;     $A22 =  $p-$a;

$v1 = $b;    $w1 = 0;
$v2 = $a;    $w2 = -1;

$l =  $p;

#############################
# setup forcing functions

$e = Formula("e^($l t)")->reduce;

$a = non_zero_random(-4,4,1);
do { $b = non_zero_random(-4,4,1); }
 until (abs($a) != abs($b));

$g1 = Formula("$a$e")->reduce;
$g2 = Formula("$b$e")->reduce;

#############################
# answer evaluators

my $P11, $P12, $P21, $P22;

$P11=Formula("$v1$e")->reduce; $P12=Formula("($w1+$v1*t)$e")->reduce;
$P21=Formula("$v2$e")->reduce; $P22=Formula("($w2+$v2*t)$e")->reduce;

# A -- Fundamental Matrix, Psi

my $Spsi11, $Spsi12, $Spsi21, $Spsi22, $sDet;

$Spsi11 = Formula(""); $Spsi12 = Formula(""); # Keep track of inverse of
$Spsi21 = Formula(""); $Spsi22 = Formula(""); # student's version of Psi
$sDet = Formula("");

$multiansA = MultiAnswer($P11, $P12, $P21, $P22)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0'><tr>
    <td> \( \boldsymbol{\Psi} = \biggl\lbrack \)</td>
    <td><table border='0' cellspacing='10'>
           <tr><td> %s $SPACE </td><td> $SPACE %s</td></tr>
           <tr><td> %s $SPACE </td><td> $SPACE %s</td></tr></table></td>
    <td> \(\biggr\rbrack\)</td>
   </tr></table>",
  tex_format => "\boldsymbol{\Psi} = 
                 \left\lbrack\begin{array}{rr}
                      %s & %s \\ %s & %s
                 \end{array}\right\rbrack",
  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = (0,0);
      my $msg = "";
      if ($c[0]->typeMatch($s[0]) && $c[1]->typeMatch($s[1]) &&
          $c[2]->typeMatch($s[2]) && $c[3]->typeMatch($s[3])) {
        if ((($s[0] != 0) || ($s[2] != 0)) &&
            (Formula("$A11*$s[0] + $A12*$s[2]") == Formula("$s[0]")->D('t')) &&
            (Formula("$A21*$s[0] + $A22*$s[2]") == Formula("$s[2]")->D('t')) ) {
          $score[0] = 1;
        } else {
          $msg = "First fundamental solution is incorrect!$BR";
        } if ((($s[1] != 0) || ($s[3] != 0)) &&
            (Formula("$A11*$s[1] + $A12*$s[3]") == Formula("$s[1]")->D('t')) &&
            (Formula("$A21*$s[1] + $A22*$s[3]") == Formula("$s[3]")->D('t')) &&
            ($s[1]*$s[2] != $s[0]*$s[3]) ) {
          $score[1] = 1;
        } else {
          $msg = $msg . "Second fundamental solution is incorrect!$BR";
        }
        $sDet = Formula("$s[0]*$s[3]-$s[1]*$s[2]");
        if ($sDet == Formula("0")) {
          $msg = $msg . "Fundamental solutions not independent!";
        } else {
          $Spsi11 = Formula("$s[3]")->reduce; $Spsi12 = Formula("-$s[1]")->reduce;
          $Spsi21 = Formula("-$s[2]")->reduce;$Spsi22 = Formula("$s[0]")->reduce;
          if ($score[0]*$score[1] == 1) { return 1; }
          else { Value::Error($msg); }
        }
      }
      if (length($msg) > 1) { Value::Error($msg); }
      return 0;
  }
);

# B -- Phi inverse

my $Pdet, $Pi11, $Pi12, $Pi21, $Pi22;

$Pdet=Formula("($v1*$w2 - $w1*$v2)$e")->reduce;

$Pi11=Formula("($w2+$v2*t)/$Pdet")->reduce; $Pi12=Formula("-($w1+$v1*t)/$Pdet")->reduce;
$Pi21=Formula("-$v2/$Pdet")->reduce;        $Pi22=Formula("$v1/$Pdet")->reduce;

$multiansB = MultiAnswer($Pi11, $Pi12, $Pi21, $Pi22)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "<table border='0'><tr>
    <td> \( \boldsymbol{\Psi}^{\mathrm{-}1} \ = \ \biggl\lbrack \)</td>
    <td><table border='0' cellspacing='10'>
           <tr><td> %s $SPACE </td><td> $SPACE %s</td></tr>
           <tr><td> %s $SPACE </td><td> $SPACE %s</td></tr></table></td>
    <td> \(\biggr\rbrack\)</td>
   </tr></table>",
  tex_format => "\boldsymbol{\Psi}^{\mathrm{-}1} = \left\lbrack\begin{array}{rr}
                      %s & %s \\ %s & %s
                 \end{array}\right\rbrack",
  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      if ($c[0]->typeMatch($s[0]) && $c[1]->typeMatch($s[1]) &&
          $c[2]->typeMatch($s[2]) && $c[3]->typeMatch($s[3])) {
        if ((Formula("$s[0]*$sDet") == Formula("$Spsi11")) &&
            (Formula("$s[1]*$sDet") == Formula("$Spsi12")) &&
            (Formula("$s[2]*$sDet") == Formula("$Spsi21")) &&
            (Formula("$s[3]*$sDet") == Formula("$Spsi22"))) {
          return 1;
        }
      }
      return 0;
  }
);

# C -- integral of Phi inverse g

my $iPig1, $iPig2;

$Pdet = Formula("2($v1*$w2 - $w1*$v2)");

$iPig1 = Formula("(2($a*$w2-$b*$w1)t+($a*$v2-$b*$v1)(t^2))/$Pdet")->reduce;
$iPig2 = Formula("                  2($b*$v1-$a*$v2)t/$Pdet")->reduce;

$multiansC = MultiAnswer($iPig1, $iPig2)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0'><tr>
    <td> \( \biggl\lbrack\)</td>
    <td><table border='0' cellspacing='10'>
           <tr><td> %s </td></tr>
           <tr><td> %s </td></tr></table></td>
    <td> \(\biggr\rbrack\)</td>
   </tr></table>",
  tex_format => "\left\lbrack\begin{array}{r}
                      %s \\ %s 
                 \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      if ($c[0]->typeMatch($s[0]) && $c[1]->typeMatch($s[1]) &&
          ($Spsi11 != Formula("")) && ($Spsi12 != Formula("")) &&
          ($Spsi21 != Formula("")) && ($Spsi22 != Formula("")) ) {
        my @dS = ( Formula("$s[0]")->D('t'), Formula("$s[1]")->D('t'));
        if ((Formula("$Spsi11*$g1 + $Spsi12*$g2") == Formula("$dS[0]*$sDet")) &&
            (Formula("$Spsi21*$g1 + $Spsi22*$g2") == Formula("$dS[1]*$sDet"))){
          return 1;
        }
      }
      return 0;
  }
);

# D -- solution to nonhomogeneous system

my $yP1 = Formula("($v1*$iPig1 + ($w1+$v1*t)*$iPig2)$e")->reduce;
my $yP2 = Formula("($v2*$iPig1 + ($w2+$v2*t)*$iPig2)$e")->reduce;

$multiansD = MultiAnswer($P11, $P21, $P12, $P22, $yP1, $yP2)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0'><tr>
    <td> \( \boldsymbol{\vec{y}} = 
         \biggl\lbrack \)</td>
    <td><table border='0' cellspacing='10'>
           <tr><td> %s </td></tr>
           <tr><td> %s </td></tr></table></td>
    <td> \(\biggr\rbrack c_1 + \biggl\lbrack\) </td>
    <td><table border='0' cellspacing='10'>
           <tr><td> %s </td></tr>
           <tr><td> %s </td></tr></table></td>
    <td> \(\biggr\rbrack c_2 + \biggl\lbrack\) </td>
    <td><table border='0' cellspacing='10'>
           <tr><td> %s </td></tr>
           <tr><td> %s </td></tr></table></td>
    <td> \(\biggr\rbrack \) </td>
   </tr></table>",
  tex_format => "\boldsymbol{\vec{y}} = 
                \left\lbrack\begin{array}{r} 
                      %s \\ %s
                 \end{array}\right\rbrack 
                 c_1 + 
                 \left\lbrack\begin{array}{r} 
                      %s \\ %s
                 \end{array}\right\rbrack 
                 c_2 + 
                 \left\lbrack\begin{array}{r} 
                      %s \\ %s
                 \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = (0,0);
      if ($c[0]->typeMatch($s[0]) && $c[1]->typeMatch($s[1]) &&
          $c[2]->typeMatch($s[2]) && $c[3]->typeMatch($s[3])) {
        if ((($s[0] != 0) || ($s[1] != 0)) &&
            (Formula("$A11*$s[0] + $A12*$s[1]") == Formula("$s[0]")->D('t')) &&
            (Formula("$A21*$s[0] + $A22*$s[1]") == Formula("$s[1]")->D('t')) ) {
          $score[0] = 1;
        } else {
          Value::Error("Homogeneous solution is incorrect!");
        } if ((($s[2] != 0) || ($s[3] != 0)) &&
            (Formula("$A11*$s[2] + $A12*$s[3]") == Formula("$s[2]")->D('t')) &&
            (Formula("$A21*$s[2] + $A22*$s[3]") == Formula("$s[3]")->D('t')) &&
            ($s[0]*$s[3] != $s[1]*$s[2]) ) {
          $score[1] = 1;
        } else {
          Value::Error("Homogeneous solution is incorrect!");
        }
      }
      if ($c[4]->typeMatch($s[4]) && $c[5]->typeMatch($s[5])) {
        if ((Formula("$A11*$s[4] + $A12*$s[5] + $g1") == Formula("$s[4]")->D('t')) &&
            (Formula("$A21*$s[4] + $A22*$s[5] + $g2") == Formula("$s[5]")->D('t'))) {
          return (($score[0] == 1) && ($score[1] == 1)) ? 1 : 0;
        } else {
          Value::Error("Particular solution is incorrect!");
        }
      }
      return 0;
    }
);



#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
In this problem you will solve the ${BBOLD}nonhomogeneous$EBOLD system
\[
\boldsymbol{\vec{y}\,'}\  = \ 
 \left\lbrack
   \begin{array}{rr}
    $A11 & $A12 \\
    $A21 & $A22
   \end{array}
 \right\rbrack \,
\boldsymbol{\vec{y}}
\ + \ 
 \left\lbrack
   \begin{array}{r}
    $g1 \\ $g2
   \end{array}
 \right\rbrack\]

$PAR
$HR
$PAR

$BBOLD A. $EBOLD
Write a fundamental matrix for the associated homogeneous system
$BR
\{ mbox( [ 
'\( \boldsymbol{\Psi}\ =\ \) ',
display_matrix(
   [[$multiansA->ans_rule(10), $multiansA->ans_rule(10)],
    [$multiansA->ans_rule(10), $multiansA->ans_rule(10)]])
] ) \}

$PAR

$BBOLD B. $EBOLD
Compute the inverse
$BR
\{ mbox( [ 
'\( \boldsymbol{\Psi}^{\mathrm{-}1}\ =\ \) ',
display_matrix(
   [[$multiansB->ans_rule(15), $multiansB->ans_rule(15)],
    [$multiansB->ans_rule(15), $multiansB->ans_rule(15)]])
] ) \}


$PAR

$BBOLD C. $EBOLD
Multiply by \(\boldsymbol{\vec{g}}\) and integrate
$BR

\{ mbox( [ 
'\(\displaystyle  \int \boldsymbol{\Psi}^{\mathrm{-}1}\; \boldsymbol{\vec{g}}\; dt\ =\ \) ',
display_matrix(
   [[$multiansC->ans_rule(15), '\(+ c_1\)'],
    [$multiansC->ans_rule(15), '\(+ c_2\)']])
] ) \}

(Do not include \(c_1\) and \(c_2\) in your answers).

$PAR

$BBOLD D. $EBOLD
Give the solution to the system
$BR
\{ mbox( [ 
'\(\displaystyle \boldsymbol{\vec{y}}\ =\ \) ',
display_matrix(
   [[$multiansD->ans_rule(10)],
    [$multiansD->ans_rule(10)]]),
'\( c_1 + \)',
display_matrix(
   [[$multiansD->ans_rule(10)],
    [$multiansD->ans_rule(10)]]),
'\( c_2 \)'
]) \}
$BR
\{ mbox([
'\(\phantom{y = } \quad + \)',
display_matrix(
   [[$multiansD->ans_rule(40)],
    [$multiansD->ans_rule(40)]])
] ) \}

(Do not include \(c_1\) and \(c_2\) in your answers).

END_TEXT

$showHint = 2;

BEGIN_TEXT
$HR
$PAR
If you don't get this in $showHint tries, you can get a hint.
END_TEXT

BEGIN_HINT
$PAR
$SPACE $SPACE \(\bullet\) The matrix 
    \(\left\lbrack
      \begin{array}{rr}
       $A11 & $A12 \\
       $A21 & $A22
      \end{array}
    \right\rbrack\) 
  has repeated eigenvalue \(\lambda = $l, $l\) with
 $BR \(\qquad\) 
   eigenvector 
    \(\boldsymbol{\vec{v}} = \left\lbrack
      \begin{array}{r}
       $v1 \\ $v2
      \end{array}
    \right\rbrack\)
  and
 $BR \(\qquad\)
   \(\boldsymbol{\vec{v}}\)-generalized eigenvector
    \(\boldsymbol{\vec{w}} = \left\lbrack
      \begin{array}{r}
       $w1 \\ $w2
      \end{array}
    \right\rbrack\).
$PAR
$SPACE $SPACE \(\bullet\) The matrix
   \(\displaystyle \boldsymbol{\Psi} =
    \left\lbrack \begin{array}{cc}
         {\scriptsize |}     &   {\scriptsize |}    \\
       \boldsymbol{\vec{v}}  & \boldsymbol{\vec{w}} \\
         {\scriptsize |}     &   {\scriptsize |}
    \end{array} \right\rbrack \,
    \left\lbrack \begin{array}{rr}
        1 & t  \\
        0 & 1
    \end{array}\right\rbrack \,
    e^{\alpha\, t} \)
$BR
 \(\qquad\) has determinant
  \(\displaystyle \mathrm{det} \bigl(\boldsymbol{\Psi}\bigr) =
     \mathrm{det}\left(
     \left\lbrack \begin{array}{cc}
         {\scriptsize |}     &   {\scriptsize |}    \\
       \boldsymbol{\vec{v}}  & \boldsymbol{\vec{w}} \\
         {\scriptsize |}     &   {\scriptsize |}
    \end{array} \right\rbrack\right) \,
   e^{2\alpha\, t} \)
$PAR
$SPACE $SPACE \(\bullet\) This problem has been carefully designed so that
 integration by parts is not needed.
END_HINT

Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multiansA->cmp() );
ANS( $multiansB->cmp() );
ANS( $multiansC->cmp() );
ANS( $multiansD->cmp() );

COMMENT('MathObject version.  Randomized with REPEATED integer eigenvalues.');

ENDDOCUMENT(); 
