## DESCRIPTION
## College Algebra
## ENDDESCRIPTION
 
 
## DBsubject()
## DBchapter()
## DBsection()
## Date(6/21/2017)
## Institution(Community College of Denver)
## Author(Michael Sulock)
## MO(1)
## KEYWORDS('college algebra')
 
 
###########################
#  Initialization
 
DOCUMENT();
 
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"contextLimitedPoint.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
);
 
TEXT(beginproblem());

 
###########################
#  Setup
Context("Numeric");
install_problem_grader(~~&std_problem_grader);
 

$var1=non_zero_random(2,6,1);
$var2=non_zero_random(2,6,1);

# The two sides of the correct answer
$ans1=Formula("log$var1($var2)");
$ans2=Formula("x");
 

# Allows for the correct equation to be given in either order
$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2 ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $stu1 && $f2 == $stu2) ||
           ($f1 == $stu2 && $f2 == $stu1) ) {
          return [1,1];
      } else {
          if ($f1 == $stu1 || $f2 == $stu1) {
              return [1,0];
          } elsif ($f1 == $stu2 || $f2 == $stu2) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
 
 
###########################
#  Main text
 
BEGIN_PGML


Rewrite the following equation in logarithmic form. 

[`x^[$var1]=[$var2] `] 

 
 [________]{$multians} [`=`] [__________]{$multians} [@ AnswerFormatHelp("logarithms") @]*
 

 
END_PGML
 
 
 
############################
#  Solution
 
#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION
 
COMMENT('MathObject version. Uses PGML.');
 
ENDDOCUMENT();