##DESCRIPTION
## Algebra: Analytic Trigonometry
##ENDDESCRIPTION

 
## tcao , PAID on 11-24-2003

## DBsubject(Trigonometry)
## DBchapter(Analytic trigonometry)
## DBsection(Solving trigonometric equations exactly)
## Date(6/3/2002)
## MLT(factorForm)
## Level(3)
## Static(1)
## MO(1)
## TitleText1('Precalculus')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('7.5')
## Problem1('25')
## KEYWORDS('algebra', 'analytic trigonometry', 'trigonometric equation')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
#  "PGchoicemacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

Context("Numeric");
$ans = Real(1/3);

Context()->texStrings;
BEGIN_TEXT
Find all solutions of the equation \( \tan^5 x - 9\tan x =0. \)
$BR
The answer has the form \(x= Ak\pi\) where \(k=0,\pm 1, \pm 2,\pm 3,\ldots \) ranges over the integers and
$BR
the constant \( A= \) \{ans_rule(15)\}.
$BR
END_TEXT
Context()->normalStrings;

ANS($ans->cmp(
  checker=> sub{
    my ($corr, $stud, $self )=@_;
    return  ($stud == $corr or $stud == -$corr  ? 1 : 0 ); 
  }
));

#$ans1= 1/3;
#ANS(num_cmp($ans1));

ENDDOCUMENT();        # This should be the last executable line in the problem.
