# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Functions')
## DBsection('')
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1()
## Problem1()
## Author('Alex Jordan')
## Institution('PCC')



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "bizarroArithmetic.pl",
  "PGML.pl",
  "PGcourse.pl",

);


TEXT(beginproblem());

##############################################


Context("Numeric");
Context()->flags->set(reduceConstants=>0);
Context()->variables->are(t=>'Real');

$a = random(0.1,0.9,0.01);
$k = random(0.1,0.9,0.1);

$b = exp($k);


Context()->operators->set(
'*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
);




##############################################

BEGIN_PGML

Convert this exponential expression to the form [`Q=ab^{t}`]. Write the base [`b`] in exact form, not as a decimal approximation.

    [`Q=[$a]\cdot e^{[$k] t}`]

    [`Q={}`][_____________________]

END_PGML

##############################################

$showPartialCorrectAnswers = 1;

ANS(Formula("$a*(e^$k)^t") -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroMul=>1);
delete $correct->{test_values}, $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroMul=>0);
Value::Error("Your answer is the same as the given expression, but not in the form that was asked for") unless $OK;
return $OK;
})); 

##############################################

BEGIN_PGML_SOLUTION

Recall the exponent rule that has to do with a product in the exponent:

    [`\begin{aligned}
Q &= [$a]\cdot e^{[$k] t}\\
Q &= [$a]\cdot \left(e^{[$k]}\right)^t\\
Q &\approx [$a]\cdot[@exp([$k])@]^t\\
\end{aligned}`]

END_PGML_SOLUTION

ENDDOCUMENT();