## DESCRIPTION
## Equations for lines in the plane
## ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## BookChapter('Linear functions, expressions, and equations')
## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: point-slope form)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(ptslope1)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Algebra: Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('5.4')
## Problem1('14')
## KEYWORDS('linear equations')

##############################################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

##############################################
#  Setup

Context("Numeric");

($x,$y) = (random(2,10,1),random(2,10,1));
$m = random(-1,1,2) * random(3,7,1);

$fac1 = Formula("$y");
$fac2 = Formula("$m * (x - $x)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);



##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Write an equation for the line through the point
\( ($x,$y) \) with slope \( $m \) in point-slope
form \( y = y_0 + m(x-x_0) \).
$BR
$BR
\( y = \) \{ $multians->ans_rule(10) \} + \{ $multians->ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;

ANS($multians->cmp() );


;
ENDDOCUMENT();
