## DESCRIPTION
## Equations for lines in the plane
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: point-slope form)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(ptslope1)
## Level(2)
## MO(1)
## TitleText1('Algebra Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('5.4')
## Problem1('14')
## KEYWORDS('linear equations')

##############################################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "parserAssignment.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


##############################################
#  Setup

Context("Numeric")->variables->add(
'y' => 'Real',
'x0'=> 'Real',
'y0'=> 'Real',
'm' => 'Real',
);
parser::Assignment->Allow;

$eqn = Compute("y = y0 + m (x - x0)");

($x,$y) = (random(2,10,1),random(2,10,1));
$m = random(-1,1,2) * random(3,7,1);

$fac1 = Formula("$y");
$fac2 = Formula("$m * (x - $x)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);



##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
(a) Write an equation for the line through 
the point \( (x_0,y_0) \) with slope \( m \)
in point-slope form.  Enter \( x_0 \) as ${BTT}x0${ETT}
and \( y_0 \) as ${BTT}y0${ETT}.    Use 
\( x \) and \( y \) for the variable names.
Your equation should be of the form \( y = \ldots\). 
$BR
$BR
\{ ans_rule(30) \} 
\{ AnswerFormatHelp("equations") \}
$BR
$BR
(b) Write an equation for the line through the point
\( ($x,$y) \) with slope \( $m \).
$BR
$BR
\( y = \) \{ $multians->ans_rule(10) \} + \{ $multians->ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;

ANS( $eqn->cmp() );
ANS( $multians->cmp() );

ENDDOCUMENT();  
