## DESCRIPTION
## Equations for lines in the plane
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: point-slope form)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(ptslope2)
## Level(2)
## MO(1)
## TitleText1('Algebra Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('5.4')
## Problem1('16')
## KEYWORDS('linear equations')

##############################################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "parserAssignment.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

##############################################
#  Setup

Context("Numeric")->variables->add(
'y' => 'Real',
'x0'=> 'Real',
'y0'=> 'Real',
'm' => 'Real',
);
parser::Assignment->Allow;

($x,$y) = (random(2,10,1),random(2,10,1));
($p,$q) = (random(2,10,1),random(2,10,1));

while ($p==$x || $q==$y) {
  ($p,$q) = (random(2,10,1),random(2,10,1));
}

$dy = $q - $y;
$dx = $p - $x;

$fac1 = Formula("$y");
$fac2 = Formula("($dy / $dx) * (x - $x)");

$fac3 = Formula("$q");
$fac4 = Formula("($dy / $dx) * (x - $p)");


$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($fac1 == $f1stu && $fac2 == $f2stu) ||
           ($fac1 == $f2stu && $fac2 == $f1stu) ||
           ($fac3 == $f1stu && $fac4 == $f2stu) ||
           ($fac3 == $f2stu && $fac4 == $f1stu)
         ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu || $fac3 == $f1stu || $fac4 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu || $fac3 == $f2stu || $fac4 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);

$eqn = Compute("y = y0 + m (x - x0)");

$slope = Compute("$dy / $dx");


##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
(a) Write an equation for the line through 
the point \( (x_0,y_0) \) with slope \( m \)
in point-slope form.  Enter \( x_0 \) as ${BTT}x0${ETT}
and \( y_0 \) as ${BTT}y0${ETT}.  Use 
\( x \) and \( y \) for variable names.
Your equation should be of the form \( y = \ldots\). 
$BR
$BR
\{ ans_rule(30) \} 
\{ AnswerFormatHelp("equations") \}
$BR
$BR
(b)  What is the slope of the line through the points
\( ($x,$y) \) and \( ($p,$q) \)?
$BR
$BR
\( m = \) \{ ans_rule(10) \}
\{ AnswerFormatHelp("numbers") \}
$BR
$BR
(c) Write an equation for the line through the points
\( ($x,$y) \) and \( ($p,$q) \) in point-slope
form.
$BR
$BR
\( y = \) \{ $multians->ans_rule(10) \} + \{ $multians->ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


ANS( $eqn->cmp() );

ANS( $slope->cmp() );

ANS( $multians->cmp() );

ENDDOCUMENT();  
