# DESCRIPTION
# This problem involves finding the moment generating function of the sum of
# identically, independently distributed binomial random variables, and then
# finding the resulting distribution of the sum.
# WebWork problem written by Vladislav Emelyanov
# ENDDESCRIPTION

## DBsubject('Probability')
## DBchapter('Random variables')
## DBsection('Generating function')
## KEYWORDS('moments','expected value','expectation','standard deviation','variance', 'sampling distributions','binomial distribution','binomial random variable')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Vladislav Emelyanov')
## Institution('UMASS-Amherst-Math-Stat')
## Date('3/31/2017')

DOCUMENT(); 
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 0;

Context("Numeric");
Context()->variables->are(binomial=>Real, m=>"Real", n=>"Real", p=>"Real", t=>"Real", q=>"Real");

$mc = new_checkbox_multiple_choice();
$mc -> qa (
    "Select all that apply.  There may be more than one correct answer.", 
    "\( binomial(m*n,p) \)$BR",
);
$mc -> extra(
    "\( binomial(n, m*p) \)$BR",
    "\( negative \) \( binomial(m*n,p) \)$BR",
    "\( negative \) \( binomial(n,m*p) \)$BR",
    "\( negative \) \( binomial(m,p) \)$BR",
    "\( negative \) \( binomial(n,p) \)$BR",
    "\( binomial(n,p) \)$BR",
    "\( binomial(m,p) \)$BR",
);
$mc -> makeLast("None of the above");

Context()->texStrings;
BEGIN_TEXT
If \( Y =\sum_{i=1}^{m} X_i \) and every \( X_i \) is i.i.d with distribution \( binomial(n,p) \), find the MGF of \( Y \).
$BR
\( M(t)= \)  \{ans_rule(20) \} 
$BR
What is the distribution of \( Y \)?
$BR
\{ $mc -> print_q() \}
$BR
\{ $mc -> print_a() \}
END_TEXT

$mgf1 = Formula("(p e^t + q)^(m*n)");
$mgf2 = Formula("(p e^t + (1-p))^(m*n)");


ANS($mgf2 ->cmp(
   checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ($ans) = $student->value;
      if ( $mgf1 == $ans || $mgf2 == $ans ) {
         return 1;
      } else {
         return 0;
      }
   }
));

ANS( checkbox_cmp( $mc->correct_ans() ) );

ENDDOCUMENT();