## DESCRIPTION
## Systems of Linear Equations
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Systems of linear equations)
## DBsection(Systems with 2 variables)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('Algebra' 'Linear Equations' 'Matrix' 'Matrices')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem());

Context('Numeric')->variables->add('y'=>'Real');
Context()->noreduce('(-x)-y');

$a = non_zero_random(-5,5,1);
$b = non_zero_random(-5,5,1);
$x0 = non_zero_random(-5,5,1);

$m = random(-1,1,2) * random(2,5,1);

$c = $m * $a;
$d = $m * $b;
$y0 = $m * $x0;

$x = Compute("x");
$y = Compute("($x0 - $a * x) / $b");

$expr1 = Formula("$a x + $b y")->reduce;
$expr2 = Formula("$c x + $d y")->reduce;


$multians = MultiAnswer($x, $y)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      $s[0] = Formula($s[0]);
      $s[1] = Formula($s[1]);
      if ( ($s[0] == $c[0]) and ($s[1] == $c[1])) {
          return [1,1];
      } elsif ( ($s[1] == Formula("y")) and ($s[0] == Formula("($x0 - $b * y) / $a")) ) {
          return [1,1];
      } else {
          return [0,0];
      }
  }
);

Context()->texStrings;
BEGIN_TEXT
Solve the linear system below by substitution.
\[
\left\lbrace
\begin{array}{rcr}
$expr1 & = & $x0 \\
$expr2 & = & $y0 
\end{array}
\right.
\]
If the answer is a line such as \( y = x + 1 \), enter \( x = x \) and \( y = x+1 \).
If there is no solution, enter ${BITALIC}DNE${EITALIC} for both answers.
$BR
$BR
\( x = \) \{ $multians->ans_rule(15) \}
$BR
\( y = \) \{ $multians->ans_rule(15) \}
END_TEXT
Context()->normalStrings;

ANS($multians->cmp);

COMMENT('MathObject version.  Always infinitely many solutions.');
ENDDOCUMENT();
