# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Applied Calculus')
## TitleText5('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## AuthorText5('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('4')
## EditionText5('7')
## Section1('11.1')
## Section2('11.1')
## Section3('11.1')
## Section4('10.2')
## Section5('11.1')
## Problem1('7')
## Problem2('14')
## Problem3('18')
## Problem4('19')
## Problem5('27')
## KEYWORDS('calculus', 'integral', 'differential equations')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->strings->add("anything"=>{});
$showPartialCorrectAnswers = 1;

$r = random(-0.09,-0.01,0.01);

$ans1 = String("anything");
$ans2 = Compute($r);
$ans3 = Compute(0);
$ans4 = String("anything");

$multipart = MultiAnswer($ans1, $ans2, $ans3, $ans4)->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checkTypes => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $a1stu, $a2stu, $a3stu, $a4stu ) = @{$student};
        my ( $a1, $a2, $a3, $a4 ) = @{$correct};
	my @ret = ( 0, 0, 0, 0 );

	if ( ref($a1stu) eq ref($a1) && $a1stu == $a1 ) {
	    @ret = ( 1, 0, 0, 0 );
	    $ret[1] = 1 if ( ref($a2stu) eq ref($a2) && $a2stu == $a2 );
	    $ret[2] = 1 if ( ref($a3stu) eq ref($a3) && $a3stu == $a3 );
	    $ret[3] = 1 if ( ref($a4stu) eq ref($a4) && $a4stu == $a4 );

	    ## can we get a better score by remarking the second two?
	    if ( ( ref($a3stu) eq ref($a3) && $a3stu == $a3 ) ||
		 ( ref($a3stu) eq ref($a1) && $a3stu == $a1 ) ) {
		my @altret = ( 0, 0, 1, 0 );
		if ( ref($a3stu) eq ref($a1) && $a3stu == $a1 ) {
		    $altret[1] = 1 if ( ref($a2stu) eq ref($a4) &&
					$a2stu == $a4 );
		    $altret[3] = 1 if ( ref($a4stu) eq ref($a2) &&
					$a4stu == $a2 );
		} else {
		    $altret[1] = 1 if ( ref($a2stu) eq ref($a2) &&
					$a2stu == $a2 );
		    $altret[3] = 1 if ( ref($a4stu) eq ref($a4) &&
					$a4stu == $a4 );
		}
		@ret = @altret if ( $ret[0] + $ret[1] + $ret[2] + $ret[3] <
				    $altret[0] + $altret[1] + $altret[2] +
				    $altret[3]);
	    }
	} elsif ( ref($a1stu) eq ref($a3) && $a1stu == $a3 ) {
	    @ret = ( 1, 0, 0, 0 );
	    $ret[1] = 1 if ( ref($a2stu) eq ref($a4) && $a2stu == $a4 );
	    $ret[2] = 1 if ( ref($a3stu) eq ref($a1) && $a3stu == $a1 );
	    $ret[3] = 1 if ( ref($a4stu) eq ref($a2) && $a4stu == $a2 );

	    ## can we get a better score by remarking the second two?
	    if ( ( ref($a3stu) eq ref($a3) && $a3stu == $a3 ) ||
		 ( ref($a3stu) eq ref($a1) && $a3stu == $a1 ) ) {
		my @altret = ( 0, 0, 1, 0 );
		if ( ref($a3stu) eq ref($a1) && $a3stu == $a1 ) {
		    $altret[1] = 1 if ( ref($a2stu) eq ref($a4) &&
					$a2stu == $a4 );
		    $altret[3] = 1 if ( ref($a4stu) eq ref($a2) &&
					$a4stu == $a2 );
		} else {
		    $altret[1] = 1 if ( ref($a2stu) eq ref($a2) &&
					$a2stu == $a2 );
		    $altret[3] = 1 if ( ref($a4stu) eq ref($a4) &&
					$a4stu == $a4 );
		}
		@ret = @altret if ( $ret[0] + $ret[1] + $ret[2] + $ret[3] <
				    $altret[0] + $altret[1] + $altret[2] +
				    $altret[3]);
	    }
	} else {
	    ## the first answer blank isn't correct; check to see if the
	    ##    third and fourth (and second) are correct
	    if ( ref($a3stu) eq ref($a1) && $a3stu == $a1 ) {
		@ret = ( 0, 0, 1, 0 );
		$ret[1] = 1 if ( ref($a2stu) eq ref($a4) && $a2stu == $a4 );
		$ret[3] = 1 if ( ref($a4stu) eq ref($a2) && $a4stu == $a2 );
	    } elsif ( ref($a3stu) eq ref($a3) && $a3stu == $a3 ) {
		@ret = ( 0, 0, 1, 0 );
		$ret[1] = 1 if ( ref($a2stu) eq ref($a2) && $a2stu == $a2 );
		$ret[3] = 1 if ( ref($a4stu) eq ref($a4) && $a4stu == $a4 );
	    } else {
		## neither of the first or third answer blanks are correct;
		##    check to see if the second or fourth are independently
		##    correct
		if ( ref($a2stu) eq ref($a2) && $a2stu == $a2 ) {
		    @ret = ( 0, 1, 0, 0 );
		    $ret[3] = 1 if (ref($a4stu) eq ref($a4) && $a4stu == $a4);
		} elsif ( ref($a2stu) eq ref($a4) && $a2stu == $a4 ) {
		    @ret = ( 0, 1, 0, 0 );
		    $ret[3] = 1 if (ref($a4stu) eq ref($a2) && $a4stu == $a2);
		} elsif ( ( ref($a4stu) eq ref($a2) && $a4stu == $a2 ) ||
			  ( ref($a4stu) eq ref($a4) && $a4stu == $a4 ) ) {
		    @ret = ( 0, 0, 0, 1 );
		}
	    }
	}
	return [ @ret ];
    }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Suppose \( Q=Ce^{kt} \) satisfies the differential equation
\[
\frac{dQ}{dt}=$r Q.
\]
What (if anything) does this tell you about the values of \(C\) and \(k\)?

$PAR
Either:
$BR
\(C = \) \{ $multipart->ans_rule(25) \}, and
$BR
\(k = \) \{ $multipart->ans_rule(25) \}
$BR
${BITALIC}(For each of C and k, enter the value that the parameter must
take, or enter the word ${BBOLD}anything$EBOLD if there are no restrictions
on the value.)$EITALIC

$PAR
or:
$BR
\(C = \) \{ $multipart->ans_rule(25) \}, and
$BR
\(k = \) \{ $multipart->ans_rule(25) \}
$BR
${BITALIC}(For each of C and k, enter the value that the parameter must
take, or enter the word ${BBOLD}anything$EBOLD if there are no restrictions
on the value.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS($multipart->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

If \(Q=Ce^{kt}\), then
\[
\frac{dQ}{dt} = Cke^{kt} = k (Ce^{kt}) = kQ.
\]
We are given that \(\frac{dQ}{dt}=$r Q\), so we know that
\(kQ = $r Q\).
Thus we either have \(Q=0\) (in which case \(C=0\) and \(k\) is
anything) or \(k=$r\).  Notice that if \(k=$r\), then \(C\)
can be any number.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
