# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Integrating factor)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('11.1')
## Section2(11.1)
## Section3(11.1)
## Section4('11.1')
## Problem1('13')
## Problem2(17)
## Problem3(21)
## Problem4('30')
## KEYWORDS('calculus', 'integral', 'differential equations')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( C=>"Real", n=>"Real" );

$showPartialCorrectAnswers = 1;

$a = random(2,4,1);
$x0 = random(2,5,1);
$y0 = random(10,50,5);

$ans0 = Formula("0 x^n");
$ans0->{test_points} = [ [-1.513,2,-1.125], [0,1,0], [0.3375,2,0.5153], 
			 [1.5,1,1.5], [2.115,2,2.115] ];
# $ans0->{test_points} = [ 0.3375, 1.5, 2.115, 3.5 ];
$ans1 = Formula("C x^($a)");

$multipart = MultiAnswer($ans0, $ans1)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f2stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

${BBOLD}(a)$EBOLD
Consider the differential equation
\[ x\frac{dy}{dx} - $a y = 0. \]
$BR
Find a general solution to this differential equation that
has the form \( y = Cx^n \).
$BR
\( y = \) \{ ans_rule(35) \}.
$BR
Find a second solution \( y = Cx^n \) that might not be a
general solution and which may have a different value of \(n\)
than your first solution.
$BR
\( y = \) \{ ans_rule(35) \}.

$PAR
${BBOLD}(b)$EBOLD
If the solution additionally satisfies  \(y=$y0\) when \(x=$x0\),
what is the solution?
$BR
\(y = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

# ANS($multipart->cmp() );
ANS($ans1->cmp() );
ANS($ans0->cmp() );
ANS(Compute("$y0*(x/$x0)^($a)")->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
If \(y=C x^n\) is a solution to the given differential equation,
then we must have
\[
x \frac{d \ (C x^n)}{dx} - $a (C x^n) = 0,
\]
so
\[
x(C n x^{n-1}) - $a (Cx^n) = 0,
\]
or
\[
Cnx^n - $a Cx^n = C(n - $a) x^n = 0.
\]
Thus, if \(C = 0\) we have a solution, \(y = 0\); and if \(C\ne0\)
then we must have \(n=$a\), \(C\) may have any value (other than zero),
and the solution is \(y = C x^{$a}\).

$PAR
${BBOLD}(b)$EBOLD
If \(y=$y0\) for \(x=$x0\), we cannot have \(C=0\). Thus, by part (a),
we get \(n=$a\).  The solution to the differential equation is
\(y = C x^{$a}\).  Thus, if \(y = $y0\) when \(x = $x0\),
\[ $y0 = C ($x0)^{$a}, \]
so that \(C = \frac{$y0}{($x0)^{$a}}\), and
\(y = \frac{$y0}{($x0)^{$a}}\cdot x^{$a}\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
