# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 01/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Extreme values and optimization)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(ListCritPtsAndClassify)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('15.1')
## Section2('15.1')
## Section3('15.1')
## Problem1('14')
## Problem2('19')
## Problem3('25')
## KEYWORDS('extrema', 'multivariable', 'functions', 'calculus', 'derivatives')

## Textbook tags
## HHChapter1('Optimization: Local and Global Extrema')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Point");
Context()->strings->add( maximum=>{}, minimum=>{}, 'saddle point'=>{},
			 none=>{} );
$showPartialCorrectAnswers = 1;

$m = list_random( 1, 4 );
$sqrtm = sqrt($m);

$b = random(1,5,1);
if ( $b == 2 ) {
    $bcoef = "1/2";
} elsif ( $b == 4 ) {
    $bcoef = "1";
} else {
    $bcoef = "$b/4";
}
$a = $m*8*$b;

$func = Compute( "$a*x*y - ($bcoef)*(x + y)^4" )->reduce();

$critPts = List( Point("(0,0)"), Point("($sqrtm,$sqrtm)"),
		 Point("(-$sqrtm,-$sqrtm)") );
$classify = List( String("saddle point"), String("maximum"),
		  String("maximum") );

$anschek = MultiAnswer( $critPts, $classify )->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checkTypes => 0,
    checker => sub {
	my ( $cor, $stu, $self ) = @_;
	my ($corList, $corClass) = @$cor;
	my ($stuList, $stuClass) = @$stu;
	$stuList = List( $stuList ) if ( ref( $stuList ) ne 'Value::List' );
	$stuClass = List( $stuClass ) if ( ref( $stuClass ) ne 'Value::List' );

	## now $corList and $stuList should be lists of points,
	##    and the class veriables lists of strings
	my @corPts = $corList->value;
	my @stuPts = $stuList->value;
	my @corStr = $corClass->value;
	my @stuStr = $stuClass->value;

	if ( @corPts != @stuPts ) {
	    $self->setMessage( 1, "You have given too " .
			       ((@corPts > @stuPts) ? "few " : "many ") .
			       "points." );
	}

	my @ret = ( 0, 0 );
	my $numpts = ( @corPts > @stuPts ) ? @corPts : @stuPts;
	my $retInc = 1/$numpts;

	my $i = 0; my $j = 0;
	my %matched = ();
	foreach my $cpt ( @corPts ) {
	    $j = 0;
	    foreach my $spt ( @stuPts ) {
		$matched{$j} = 0 if ( ! defined( $matched{$j} ) );
		if ( $matched{$j} ) {
		    $j++;
		    next;
		}

		if ( ref($cpt) eq ref($spt) && $cpt == $spt  ) {
		    $ret[0] += $retInc;
		    $matched{$j} = 1;

		    if ( $j < @stuStr ) {
			if ( ref($corStr[$i]) eq ref($stuStr[$j]) &&
			     $corStr[$i] == $stuStr[$j] ) {
			    $ret[1] += $retInc;
			}
		    }
		}
		$j++;
	    }
	    $i++;
	}
	$ret[0] = sprintf("%.2f", $ret[0]);
	$ret[1] = sprintf("%.2f", $ret[1]);
	return [ @ret ];
    } );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find the critical points for the function
\[ f(x,y) = $func \]
and classify each as a local maximum, local minimum, saddle point, or
none of these.
$PAR
critical points: \{ $anschek->ans_rule(55) \}
$BR
${BITALIC}(give your points as a comma separated list of (x,y)
coordinates.)$EITALIC
$PAR
classifications: \{ $anschek->ans_rule(55) \}
$BR
${BITALIC}(give your answers in a comma separated list, specifying
${BBOLD}maximum$EBOLD, ${BBOLD}minimum$EBOLD, ${BBOLD}saddle point$EBOLD,
or ${BBOLD}none$EBOLD for each, in the same order as you entered your
critical points)$EITALIC

END_TEXT
Context()->normalStrings;

ANS($anschek->cmp() );

$eightb = 8*$b;
$threeb = 3*$b;
$ninebsq = $threeb*$threeb;
$d00 = -1*$a*$a;
$d11 = $ninebsq*16*$m*$m - ($a - $threeb*4*$m)**2;
$fxx11 = -1*$threeb*4*$m;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

At a critical point, \(f_x=0\) and \(f_y=0\).  Thus
\[ f_x = $a y - $b(x+y)^3 = 0\]
and
\[ f_y = $a x - $b (x+y)^3 = 0. \]
Thus
\[ $a y = $b(x + y)^3 = $a x, \]
and \(x = y\).  Thus \($a x + $eightb x^3 = 0\), and solving gives
\(x = 0, \pm $sqrtm\).  Thus the critical points are
\((0,0)\), \(($sqrtm,$sqrtm)\) and \((-$sqrtm,-$sqrtm)\).

$PAR
Next,
\( f_{yy} = f_{xx} = -$threeb (x+y)^2\), and
\( f_{xy} = $a - $threeb(x+y)^2\).
The discriminant is
\( D(x,y) = f_{xx}f_{yy}- f_{xy}^2\), so
\[ D = $ninebsq(x+y)^4 - \left($a - $threeb(x + y)^2\right)^2. \]
Thus
\( D(0,0) = $d00 < 0\), so \((0,0)\) is a saddle point.
$BR
\( D($sqrtm,$sqrtm) = $d11 > 0\) and
\(f_{xx}($sqrtm,$sqrtm) = $fxx11 < 0\), so \(($sqrtm,$sqrtm)\) is a local
maximum.
$BR
\( D(-$sqrtm,-$sqrtm) = $d11 > 0\) and
\(f_{xx}(-$sqrtm,-$sqrtm) = $fxx11 < 0\), so \((-$sqrtm,-$sqrtm)\) is a local
maximum.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
