## DESCRIPTION
##   Use Substitution to Evaluate an Indefinite Integral
## ENDDESCRIPTION

## Tagged by nhamblet

## DBsubject(Calculus - single variable)
## DBchapter(Integrals)
## DBsection(Indefinite integrals (without trigonometric functions))
## Date(8/23/07)
## Institution(Union)
## Author(K. Lesh)
## MLT(indefinite_kx_exponential_only)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Anton')
## EditionText1('7')
## Section1('6.3')
## Problem1('01')
## TitleText2('Calculus: Early Transcendentals')
## AuthorText2('Rogawski')
## EditionText2('1')
## Section2('5.7')
## Problem2('39')
## KEYWORDS('Indefinite', 'Integral', 'Substitution')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

###################################
# Setup

Context("Numeric");
$a = random(2,10,1);

$integrand=Formula("e^{$a x}");

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT
Evaluate the indefinite integral.
$PAR
\( \displaystyle\int $integrand \, dx \)
             = \{ans_rule(50)\} \( + C\).
END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;
Context()->flags->set(reduceConstants=>0);

$antideriv = Formula("(1/$a)e^(${a}x)");

# check the derivative of the student's answer

ANS($antideriv->cmp( checker => sub{
  my ($correct,$student) = @_;
  my $df = $student->D('x');
  return $integrand == $df;
}));

###################################

ENDDOCUMENT();
