##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Partial differential equations)
## DBsection(Heat equation)
## Date(7/25/2011)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(5)
## TitleText1('Differential Equations')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('12.3')
## Problem1('')
## KEYWORDS('PDE', 'heat', 'equation')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "answerHints.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;
##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->functions->undefine("sin","cos","tan","sinh","cosh","tanh");
Context()->variables->add(a=>"Real",b=>"Real",alpha=>"Real");
Context()->variables->set(alpha => {TeX => '{\alpha}'});
Context()->functions->undefine("acos","acosh");
Context()->{error}{msg}{"Function 'acosh' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cosh function, you need to put the multiplication symbol between the a and the cosh";
Context()->{error}{msg}{"Function 'acos' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cos function, you need to put the multiplication symbol between the a and the cos";
$simplify = Context()->copy;

Context("Numeric");
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("X"=>"Real","X'"=>"Real","X''"=>"Real","T"=>"Real","T'"=>"Real","L"=>"Real","lambda"=>"Real","a"=>"Real","k"=>"Real");
Context()->variables->set(lambda => {TeX => '{\lambda}'});
Context()->functions->undefine("acos","acosh");
Context()->{error}{msg}{"Function 'acosh' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cosh function, you need to put the multiplication symbol between the a and the cosh";
Context()->{error}{msg}{"Function 'acos' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cos function, you need to put the multiplication symbol between the a and the cos";
$DEcontext = Context()->copy;

Context("Numeric");
Context()->variables->are("x"=>"Real","y"=>"Real","t"=>"Real","L"=>"Real","n"=>"Real","alpha"=>"Real","a"=>"Real",b=>"Real","c"=>"Real","d"=>"Real","k"=>"Real");
Context()->variables->set(alpha => {TeX => '{\alpha}'});
Context()->functions->undefine("acos","acosh");
Context()->{error}{msg}{"Function 'acosh' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cosh function, you need to put the multiplication symbol between the a and the cosh";
Context()->{error}{msg}{"Function 'acos' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cos function, you need to put the multiplication symbol between the a and the cos";
##########
##
## Tranlate PDE to ODE with IC

$Xdiff1 = Formula($DEcontext,"X''/X")->reduce;
$Tdiff1 = Formula($DEcontext,"T'/(k*T)")->reduce;

$Xdiff2 = Formula($DEcontext,"X''+lambda*X")->reduce;
$Tdiff2 = Formula($DEcontext,"T'+k*lambda*T")->reduce;

$left = Formula($DEcontext,"0");
$right = Formula($DEcontext,"L");
$homogeneousleft = Formula($DEcontext,"0");
$homogeneousright = Formula($DEcontext,"0");

$multians = MultiAnswer($left,$homogeneousleft,$right,$homogeneousright)->with(
  singleResult => 0,
  checker => sub 
    {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2,$stu3,$stu4 ) = @{$student};
      my ( $c1, $c2, $c3, $c4 ) = @{$correct};
      if ( ($c1 == $stu1) && ($c3 == $stu3) ) 
         {if ( ($c2 == $stu2) && ($c4 == $stu4) )
            {return [1,1,1,1];}
          elsif ($c2 == $stu2) {return [1,1,1,0]}
          elsif ($c4 == $stu4) {return [1,0,1,1]}
          else {return [1,0,1,0]};
         }
      elsif ( ($c3 == $stu1) && ($c1 == $stu3) )
         {if ( ($c4 == $stu2) && ($c2 == $stu4) )
             {return [1,1,1,1];}
         elsif ($c4 == $stu2) {return [1,1,1,0]}
         elsif ($c2 == $stu4) {return [1,0,1,1]}
         else {return [1,0,1,0]};
         }
      else {return [0,0,0,0];}
  }
);

## Sturm-Louiville problem
$sn = Formula("sin(n*pi*x/L)");
$cn = Formula("cos(n*pi*x/L)");

$Xequal = Formula("a+bx");
$Xequalprime = $Xequal->D('x');

$Xless = Formula("a*cosh(alpha*x)+b*sinh(alpha*x)");
$Xlessprime = $Xless->D('x');
$Xlessleft = Formula($simplify,"b*alpha");
$Xsimplified = Formula("a*cosh(alpha*x)");
$Xlessright = Formula("a*alpha*sinh(alpha*L)");

$Xgreater =  Formula("a*cos(alpha*x)+b*sin(alpha*x)");
$Xgreaterprime = $Xgreater->D('x');
$Xgreaterleft = Formula($simplify,"b*alpha");
$Xcos = Formula("a*cos(alpha*x)");
$Xgreaterright = $Xgreaterprime->substitute(x=>"L");
$Tgreater = Formula("c*e^(-k*(n*pi/L)^2*t)");
$ugreater = $cn*$Tgreater;

$un = Formula("cos(n*pi*x/L)*e^(-k*(n*pi/L)^2*t)");

$f1 = Formula("1");
$f2 = Formula("0");

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
For the general solution of the differential equation in the following cases use a and b for your constants and list the functions in alphabetical order, for example \(y = a\cos (x) + b \sin(x)\). For the variable \(\lambda\) type the word lambda, for \(\alpha\) type alpha, otherwise treat these as you would any other variable.
$BR
$BR
$BR
Solve the heat equation 

\[\displaystyle
k\frac{\partial^2 u}{\partial x^2} = \frac{\partial u}{\partial t}, 0< x < L, t > 0\]

\[\displaystyle\frac{\partial u}{\partial x}(0,t) = 0, \frac{\partial u}{\partial x}(L,t) = 0, t > 0\]

\[\displaystyle
u(x,0) = 
\begin{cases}
$f1, &0 < x < \frac{L}{2}\\
$f2, &\frac{L}{2}<x<L\\
\end{cases}
\]
$BR
$BR
This models the heat equation in a long thin rod with insulated ends. Since we have not solved this boundary value problem, we have to start from scratch
$BR
We look for solutions of the form \(u(x,t) = X(x)T(t)\). 
$BR
The PDE can be rewritten using this solution as (placing constants in the DE for T) into 
$BR
\{$Xdiff1->ans_rule\} = \{$Tdiff1->ans_rule\} = \(-\lambda\)
$BR
$BR
$BBOLD Note: $EBOLD Use the prime notation for derivatives, so the derivative of \(X\) is written as \(X^\prime\). Do NOT use \(X^\prime(x)\)
$BR
The boundary conditions in the PDE translate into boundary conditions for the differential equation in X:
$BR

Since these differential equations are independent of each other, they can be separated 
$BR
DE in X: \{$Xdiff2->ans_rule\} \( = 0\)
$BR BCs in X: $BBOLD Note: $EBOLD You have to input all four answers to this question or WeBWorK will mark it wrong.
$BR
$BCENTER
\(X^\prime(\)\{$multians->ans_rule(1)\}\() = \)\{$multians->ans_rule(1)\}
$BR
\(X^\prime(\)\{$multians->ans_rule(1)\}\() = \)\{$multians->ans_rule(1)\}
$ECENTER
$BR
DE in T: \{$Tdiff2->ans_rule\} \( = 0\)
$BR
$BR
This leads us to a Sturm-Louiville problem. In each case the general solution in X is written with constants a and b and the general solution in Y is written with constants c and d.
$BR
$BBOLD Case 1: $EBOLD \(\lambda = 0\)
$BR
\(X(x) = \) \{$Xequal->ans_rule(70)\}
$BR
Using this formula
$BR
$BCENTER
\(X^\prime(0) = \)\{ans_rule(1)\} \( = 0\)
$BR
\(X^\prime(L) = \)\{ans_rule(1)\} \( = 0\)
$ECENTER
$BR
Making X a constant function. We get the constant solution another way, so we will discard this case.
$BR
$BR
$BBOLD Case 2: $EBOLD \(\lambda = -\alpha^2\)
$BR
\(X(x) = \) \{$Xless->ans_rule(70)\}
$BR
\(X^\prime(0) = \)\{ans_rule\} \( = 0\)
$BR
That simplifies X down to \{ans_rule\}. Using the simplified X:
$BR
\(X^\prime(L) = \)\{ans_rule\} \( = 0\)
$BR
Therefore \(a = \) \{ans_rule\} and
$BR
$BR
\(X = \) \{ans_rule\}
$BR
$BR
$BBOLD Case 3: $EBOLD \(\lambda = \alpha^2 \)
$BR
\(X(x) = \) \{$Xgreater->ans_rule(70)\}
$BR
\(X^\prime(0) = \)\{ans_rule\} \( = 0\)
$BR
So X simplifies to \{$Xcos->ans_rule\}. Using the simplified X:
$BR
\(X^\prime(L) = \)\{ans_rule\} \( = 0\)
$BR
Which leads us to the eigenvalues \(\alpha =\)\{ans_rule\}
$BR 
$BBOLD Note $EBOLD n = 0 is the constant from above, including it would have been redundant.
$BR
Since this case has nontrivial solutions we now solve the differential equation in T using c as the constant: 
$BR
\(T = \) \{$Tgreater->ans_rule\}
$BR
So for each n the product \(u = XT = \)\{$un->ans_rule(30)\} (without constants)is a solution of the PDE with boundary conditions: 
$BR

\(\displaystyle k\frac{\partial^2 u}{\partial x^2} = \frac{\partial u}{\partial t}, 0< x < L, t > 0\)
\(\displaystyle\frac{\partial u}{\partial x}(0,t) = 0, \frac{\partial u}{\partial x}(L,t) = 0, t > 0\)

$BR
$BR
To get the initial conditions we need a series solution of the form
$BR
\(u(x,t)=\displaystyle\sum\limits_{n=0}^\infty A_n\)\{$un->ans_rule(30)\}
$BR
\(u(x,0)=\displaystyle\sum\limits_{n=0}^\infty A_n\)\{$cn->ans_rule(30)\}\( = \begin{cases}
$f1, &0 < x < \frac{L}{2}\\
$f2, &\frac{L}{2}<x<L\\
\end{cases}\)
$BR
which we recognize as a Fourier series:
$BR 
\(\displaystyle
A_n=\frac{2}{L}\int_0^{\frac{L}{2}}\)\{ans_rule(30)\} \(\hskip 3pt dx\)
$BR
$BR
\(\hskip 14pt =\) \{ans_rule(30)\}


END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

#####
# Translate to ODE
#
ANS($Xdiff1->cmp);
ANS($Tdiff1->cmp);
ANS($Xdiff2->cmp);
ANS( $multians->cmp() ); 
ANS($Tdiff2->cmp);
## Sturm-Louiville
ANS($Xequal->cmp);
ANS(Formula("b")->cmp);
ANS(Formula("b")->cmp);

ANS($Xless->cmp);
ANS($Xlessleft->cmp);
ANS($Xsimplified->cmp);
ANS($Xlessright->cmp);
ANS(Compute("0")->cmp);
ANS(Compute("0")->cmp);


ANS($Xgreater->cmp);
ANS($Xgreaterleft->cmp);
ANS($Xcos->cmp);
ANS(Compute("-a*alpha*sin(alpha*L)")->cmp);
ANS(Compute("n*pi/L")->cmp);
ANS($Tgreater->cmp);

ANS($un->cmp());
ANS($un->cmp());
ANS($cn->cmp());

# Compute a_n
ANS(Compute("$f1*$cn")->cmp); 
ANS(Compute("2/(n*pi)sin(n*pi/2)")->cmp);
ENDDOCUMENT();        
