## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Complex eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Complex-Matrix");

#$Q = GL2Z();

$A = Matrix([
[1,-2],
[1,3],
]);

#$A = $Q * $A * ($Q->inverse);

$e1 = 2 + 1 * i;
$e2 = 2 - 1 * i;

#$V = Matrix([
#[2,2],
#[-1-1*i,-1+1*i],
#]);

#$QV = $Q * $V;

#$v1 = $QV->column(1);
#$v2 = $QV->column(2);

$v1 = Matrix([[2],[-1-1*i]]);
$v2 = Matrix([[2],[-1+1*i]]);

Context()->texStrings;
BEGIN_TEXT
Find the eigenvalues and eigenvectors for \( A = $A \).
$BR
$BR
The eigenvalue \( a + b i = \) \{ $e1->ans_rule(5) \} has an eigenvector \{ $v1->ans_array(5) \}.
$BR
$BR
The eigenvalue \( a - b i = \) \{ $e2->ans_rule(5) \} has an eigenvector \{ $v2->ans_array(5) \}.
END_TEXT
Context()->normalStrings;

ANS( $e1->cmp );
ANS( $v1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0]]);
return $A * $student == $e1 * $student;
}) );

ANS( $e2->cmp );
ANS( $v2->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0]]);
return $A * $student == $e2 * $student;
}) );

COMMENT('MathObject version.  Not randomized.  NEEDS MORE WORK!');

ENDDOCUMENT();
