## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Complex eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Complex-Matrix");

$A = Matrix([
[5,-5],
[1,1],
]);

$e1 = 3 + 1 * i;
$e2 = 3 - 1 * i;

$v1 = Matrix([[5],[2-1*i]]);
$v2 = Matrix([[5],[2+1*i]]);

Context()->texStrings;
BEGIN_TEXT
Suppose \( f : \mathbb{R}^2 \to \mathbb{R}^2 \) is the linear transformation 
defined by \( f(x,y) = \langle 5x-5y, x+y \rangle \).
Find the eigenvalues and eigenvectors for \( f \).
$BR
$BR
The eigenvalue \( a + b i = \) \{ $e1->ans_rule(5) \} has an eigenvector \{ $v1->ans_array(5) \}.
$BR
$BR
The eigenvalue \( a - b i = \) \{ $e2->ans_rule(5) \} has an eigenvector \{ $v2->ans_array(5) \}.
END_TEXT
Context()->normalStrings;

ANS( $e1->cmp );
ANS( $v1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0]]);
return $A * $student == $e1 * $student;
}) );

ANS( $e2->cmp );
ANS( $v2->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0]]);
return $A * $student == $e2 * $student;
}) );

COMMENT('MathObject version.  Not randomized.  NEEDS MORE WORK!');

ENDDOCUMENT();
