## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Complex eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Complex-Matrix");

$A = Matrix([
[-1,3,-4],
[-2,3,-4],
[1,1,3]
]);

$e1 = 1 + 0 * i;
$e2 = 2 - 3 * i;
$e3 = 2 + 3 * i;

$v1 = Matrix([[2],[0],[-1]]);
$v2 = Matrix([[-1 - 5 * i],[-2 - 4 * i],[3]]);
$v3 = Matrix([[-1 + 5 * i],[-2 + 4 * i],[3]]);

$ans3a = List(-13,-1,1,13);
$ans3b = List(1);

Context()->texStrings;
BEGIN_TEXT
Suppose \[ A = $A. \]
$BR
\{ BeginList('OL',type=>'a') \}
$ITEM
Find the eigenvalue corresponding to the eigenvector \( $v2 \).
Note: you should be able to calculate this eigenvalue without 
using a determinant or a characteristic polynomial or a calculator. 
$BR
The eigenvalue is \{ $e1->ans_rule(5) \}.

$ITEMSEP
$ITEM
Using your answer above, find another eigenvalue and eigenvector for \( A \).
$BR
The eigenvalue \{ $e3->ans_rule(5) \} has an associated eigenvector \{ $v3->ans_array(5) \}.

$ITEMSEP
$ITEM
The characteristic polynomial for \( A \) is \( p(\lambda) = - \lambda^3 + 5 \lambda^2 - 17 \lambda + 13 \).
The rational root theorem says that the possible rational roots of \( p(\lambda) \) are 
\{ ans_rule(20) \} (enter a list of numbers).  
One of these possible rational roots is actually a root of \( p(\lambda) \).  
Find this rational root and its associated eigenvector.
$BR
The eigenvalue \{ ans_rule(10) \} has an associated eigenvector \{ $v1->ans_array(5) \}.

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;

ANS( $e2->cmp );


ANS( $e3->cmp );
ANS( $v3->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0],[0]]);
return $A * $student == $e3 * $student;
}) );

ANS( $ans3a->cmp );
ANS( $ans3b->cmp );

ANS( $v1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return 0 if Matrix($student) == Matrix([[0],[0],[0]]);
return $A * $student == $e1 * $student;
}) );

COMMENT('MathObject version.  Not randomized.  NEEDS MORE WORK!');

ENDDOCUMENT();
