## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
install_problem_grader(~~&std_problem_grader);

Context('Matrix');

$P = GL2Z();

$D = Matrix([
[random(2,3,1),0],
[0,random(-2,-1,1)],
]);

$e1 = $D->element(1,1);
$e2 = $D->element(2,2);

$Dalt = Matrix([[$e2,0],[0,$e1]]);

$A = $P * $D * ($P->inverse);

$multians = MultiAnswer($P, $D)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = (0,0); # (0) x 4; # initialize with zeros
      
      $score[1] = 1 if (($s[1] == $D) or ($s[1] == $Dalt));
      if ($s[0]->det != 0) {
        $score[0] = 1 if (($A == $s[0] * $s[1] * ($s[0]->inverse)) and ($score[1] == 1));
      }

      if ( $score[1] == 0 ) {
         $self->setMessage(1,"Since D was incorrect, P was not checked for correctness");
      }

      return @score;
  }
);


$popup = PopUp(['choose','diagonalizable','not diagonalizable'],'diagonalizable');

Context()->texStrings;
BEGIN_TEXT
Let \[ A = $A. \]
If possible, find an invertible matrix \( P \) so that \( D = P^{-1} A P \) is a diagonal matrix.
If it is not possible, enter the identity matrix for \( P \) and the matrix \( A \) for \( D \).
You must enter a number in every answer blank for the answer evaluator to work properly. 
$BR
$BR
\( P = \) \{ $multians->ans_array(5) \}.
$BR
$BR
\( D = \) \{ $multians->ans_array(5) \}.
$BR
$BR
Is \( A \) diagonalizable over \( \mathbb{R} \)? \{ $popup->menu() \}  Be sure you can explain why or why not.
END_TEXT
Context()->normalStrings;

ANS( $multians->cmp );
ANS( $popup->cmp );

COMMENT('MathObject version.  Always diagonalizable.');

ENDDOCUMENT();
