##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Algebra')
## DBchapter('Rational equations and functions')
## DBsection('Rational inequalities')
## Date('7/3/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "PGgraphmacros.pl",
   "PGgraders.pl",
   "contextFraction.pl",
   "parserMultiAnswer.pl",
   "parserPopUp.pl",
   "scaffold.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

# Give full credit if the student is able to complete the
# quotient without doing the scaffold steps...
install_problem_grader(~~&full_partial_grader);

##############################################################
#
#  Setup
#
#
Context("Fraction");
Context()->noreduce('(-x)-y');
Parser::Number::NoDecimals();
Context()->{error}{msg}{"You are not allowed to type decimal numbers in this problem"} 
  = "Decimals are not allowed for this problem, please use exact forms (such as fractions) instead of decimal approximations.";

$r[0] = random(-1,1,2)*random(5,14,1);
do { $r[1] = non_zero_random(-7,7,1); } until ( abs($r[1]) != abs($r[0]) );
do { $a[0] = list_random(2,3,5); } until (gcf(abs($r[0]),$a[0])==1);
do { $a[1] = list_random(2,3,5); } until (gcf(abs($r[1]),$a[1])==1);

# select one of the leading coeff to be negative
$nc = random(0,1,1);
$a[0] = $a[0]*(-1)**$nc;
$a[1] = $a[1]*(-1)**($nc+1);

@groot = ($r[0]/$a[0],$r[1]/$a[1]);

if ($groot[0]*$groot[1] > 0) { $r[0] = -1*$r[0]; $groot[0] = -1*$groot[0] };

$n = Formula("$a[0] x - $r[0]")->reduce;
$d = Formula("$a[1] x - $r[1]")->reduce;
$rational = Formula("$n/$d");

@root = (Fraction($r[0],$a[0]),Fraction($r[1],$a[1]));

($min,$max) = ($r[0]/$a[0]<$r[1]/$a[1]) ?
    (Fraction($r[0],$a[0]),Fraction($r[1],$a[1])):
    (Fraction($r[1],$a[1]),Fraction($r[0],$a[0]));
    
$asyLeft = $groot[1]-.001;
$asyRight = $groot[1]+.001;

($minReal,$maxReal) = ($r[0]/$a[0]<$r[1]/$a[1]) ?
    ($r[0]/$a[0],$r[1]/$a[1]):
    ($r[1]/$a[1],$r[0]/$a[0]);



#graph of rational:
$gr = init_graph(-10,-10,10,10,grid=>[20,20],axes=>[0,0],size=>[400,400]);
add_functions($gr,"$rational for x in <-10,$asyLeft> using color:blue and weight:2");
add_functions($gr,"$rational for x in <$asyRight,10> using color:blue and weight:2");
($gr->fn)[0]->steps(100);
($gr->fn)[1]->steps(100);
$gr->moveTo($groot[1],-10);
$gr->lineTo($groot[1],10,'orange',2,'dashed');
$gr->stamps(closed_circle($groot[0],0,'orange'));
$gr->stamps(open_circle($groot[1],0,'orange'));


$grDom = init_graph(-8,-1,8,1, size=>[544,68]);
# x-axis
$grDom -> moveTo(-8,0);
$grDom -> lineTo(8,0,'black','1');
# x-axis ticks entered as y-value, color, list of x-values
$grDom -> h_ticks(0,"black",-7,-6,-5,-4,-3,-2,-1,0,1,2,3,4,5,6,7);
# labels for tickmarks
foreach my $j (0..7) {
    $grDom->lb( new Label($j,-0.5,$j, 'black','center','middle'));
    $grDom->lb( new Label(-$j,-0.5,-$j, 'black','center','middle'));
}
$firstMid = (-8+$minReal)/2;
$secondMid = ($minReal+$maxReal)/2;
$thirdMid = ($maxReal+8)/2;
$grDom->lb( new Label($firstMid,0.5,"(-inf,$min)",'black','center','middle'));
$grDom->lb( new Label($secondMid,0.5,"($min,$max)",'black','center','middle'));
$grDom->lb( new Label($thirdMid,0.5,"($max,inf)",'black','center','middle'));
$grDom->stamps(closed_circle($groot[0],0,'orange'));
$grDom->stamps(open_circle($groot[1],0,'orange'));

$grDom2 = init_graph(-8,-2,8,1, size=>[544,102]);
# x-axis
$grDom2 -> moveTo(-8,0);
$grDom2 -> lineTo(8,0,'black','1');
# x-axis ticks entered as y-value, color, list of x-values
$grDom2 -> h_ticks(0,"black",-7,-6,-5,-4,-3,-2,-1,0,1,2,3,4,5,6,7);
# labels for tickmarks
foreach my $j (0..7) {
    $grDom2->lb( new Label($j,-0.5,$j, 'black','center','middle'));
    $grDom2->lb( new Label(-$j,-0.5,-$j, 'black','center','middle'));
}
$grDom2->moveTo(-8,0);
$grDom2->lineTo($minReal,0,'red',3);
$grDom2->lineTo($maxReal,0,'green',3);
$grDom2->lineTo(8,0,'red',3);
$grDom2->lb( new Label($firstMid,0.5,"negative",'red','center','middle'));
$grDom2->lb( new Label($secondMid,0.5,"positive",'green','center','middle'));
$grDom2->lb( new Label($thirdMid,0.5,"negative",'red','center','middle'));
$grDom2->moveTo($groot[0],0.55);
$grDom2->lineTo($groot[0],-1,'orange',1);
$grDom2->stamps(closed_circle($groot[0],0,'orange'));
$grDom2->lb( new Label($groot[0],0.75,"root",'orange','center','middle'));
$grDom2->lb( new Label($groot[0],-1.5,"x=$root[0]",'orange','center','middle'));
$grDom2->moveTo($groot[1],0.55);
$grDom2->lineTo($groot[1],-1,'orange',1,'dashed');
$grDom2->stamps(open_circle($groot[1],0,'orange'));
$grDom2->lb( new Label($groot[1],0.75,"asymptote",'orange','center','middle'));
$grDom2->lb( new Label($groot[1],-1.5,"x=$root[1]",'orange','center','middle'));

$x1 = Real(random(-8,floor($minReal),1));
$y1 = Fraction($n->eval(x=>$x1),$d->eval(x=>$x1));
$pn1 = PopUp(["?","Positive","Negative"],"Negative");
$xy1 = MultiAnswer($x1,$y1)->with(
    singleResult=>1, 
    checkTypes=>0,
    format=>'f(%1$s)=%2$s',
    tex_format=>'f(%1$s)=%2$s',
    checker=> sub {
        my ($correct, $student, $self) = @_;
        my ($x,$y) = @{$student};
        Value::Error("You should select a number for your x-value.") unless ($x->type eq "Number");
        Value::Error("f($x) should be a number.") unless ($y->type eq "Number");
        $self->setMessage(1,"You need to select an x-value less than $min for this interval.") unless ($x < $min);
        $self->setMessage(2,"You have incorrectly evaluated f($x).") unless( $y == $rational->eval(x=>$x));
        return ($x < $min, $x< $min && $y == $rational->eval(x=>$x));
    });

$x2 = Real(0);
$y2 = Fraction($n->eval(x=>$x2),$d->eval(x=>$x2));
$pn2 = PopUp(["?","Positive","Negative"],"Positive");
$xy2 = MultiAnswer($x2,$y2)->with(
    singleResult=>1, 
    checkTypes=>0,
    format=>'f(%1$s)=%2$s',
    tex_format=>'f(%1$s)=%2$s',
    checker=> sub {
        my ($correct, $student, $self) = @_;
        my ($x,$y) = @{$student};
        Value::Error("You should select a number for your x-value.") unless ($x->type eq "Number");
        Value::Error("f($x) should be a number.") unless ($y->type eq "Number");
        $self->setMessage(1,"You need to select an x-value between $min and $max for this interval.") unless ($x > $min && $x < $max);
        $self->setMessage(2,"You have incorrectly evaluated f($x).") unless( $y == $rational->eval(x=>$x));
        return ($x > $min && $x < $max, $x > $min && $x < $max && $y == $rational->eval(x=>$x));
    });

$x3 = Real(random(ceil($maxReal),8,1));
$y3 = Fraction($n->eval(x=>$x3),$d->eval(x=>$x3));
$pn3 = PopUp(["?","Positive","Negative"],"Negative");
$xy3 = MultiAnswer($x3,$y3)->with(
    singleResult=>1, 
    checkTypes=>0,
    format=>'f(%1$s)=%2$s',
    tex_format=>'f(%1$s)=%2$s',
    checker=> sub {
        my ($correct, $student, $self) = @_;
        my ($x,$y) = @{$student};
        Value::Error("You should select a number for your x-value.") unless ($x->type eq "Number");
        Value::Error("f($x) should be a number.") unless ($y->type eq "Number");
        $self->setMessage(1,"You need to select an x-value greater than $max for this interval.") unless ($x > $max);
        $self->setMessage(2,"You have incorrectly evaluated f($x).") unless( $y == $rational->eval(x=>$x));
        return ($x > $max, $x > $max && $y == $rational->eval(x=>$x));
    });

Context("Interval");

$nonZero = Compute("(-inf,$min) U ($min,$max) U ($max,inf)");

# set up all four inequalities to solve, align with answer
@texIneq = ('>','<','\geq','\leq');
$ansInt[0] = Compute("($min,$max)");
$ansInt[1] = Compute("(-inf,$min) U ($max,inf)");
$ansInt[2] = ($min==$root[1])?Compute("($min,$max]"):Compute("[$min,$max)");
$ansInt[3] = ($min==$root[1])?Compute("(-inf,$min)U[$max,inf)"):Compute("(-inf,$min]U($max,inf)");

$randomIneq = random(0,3,1);

##############################################################
#
#  Text
#
#

BEGIN_PGML

Solve for [`x`]:

>> [`\displaystyle [$rational] [$texIneq[$randomIneq]] 0 `] <<

END_PGML

Scaffold::Begin();
Section::Begin("Identify roots and asymptotes");
BEGIN_PGML

[`\displaystyle [$rational]`] has a root at [`x=`][_____]{$root[0]}

[`\displaystyle [$rational]`] has a vertical asymptote at [`x=`][_____]{$root[1]}

Use interval notation to represent the three separate intervals where [`` [$rational] ``] will be *non-zero*: [_________________________]{$nonZero}

* Type "inf" and "-inf" for [`\infty`] and [`-\infty`], respectively.
* Type the capital letter "U" to represent the union symbol [`\cup`].


END_PGML
Section::End();

Section::Begin("Test each interval");
BEGIN_PGML

>> [@ image( insertGraph($grDom), width=>544, height=>68, tex_size=>800 ) @]* <<

a. Select an [`x`]-value in the first interval, [`(-\infty,[$min])`]: [_____]{$xy1}  

    Evaluate the function [`f(x)`] at your selected [`x`]-value: [_____]{$xy1}  
    
    Will [`f(x)`] be positive or negative on the first interval? [_]{$pn1}
    
b. Select an [`x`]-value in the second interval, [`([$min],[$max])`]: [_____]{$xy2}  

    And evaluate the function [`f(x)`] at your selected [`x`]-value: [_____]{$xy2}  
    
    Will [`f(x)`] be positive or negative on the second interval? [_]{$pn2}
    
c. Select an [`x`]-value in the third interval, [`([$max],\infty)`]: [_____]{$xy3}  

    And evaluate the function [`f(x)`] at your selected [`x`]-value: [_____]{$xy3}  
    
    Will [`f(x)`] be positive or negative on the third interval? [_]{$pn3}

END_PGML
Section::End();

Section::Begin("Solve inequalities by interpreting your results");

BEGIN_PGML


a. On what interval(s) will [` \displaystyle [$rational] [$texIneq[0]] 0 `]? [_______________]{$ansInt[0]}

a. On what interval(s) will [` \displaystyle [$rational] [$texIneq[1]] 0 `]? [_______________]{$ansInt[1]}

a. On what interval(s) will [` \displaystyle [$rational] [$texIneq[2]] 0 `]? [_______________]{$ansInt[2]}

a. On what interval(s) will [` \displaystyle [$rational] [$texIneq[3]] 0 `]? [_______________]{$ansInt[3]}

END_PGML
Section::End();
Scaffold::End();
BEGIN_PGML

You may receive full credit for this problem by giving the correct answer here.  
(The steps above are provided to assist you with the process of finding the correct final answer.)

[`\displaystyle [$rational] [$texIneq[$randomIneq]] 0 `] on the interval(s): [_______________]{$ansInt[$randomIneq]}

* Type "inf" and "-inf" for [`\infty`] and [`-\infty`], respectively.
* Type the capital letter "U" to represent the union symbol [`\cup`].



END_PGML
##############################################################
#
#  Answers
#
#

BEGIN_PGML_SOLUTION

The graph of [`y=f(x)`] looks like:

>> [@ image( insertGraph($gr), width=>400, height=>400, tex_size=>800 ) @]* <<

>> [@ image( insertGraph($grDom2), width=>544, height=>68, tex_size=>800 ) @]* <<


END_PGML_SOLUTION

ENDDOCUMENT();        
