# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Read a word problem. Write an equation and solve for unknown.
# b is the product of c and d less than a number. What is the number? 
#
# Last updated: Kling, 7/21/13; Wherry, 7/03/13, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','english-to-math','add','multiply','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "PCCmacros.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

$a=random(2,9,1);
$b=random(1,10,1);
$ans2=random(2,15,1);
$c=$a*($ans2-$b);

Context("ImplicitPlane");
$var=RandomVariableName();
Context()->variables->are($var=>"Real");
Context()->flags->set(showExtraParens=>0);
Context()->flags->set(reduceConstants=>0);

$ans1string = "$c = $a*($var-$b)";
$ans1 = Compute($ans1string);


##############################################

TEXT(beginproblem());

BEGIN_PGML

Write an equation for the following situation, and then solve for the unknown. Please use [`[$var]`] as the unknown variable.

    [@numberWord($c,capital=>1)@] is the product of [@numberWord($a)@] and [@numberWord($b)@] less than a number. What is the number?

Your equation is: 
    [_______________________]

The unknown number is: 
    [___________]

END_PGML

##############################################

ANS($ans1->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans1string)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans1string)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  Context()->variables->are($var=>"Real");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

Context("LimitedNumeric");
Context()->variables->are($var=>"Real");
ANS(Compute($ans2)->cmp);

$s1=$b*$a;
$s2=$c+$s1;

BEGIN_PGML_SOLUTION

Let the unknown number be [`[$var]`]. 

Following "the product of", there are two expressions being multiplied: "[`[$a]`]" and "[`[$b]`] less than a number."

We need to use parentheses so that subtraction happens first.

The equation is:

    [`` [$c]=[$a]([$var]-[$b]) ``]

Next, we solve for [`[$var]`]:

    [`
\begin{aligned}
  [$c] &\;=\; [$a]([$var]-[$b]) \\
  [$c] &\;=\; [$a][$var]-[$a] \cdot [$b]   \\
  [$c] &\;=\; [$a][$var]-[$s1]    \\
  [$c]\mathbf{{} +[$s1]} &\;=\; [$a][$var]-[$s1] \mathbf{{} +[$s1]} \\
  [$s2] &\;=\; [$a][$var]          \\
  \frac{[$s2]}{[$a]} &\;=\; \frac{[$a][$var]}{[$a]} \\
  [$ans2] &\;=\; [$var]
\end{aligned}
`]

The unknown number is [`[$ans2]`].

END_PGML_SOLUTION

ENDDOCUMENT();

