# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Read a word problem. Write an equation and solve for unknown.
# The product of b and a number is c. What is the number? 
#
# Last updated: Wherry, 7/03/13, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','english-to-math','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "PCCmacros.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

$a=random(2,9,1);
$ans2=random(5,15,1);
$b=$a*$ans2;

Context("ImplicitPlane");
Context()->flags->set(showExtraParens=>0);
Context()->variables->are(x=>'Real');
Context()->flags->set(reduceConstants=>0);
$ans1 = Compute("$a*x= $b");


##############################################

TEXT(beginproblem());

BEGIN_PGML

Write an equation for the following situation, and then solve for the unknown. Please use [`x`] as the unknown variable.

    The product of [@numberWord($a)@] and a number is [@numberWord($b)@]. What is the number?

Your equation is: 
    [_______________________]

The unknown number is: 
    [___________]

END_PGML

##############################################

ANS($ans1->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans1->{correct_ans})->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans1->{correct_ans})->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

Context("LimitedNumeric");
$ans2 = $b/$a;
ANS(Compute($ans2)->cmp);

BEGIN_PGML_SOLUTION

Let the unknown number be [`x`]. The equation is:

    [`` [$a]x=[$b] ``]

Next, we solve for [`x`]:

    [`
\begin{aligned}
  [$a]x   &= [$b] \\
   \frac{[$a]x}{[$a]} &= \frac{[$b]}{[$a]} \\
  x        &= [$ans2] 
\end{aligned}
`]

The unknown number is [`[$ans2]`].

END_PGML_SOLUTION

ENDDOCUMENT();


