# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Rewrite an English sentence as a mathematical equation
# The sum of (twice/n times) a number and $a is $b
#
#
# Last edited: Jordan 6/26/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('equation','english-to-math')
## DBCCSS('6.EE.2.a')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserImplicitPlane.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");

$n = random(2,8,1);

$a = random(2,30,1);
$sol = random(1,50,1);
$b = $n*$sol + $a;


Context("ImplicitPlane");
Context()->flags->set(showExtraParens=>0);
Context()->flags->set(reduceConstants=>0);
$ans = Compute("$n*x+$a=$b");

if ($n == 2) {$nString = 'twice';} else {$nString = numberWord($n).' times';};

##############################################

TEXT(beginproblem());

BEGIN_PGML

Translate the following sentence into a math equation. Use [`x`] to represent the unknown number.

    The sum of [$nString] a number and [@numberWord($a)@] is [@numberWord($b)@].

    [_______________]

END_PGML

##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));



BEGIN_PGML_SOLUTION

The word 'is' implies equality. So we have 

    The sum of [$nString] a number and [@numberWord($a)@] [`=`] [@numberWord($b)@].

The word 'sum' implies addition. The two objects being added are '[$nString] a number' and '[@numberWord($a)@]'. So we have:

    The sum of [`[$n]x`] and [`[$a]`] [`=`] [`[$b]`].

Lastly we write the additions symbolically:

    [`[@Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX@]`]


END_PGML_SOLUTION

ENDDOCUMENT();


