# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Read a word problem. Write an equation and solve for unknown.
# b less than the quotient of a number and c is the same as 
# the quotient of the number and d. 
# What is the number? 
#
# Last updated: Wherry, 7/03/13, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','english-to-math','multiply','subtract','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "PCCmacros.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

$a=random(2,5,1);
$c=random(6,10,1);
$com = lcm($a,$c);
$ans2=$com*list_random((1,2,3));
$b = $ans2/$a-$ans2/$c;

Context("ImplicitPlane");
$var=RandomVariableName();
Context()->variables->are($var=>"Real");
Context()->flags->set(showExtraParens=>0);
Context()->flags->set(reduceConstants=>0);

$ans1string = "$var/$a-$b=$var/$c";
$ans1 = Compute($ans1string);


##############################################

TEXT(beginproblem());

BEGIN_PGML

Write an equation for the following situation, and then solve for the unknown. Please use [`[$var]`] as the unknown variable.

    [@numberWord($b,capital=>1)@] less than the quotient of a number and [@numberWord($a)@] is the same as the quotient of the number and [@numberWord($c)@]. What is the number?

Your equation is: 
    [_______________________]

The unknown number is: 
    [___________]

END_PGML

##############################################

ANS($ans1->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans1string)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans1string)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  Context()->variables->are($var=>"Real");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

Context("LimitedNumeric");
Context()->variables->are($var=>"Real");
ANS(Compute($ans2)->cmp);

$s1 = $com/$a;
$s2 = $com/$c;
$s3 = $com*$b;
$s4 = $s2-$s1;

if ($s2 == 1) {$s2 = ''};
if ($s4 == -1) {$s4b = '-'};

BEGIN_PGML_SOLUTION

Let the unknown number be [`[$var]`]. The equation is:

    [`` \frac{[$var]}{[$a]} -[$b] = \frac{[$var]}{[$c]} ``]

Next, we solve for [`[$var]`]. To clear denominators from both fractions, we multiply each term in the equation with a common denominator, which could be [`[$com]`].

    [`
\begin{aligned}
  \frac{[$var]}{[$a]} -[$b]   &=  \frac{[$var]}{[$c]} \\
  \frac{[$var]}{[$a]} +(-[$b])   &=  \frac{[$var]}{[$c]} \\
  [$com] \cdot \frac{[$var]}{[$a]} +[$com] \cdot (-[$b])   &=  [$com] \cdot \frac{[$var]}{[$c]} \\
  [$s1][$var] - [$s3] &= [$s2][$var] \\
  [$s1][$var] - [$s3]\mathbf{{}-[$s1][$var]} &= [$s2][$var]\mathbf{{}-[$s1][$var]} \\
         -[$s3] &= [$s4b][$var] \\
 \frac{-[$s3]}{[$s4]} &= \frac{[$s4b][$var]}{[$s4]} \\
  [$ans2] &= [$var]\\
  [$var] &= [$ans2]
\end{aligned}
`]

The unknown number is [`[$ans2]`].

END_PGML_SOLUTION

ENDDOCUMENT();


