# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Read a word problem. Write an equation and solve for unknown.
# The difference between a number and b is c. What is the number? 
#
# Last updated: Kling, 7/21/13; Wherry, 7/03/13, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','english-to-math','add')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "PCCmacros.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

$a=random(11,19,1);
$ans2=random(30,50,1);
$b=$ans2-$a;

Context("ImplicitPlane");
Context()->flags->set(showExtraParens=>0);
Context()->variables->are(x=>'Real');
Context()->flags->set(reduceConstants=>0);

$ans1string = "x - $a = $b";
$ans1 = Compute($ans1string);


##############################################

TEXT(beginproblem());

BEGIN_PGML

Write an equation for the following situation, and then solve for the unknown. Please use [`x`] as the unknown variable.

    The difference between a number and [@numberWord($a)@] is [@numberWord($b)@]. What is the number?

Your equation is: 
    [_______________________]

The unknown number is: 
    [___________]

END_PGML

##############################################

ANS($ans1->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans1string)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans1string)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

Context("LimitedNumeric");
$ans2 = $b+$a;
ANS(Compute($ans2)->cmp);

BEGIN_PGML_SOLUTION

Let the unknown number be [`x`]. The equation is:

    [`` x-[$a]=[$b] ``]

Next, we solve for [`x`]:

[`
\begin{aligned}
  x-[$a]   &= [$b] \\
   x-[$a]\mathbf{{}+[$a]}   &= [$b]\mathbf{{}+[$a]} \\
  x        &= [$ans2] 
\end{aligned}
`]

The unknown number is [`[$ans2]`].

END_PGML_SOLUTION

ENDDOCUMENT();

