# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Rewrite an English sentence as a mathematical equation
# The difference between/sum of $a and a number is $b
#
# Last edited: Jordan 6/26/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('equation','english-to-math')
## DBCCSS('6.EE.2.a')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserImplicitPlane.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->flags->set(showExtraParens=>0);

$index = random(0,1,1);
@operations = ('sum of','difference between');
$operation = $operations[$index];


$sol = random(1,20,1);
$a = $sol+random(1,40,1);
$b = $a+(-1)**$index*$sol;

$ans = Compute("$a+x=$b") if ($index == 0);
$ans = Compute("$a-x=$b") if ($index == 1);


##############################################

TEXT(beginproblem());

BEGIN_PGML

Translate the following sentence into a math equation. Use [`x`] to represent the unknown number.

    The [$operation] [@numberWord($a)@] and a number is [@numberWord($b)@].

    [____________]

END_PGML

##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

@solnwords = ('addition','subtraction');
$solnword = $solnwords[$index];

BEGIN_PGML_SOLUTION

The word 'is' implies equality. So we have 

    The [$operation] [@numberWord($a)@] and a number [`=`] [@numberWord($b)@].

The phrase '[$operation]' implies [$solnword]. So the equation we want is

    [`[@Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX@]`]

END_PGML_SOLUTION

ENDDOCUMENT();


