## DESCRIPTION
## Precalculus
## ENDDESCRIPTION

## Tagged by cmd6a 3/12/06

## DBsubject(Algebra)
## DBchapter(Exponential and logarithmic expressions and functions)
## DBsection(Applications and models - general)
## Date(6/3/2002)
## Institution(ASU)
## MLT(GivenExponential)
## Level(5)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('6.2')
## Problem1('17')
## KEYWORDS('algebra', 'natural exponential function')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "Parser.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$p=random(900,1000,5);
$a=random(10,50,5);
$r=$a/100;

BEGIN_TEXT
The number of bacteria in a culture is given by the function
\[
n(t)=$p e^{$r t}
\]
where \(t\) is measured in hours.
$BR
$BR
(a) What is the relative rate of growth of this bacterium population?
$BR
Your answer is \{ans_rule(10) \} percent
$BR
$BR 
(b) What is the initial population of the culture? 
$BR
Your answer is \{ans_rule(15) \}
$BR
$BR
(c) How many bacteria will the culture contain at time \(t=5\)? 
$BR
Your answer is (round off to the nearest integer) \{ans_rule(15) \}
END_TEXT

$ans1 = $a;
$ans2 = $p;
$ans3 = $p*exp($r*5);
$ans3 = round($ans3);

ANS(num_cmp($ans1));
ANS(num_cmp($ans2));
Context("Numeric")->flags->set(
tolerance => 0.00001,
            tolType => 'absolute',
          );

$a3= Real($ans3);
ANS($a3->cmp(checker=> sub {
 my $correct = shift; my $student = shift;
 if (abs($student->value - round($student->value))>0.0000001) {
   Value::Error("Your answer should be an integer") ;
   return 0;
 }
 return 1 if $correct == $student;
 return 0;
}));

#ANS(num_cmp($ans3, mode=>'strict', tolType=> 'absolute', tol=>0.00001));
#ANS(Real($ans3)->cmp);

ENDDOCUMENT();        # This should be the last executable line in the problem.
