# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Build a table of ordered pairs for a function x=c
# 
# Last updated: Jordan, 8/7/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Coordinate Geometry')
## KEYWORDS('coordinates','linear equation')
## DBCCSS('8.F.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "pccTables.pl",
  "parserMultiAnswer.pl",
,
  "PGcourse.pl",
);

##############################################
TEXT(beginproblem());
Context("Numeric");
Context()->variables->add(y=>'Real');
Context()->flags->set(showExtraParens=>0, reduceConstants=>0);

$c = list_random(-10..-2,2..10);

$left = Compute("x");
$right = Formula("$c");

@defaultAnswers = ();
for my $i (0..4) {
  $x[$i] = $c;
  $y[$i] = $i-1;
  push(@defaultAnswers,$x[$i]);
  push(@defaultAnswers,$y[$i]);
}

$multians = MultiAnswer(@defaultAnswers)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @sxy = @{$student};
      my @xs, @ys;
     for my $k (0..4) 
       {$xs[$k] = $sxy[2*$k];
        $ys[$k] = $sxy[2*$k+1];}
     #
     #  Loop though the student answers
     ##
     for ($i = 0; $i <= 4; $i++) {
       my $ith = Value::List->NameForNumber($i+1);
       #
       #  Check that the student's answers are Reals
       #
       Value->Error("Your $ith x-value is not a number") if ($xs[$i]->type ne "Number");
       Value->Error("Your $ith y-value is not a number") if ($ys[$i]->type ne "Number");
       #
       #  Check that the point hasn't been given before
       #
       for ($j = 0, $used = 0; $j < $i; $j++) {
         my $jth = Value::List->NameForNumber($j+1);
         if ($xs[$j] == $xs[$i] and $ys[$j] == $ys[$i]) {
           Value->Error("Your $ith row in the table is the same as your $jth row") unless $ansHash->{isPreview};
           $used = 1; last;
         }
       }
       #
       #  If not already used, check that it satisfies the equation
       #    and increase the score if so.
       #
       if (!$used) {
         Value->Error("Your $ith row is not correct") unless $ansHash->{isPreview} or $left->eval(x=>$xs[$i],y=>$ys[$i]) == $right->eval(x=>$xs[$i],y=>$ys[$i]);
         }
       }
     return 1;
});



##############################################

BEGIN_PGML

Make a table of solutions for the equation [`[$left]=[$right]`].

END_PGML
Context()->texStrings;
BEGIN_TEXT

\{ 
DataTable([[{data=>"\(x\)",header=>"CH"},{data=>"\(y\)",header=>"CH"},],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
[$multians->ans_rule(8), $multians->ans_rule(8)],
], allcellcss=>"border:1px solid;padding:4pt;text-align:center;", caption => "Values of \(x\) and \(y\) satisfying the equation \($left=$right\)", captioncss =>"padding:4pt; text-decoration:none;");
\}  


$PAR

END_TEXT

Context()->normalStrings;

##############################################

ANS( $multians->cmp() ); 

##############################################

sub format_list {
   return "" if !@_;
   my $last = pop(@_);
   return '\('.$last.'\)' if !@_;
   return '\('.join('\), \(', @_) .'\) and \('. $last.'\)';
}

$yvals = format_list(@y);

Context()->texStrings;
BEGIN_SOLUTION

Since this equation does not involve \(y\) and demands that \(x\) be equal to \($c\), any \(y\)-value can be paired up with \($c\) for \(x\) to make a valid solution. 

This solution will use the \(y\)-values $yvals. The choice to use these \(y\)-values is arbitrary.


So the table may be completed as:

\{ 
DataTable([[{data=>"\(x\)",header=>"CH"},{data=>"\(y\)",header=>"CH"},],
 ["\($x[0]\)","\($y[0]\)"],
["\($x[1]\)","\($y[1]\)"],
["\($x[2]\)","\($y[2]\)"],
["\($x[3]\)","\($y[3]\)"],
["\($x[4]\)","\($y[4]\)"],
], allcellcss=>"border:1px solid;padding:4pt;text-align:center; width:60pt;", caption => "Values of \(x\) and \(y\) satisfying the equation \( $left=$right \)", captioncss =>"padding:4pt; text-decoration:none;");
\}  




END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();


