##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('PDE', 'laplacian', 'polar')

## DBsubject('Calculus')
## DBchapter('Partial Differential Equations')
## DBsection('Other Coordinate Systems')
## Date('9/12/2011')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('13.1')
## Problem1('')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "contextFraction.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   "PGcourse.pl",      # Customization file for the course
   "Alfredmacros.pl",
   "answerHints.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;
#force the use of jsmath mode for the inline entry answer boxes
jsmathmode();
##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->variables->add(y=>'Real',r=>'Real',theta=>'Real',rx=>'Real',thetax=>'Real',thetay=>'Real',ur=>'Real',utheta=>'Real',urr=>'Real',urtheta=>'Real',uthetar=>'Real',uthetatheta=>'Real' );

 
##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
For this problem you will need to input partial derivatives. We will use the notation dependent variable followed by the independent variables that are being differentiated with respect to. For example
$BR
\(\frac{\partial^2 u}{\partial r^2} = urr\)
$BR
in WeBWorK notation. \(\theta\) is typed as theta.
$BR
$BR
$BR
In this problem we will derive the Laplacian in polar coordinates using the identities
$BR
$BCENTER
\(\text{(i)}\ x = r\cos\left(\theta\right),\ \text{(ii)}\ y = r\sin\left(\theta\right),\ \text{(iii)}\  x^2+y^2=r^2,\ \text{(iv)}\  \tan\left(\theta\right) = \frac{y}{x}\)
$ECENTER
$BR
$BR
$BBOLD Part 1 - Find \(\displaystyle\frac{\partial r}{\partial x},\frac{\partial r}{\partial y}, \frac{\partial \theta}{\partial x}, \frac{\partial \theta}{\partial y}\) $EBOLD
$BR
$BR
First we need to find \(\displaystyle\frac{\partial r}{\partial x},\frac{\partial r}{\partial y}, \frac{\partial \theta}{\partial x}, \frac{\partial \theta}{\partial y}.\) To do this we will implicitly differentiate the above equations. 
$BR
Implicitly differentiating both sides of (iii) with respect to x, without simplifying, gives us
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Solving for \(\frac{\partial r}{\partial x}\) we get 
$BR
$BCENTER
\(\displaystyle\frac{\partial r}{\partial x} = \) \{ans_rule\} 
$ECENTER
$BR
Similarly, implicitly differentiating both sides of (iii) with respect to y, without simplifying, gives us
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Solving for \(\frac{\partial r}{\partial y}\) we get 
$BR
$BCENTER
\(\displaystyle\frac{\partial r}{\partial y} = \) \{ans_rule\} 
$ECENTER
$BR
$BR
Now we need to find \(\frac{\partial \theta}{\partial x}, \frac{\partial \theta}{\partial y},\) which take a bit more work. Implicitly differentiating both sides of (iv) with respect to x gives
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Substitute (i) and (ii) into the right hand side of this equation for x and y respectively, and solve for \(\frac{\partial \theta}{\partial x}\)
$BR
$BCENTER
\(\displaystyle\frac{\partial \theta}{\partial x} = \) \{ans_rule\}
$ECENTER
$BR
Implicitly differentiating both sides of (iv) with respect to y gives
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Substitute (i) and (ii) into the right hand side of this equation for x and y respectively, and solve for \(\frac{\partial \theta}{\partial y}\)
$BR
$BCENTER
\(\displaystyle\frac{\partial \theta}{\partial y} = \) \{ans_rule\}
$ECENTER
$BR
$BR
$BBOLD Part 2 - Find \(\displaystyle\frac{\partial u}{\partial x},\frac{\partial u}{\partial y}\) $EBOLD
$BR
$BR
By the chain rule we know that
$BR
\(\displaystyle\frac{\partial u}{\partial x} = \frac{\partial u}{\partial r}\frac{\partial r}{\partial x} + \frac{\partial u}{\partial \theta}\frac{\partial \theta}{\partial x} \) = \{ans_rule(30)\}
$BR
and
$BR
\(\displaystyle\frac{\partial u}{\partial y} = \frac{\partial u}{\partial r}\frac{\partial r}{\partial y} + \frac{\partial u}{\partial \theta}\frac{\partial \theta}{\partial y} \) = \{ans_rule(30)\}
$BR
$BR
$BBOLD Part 2 - Find \(\displaystyle\frac{\partial^2 u}{\partial x^2},\frac{\partial^2 u}{\partial y^2}\) $EBOLD
$BR
$BR
\(
\begin{aligned}
\displaystyle\frac{\partial^2 u}{\partial x^2} = \frac{\partial }{\partial x}\left(\frac{\partial u}{\partial x}\right) =&  \left(\cos(\theta)\frac{\partial }{\partial r} - \frac{\sin(\theta)}{r}\frac{\partial }{\partial \theta}\right)\left(\cos(\theta)\frac{\partial u}{\partial r} - \frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right)\\
=&  \cos(\theta)\frac{\partial u}{\partial r}\left(\cos(\theta)\frac{\partial }{\partial r}\right) 
+ \cos(\theta)\frac{\partial }{\partial r}\left( \frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right)
-\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\left(\cos(\theta)\frac{\partial u}{\partial r}\right)+
\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\left(\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right)\\
\end{aligned}
\)
$BR
Remembering to use the product rule
$BR
\(\displaystyle\cos(\theta)\frac{\partial u}{\partial r}\left(\cos(\theta)\frac{\partial }{\partial r}\right) = \) \{ans_rule(80)\}
$BR
\(\displaystyle\cos(\theta)\frac{\partial }{\partial r}\left( \frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right) = \) \{ans_rule(80)\}
$BR
\(\displaystyle\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\left(\cos(\theta)\frac{\partial u}{\partial r}\right) = \) \{ans_rule(80)\}
$BR
\(\displaystyle\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\left(\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right) = \) \{ans_rule(80)\}
$BR
$BR
$BBOLD Part 3 - Putting it all together $EBOLD
$BR
$BR
Using the results from Part 2 we see that
$BR
\(
\displaystyle\frac{\partial^2 u}{\partial x^2} = \cos^2(\theta)\frac{\partial^2 u}{\partial r^2} - \frac{2\sin(\theta)\cos(\theta)}{r}\frac{\partial^2 u}{\partial r \partial\theta}+\frac{\sin^2(\theta)}{r^2}\frac{\partial^2 u}{\partial \theta^2}+\frac{\sin^2(\theta)}{r}\frac{\partial u}{\partial r}+\frac{2\sin(\theta)\cos(\theta)}{r^2}\frac{\partial u}{\partial\theta}
\)
$BR
Similarly
$BR
\(
\displaystyle\frac{\partial^2 u}{\partial y^2} = \sin^2(\theta)\frac{\partial^2 u}{\partial r^2} + \frac{2\sin(\theta)\cos(\theta)}{r}\frac{\partial^2 u}{\partial r \partial\theta}+\frac{\cos^2(\theta)}{r^2}\frac{\partial^2 u}{\partial \theta^2}+\frac{\cos^2(\theta)}{r}\frac{\partial u}{\partial r}-\frac{2\sin(\theta)\cos(\theta)}{r^2}\frac{\partial u}{\partial\theta}
\)
$BR
Therefore, making use of some trig identities,
$BR
\(\displaystyle\nabla^2 u = \frac{\partial^2 u}{\partial x^2}+\frac{\partial^2 u}{\partial y^2}\) = \{ans_rule(70)\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

#dr/dx
ANS(Formula("2*r*rx")->cmp());
ANS(Formula("2*x")->cmp());
ANS(Formula("cos(theta)")->cmp());

#dr/dy
ANS(Formula("2*r*ry")->cmp());
ANS(Formula("2*y")->cmp());
ANS(Formula("sin(theta)")->cmp());

#dtheta/dx
ANS(Formula("sec^2(theta)*thetax")->cmp());
ANS(Formula("-y/x^2")->cmp());
ANS(Formula("-sin(theta)/r")->cmp());

#dtheta/dy
ANS(Formula("sec^2(theta)*thetay")->cmp());
ANS(Formula("-1/x")->cmp());
ANS(Formula("-sin(theta)/r")->cmp());

#du/dx
ANS(Formula("cos(theta)ur-sin(theta)/r*utheta")->cmp());

#du/dy
ANS(Formula("sin(theta)ur+cos(theta)/r*utheta")->cmp());

#second partials
ANS(Formula("cos^2(theta)*urr")->cmp());
ANS(Formula("-sin(theta)*cos(theta)/r^2*utheta+sin(theta)cos(theta)/r*urtheta")->cmp(
checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return (($correct == $student)||(Formula("-sin(theta)*cos(theta)/r^2*utheta+sin(theta)cos(theta)/r*uthetar")==$student));
}));
ANS(Formula("sin^2(theta)/r*ur-sin(theta)*cos(theta)/r*uthetar")->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return (($correct == $student)||(Formula("sin^2(theta)/r*ur-sin(theta)*cos(theta)/r*urtheta")==$student));
}));
ANS(Formula("sin(theta)*cos(theta)/r^2*utheta+sin^2(theta)/r*uthetatheta")->cmp());
ANS(Formula("urr+1/r*ur+1/r^2*uthetatheta")->cmp());
ENDDOCUMENT(); 