##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Partial differential equations)
## DBsection(Laplace's equation)
## Date(9/12/2011)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(5)
## TitleText1('Differential Equations')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('13.1')
## Problem1('')
## KEYWORDS('PDE', 'laplacian', 'polar')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->variables->are(r=>'Real',theta=>'Real',rx=>'Real',thetax=>'Real',thetay=>'Real',ur=>'Real',utheta=>'Real',urr=>'Real',urtheta=>'Real',uthetar=>'Real',uthetatheta=>'Real' );
Context()->{error}{msg}{"Variable 'x' is not defined in this context"} 
  = "You have to substitute x = rcos(theta) in for x";
Context()->{error}{msg}{"Variable 'y' is not defined in this context"} 
  = "You have to substitute y = rsin(theta) in for y";
Context()->variables->set(theta => {TeX => '\theta'});
Context()->variables->set(rx => {TeX => '\frac{\partial r}{\partial x}'});
Context()->variables->set(thetax => {TeX => '\frac{\partial \theta}{\partial x}'});
Context()->variables->set(thetay => {TeX => '\frac{\partial \theta}{\partial y}'});
Context()->variables->set(ur => {TeX => '\frac{\partial u}{\partial r}'});
Context()->variables->set(urr => {TeX => '\frac{\partial^2 u}{\partial r^2}'});
Context()->variables->set(utheta => {TeX => '\frac{\partial u}{\partial\theta}'});
Context()->variables->set(urtheta => {TeX => '\frac{\partial^2 u}{\partial r\partial\theta}'});
Context()->variables->set(uthetar => {TeX => '\frac{\partial^2 u}{\partial\theta\partial r}'});
Context()->variables->set(uthetatheta => {TeX => '\frac{\partial^2 u}{\partial\theta^2}'});
$polarcontext = Context()->copy;

Context()->variables->are(r=>'Real',theta=>'Real',rx=>'Real',thetax=>'Real',thetay=>'Real',ur=>'Real',utheta=>'Real',urr=>'Real',urtheta=>'Real',uthetar=>'Real',uthetatheta=>'Real' );
Context()->functions->undefine("sec","csc");
Context()->{error}{msg}{"Variable 'x' is not defined in this context"} 
  = "You have to substitute x = rcos(theta) in for x";
Context()->{error}{msg}{"Variable 'y' is not defined in this context"} 
  = "You have to substitute y = rsin(theta) in for y";
Context()->{error}{msg}{"Function 'sec' is not allowed in this context"} 
  = "sec(theta) = 1/cos(theta)";
Context()->variables->set(theta => {TeX => '\theta'});
Context()->variables->set(rx => {TeX => '\frac{\partial r}{\partial x}'});
Context()->variables->set(thetax => {TeX => '\frac{\partial \theta}{\partial x}'});
Context()->variables->set(thetay => {TeX => '\frac{\partial \theta}{\partial y}'});
Context()->variables->set(ur => {TeX => '\frac{\partial u}{\partial r}'});
Context()->variables->set(urr => {TeX => '\frac{\partial^2 u}{\partial r^2}'});
Context()->variables->set(utheta => {TeX => '\frac{\partial u}{\partial\theta}'});
Context()->variables->set(urtheta => {TeX => '\frac{\partial^2 u}{\partial r\partial\theta}'});
Context()->variables->set(uthetar => {TeX => '\frac{\partial^2 u}{\partial\theta\partial r}'});
Context()->variables->set(uthetatheta => {TeX => '\frac{\partial^2 u}{\partial\theta^2}'});
$polarcontextnosec = Context()->copy;

Context("Numeric");
Context()->variables->add(y=>'Real',r=>'Real',theta=>'Real',rx=>'Real',rx=>'Real',thetax=>'Real',thetay=>'Real',ur=>'Real',utheta=>'Real',urr=>'Real',urtheta=>'Real',uthetar=>'Real',uthetatheta=>'Real' );
Context()->variables->set(theta => {TeX => '\theta'});
Context()->variables->set(rx => {TeX => '\frac{\partial r}{\partial x}'});
Context()->variables->set(thetax => {TeX => '\frac{\partial \theta}{\partial x}'});
Context()->variables->set(thetay => {TeX => '\frac{\partial \theta}{\partial y}'});
Context()->variables->set(ur => {TeX => '\frac{\partial u}{\partial r}'});
Context()->variables->set(urr => {TeX => '\frac{\partial^2 u}{\partial r^2}'});
Context()->variables->set(utheta => {TeX => '\frac{\partial u}{\partial\theta}'});
Context()->variables->set(urtheta => {TeX => '\frac{\partial^2 u}{\partial r\partial\theta}'});
Context()->variables->set(uthetar => {TeX => '\frac{\partial^2 u}{\partial\theta\partial r}'});
Context()->variables->set(uthetatheta => {TeX => '\frac{\partial^2 u}{\partial\theta^2}'});
 
##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
$BR
$BBOLD Introduction $EBOLD
$BR
In this problem we will derive the Laplacian in polar coordinates using the identities
$BR
$BCENTER
\(\text{(i)}\ x = r\cos\left(\theta\right),\ \text{(ii)}\ y = r\sin\left(\theta\right),\ \text{(iii)}\  x^2+y^2=r^2,\ \text{(iv)}\  \tan\left(\theta\right) = \frac{y}{x}\)
$ECENTER
$BR
At the end of the problem you will find a complete list of the notations that you will need to enter and how to enter them.
$PAR
$BBOLD Part 1 - Find \(\displaystyle\frac{\partial r}{\partial x},\frac{\partial r}{\partial y}, \frac{\partial \theta}{\partial x}, \frac{\partial \theta}{\partial y}\) $EBOLD
$PAR
To find \(\displaystyle\frac{\partial r}{\partial x},\frac{\partial r}{\partial y}, \frac{\partial \theta}{\partial x}\), and \(\displaystyle\frac{\partial \theta}{\partial y}\) we will implicitly differentiate the above equations. Unlike calc I, x and y are both independent variables, so that \(\displaystyle\frac{dy}{dx} = 0\).
$BR
Implicitly differentiating both sides of (iii) with respect to x, without simplifying, gives us
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Solving for \(\displaystyle \frac{\partial r}{\partial x}\) we get 
$BR
$BCENTER
\(\displaystyle\frac{\partial r}{\partial x} = \) \{ans_rule\} 
$ECENTER
$BR
Similarly, implicitly differentiating both sides of (iii) with respect to y, without simplifying, gives us
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Solving for \(\displaystyle \frac{\partial r}{\partial y}\) we get 
$BR
$BCENTER
\(\displaystyle\frac{\partial r}{\partial y} = \) \{ans_rule\} 
$ECENTER
$PAR
Now we need to find \(\displaystyle \frac{\partial \theta}{\partial x}\) and \( \displaystyle \frac{\partial \theta}{\partial y},\) which take a bit more work. Implicitly differentiating both sides of (iv) with respect to x gives
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Substitute (i) and (ii) into the right hand side of this equation for x and y respectively, and solve for \(\displaystyle \frac{\partial \theta}{\partial x}\)
$BR
$BCENTER
\(\displaystyle\frac{\partial \theta}{\partial x} = \) \{ans_rule\}
$ECENTER
$BR
Implicitly differentiating both sides of (iv) with respect to y gives
$BR
$BCENTER
\{ans_rule\}=\{ans_rule\}
$ECENTER 
$BR
Substitute (i) and (ii) into the right hand side of this equation for x and y respectively, and solve for \(\displaystyle \frac{\partial \theta}{\partial y}\)
$BR
$BCENTER
\(\displaystyle\frac{\partial \theta}{\partial y} = \) \{ans_rule\}
$ECENTER
$PAR
$BBOLD Part 2 - Find \(\displaystyle\frac{\partial u}{\partial x},\frac{\partial u}{\partial y}\) $EBOLD
$PAR
By the chain rule, we know that
$PAR
\(\displaystyle\frac{\partial u}{\partial x} = \frac{\partial u}{\partial r}\frac{\partial r}{\partial x} + \frac{\partial u}{\partial \theta}\frac{\partial \theta}{\partial x} \) $SPACE and $SPACE
\(\displaystyle\frac{\partial u}{\partial y} = \frac{\partial u}{\partial r}\frac{\partial r}{\partial y} + \frac{\partial u}{\partial \theta}\frac{\partial \theta}{\partial y} \)
$PAR
Substituting the results from part 1 into these equations we get
$PAR
\(\displaystyle\frac{\partial u}{\partial x} =  \) \{ans_rule(30)\}
$PAR
\(\displaystyle\frac{\partial u}{\partial y} =  \)  \{ans_rule(30)\}
$PAR
$BBOLD Part 3 - Find \(\displaystyle\frac{\partial^2 u}{\partial x^2}\) $EBOLD
$PAR
\(\displaystyle\frac{\partial^2 u}{\partial x^2} = \frac{\partial }{\partial x}\left(\frac{\partial u}{\partial x}\right) =  \left(\cos(\theta)\frac{\partial }{\partial r} - \frac{\sin(\theta)}{r}\frac{\partial }{\partial \theta}\right)\left(\cos(\theta)\frac{\partial u}{\partial r} - \frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right)\)
$PAR
\(\displaystyle\hskip 85pt
=  \cos(\theta)\frac{\partial}{\partial r}\left(\cos(\theta)\frac{\partial u}{\partial r}\right) 
+ \cos(\theta)\frac{\partial}{\partial r}\left(-\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right)
-\frac{\sin(\theta)}{r}\frac{\partial}{\partial \theta}\left(\cos(\theta)\frac{\partial u}{\partial r}\right)+
\frac{\sin(\theta)}{r}\frac{\partial}{\partial \theta}\left(\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right)\)
$PAR
Remembering to use the product rule when it's appropriate
$PAR
\(\displaystyle\cos(\theta)\frac{\partial}{\partial r}\left(\cos(\theta)\frac{\partial u}{\partial r}\right) = \) \{ans_rule(80)\}
$PAR
\(\displaystyle\cos(\theta)\frac{\partial }{\partial r}\left( -\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right) = \) \{ans_rule(80)\}
$PAR
\(\displaystyle-\frac{\sin(\theta)}{r}\frac{\partial}{\partial \theta}\left(\cos(\theta)\frac{\partial u}{\partial r}\right) = \) \{ans_rule(80)\}
$PAR
\(\displaystyle\frac{\sin(\theta)}{r}\frac{\partial}{\partial \theta}\left(\frac{\sin(\theta)}{r}\frac{\partial u}{\partial \theta}\right) = \) \{ans_rule(80)\}
$PAR
$BBOLD Part 4 - Find \(\displaystyle\frac{\partial^2 u}{\partial y^2}\) $EBOLD
$PAR
\(\displaystyle\frac{\partial^2 u}{\partial y^2} = \frac{\partial }{\partial y}\left(\frac{\partial u}{\partial y}\right) =  \left(\sin(\theta)\frac{\partial }{\partial r} + \frac{\cos(\theta)}{r}\frac{\partial }{\partial \theta}\right)\left(\sin(\theta)\frac{\partial u}{\partial r} + \frac{\cos(\theta)}{r}\frac{\partial u}{\partial \theta}\right)\)
$PAR
\(\displaystyle\hskip 85pt
=  \sin(\theta)\frac{\partial}{\partial r}\left(\sin(\theta)\frac{\partial u}{\partial r}\right) 
+ \sin(\theta)\frac{\partial}{\partial r}\left(\frac{\cos(\theta)}{r}\frac{\partial u}{\partial \theta}\right)
+\frac{\cos(\theta)}{r}\frac{\partial}{\partial \theta}\left(\sin(\theta)\frac{\partial u}{\partial r}\right)+
\frac{\cos(\theta)}{r}\frac{\partial}{\partial \theta}\left(\frac{\cos(\theta)}{r}\frac{\partial u}{\partial \theta}\right)\)
$PAR
\(\displaystyle\sin(\theta)\frac{\partial}{\partial r}\left(\sin(\theta)\frac{\partial u}{\partial r}\right) = \) \{ans_rule(80)\}
$PAR
\(\displaystyle\sin(\theta)\frac{\partial}{\partial r}\left(\frac{\cos(\theta)}{r}\frac{\partial u}{\partial \theta}\right) = \) \{ans_rule(80)\}
$PAR
\(\displaystyle\frac{\cos(\theta)}{r}\frac{\partial}{\partial \theta}\left(\sin(\theta)\frac{\partial u}{\partial r}\right) = \) \{ans_rule(80)\}
$PAR
\(\displaystyle\frac{\cos(\theta)}{r}\frac{\partial}{\partial \theta}\left(\frac{\cos(\theta)}{r}\frac{\partial u}{\partial \theta}\right) = \) \{ans_rule(80)\}
$PAR
$BBOLD Part 5 - Putting it all together $EBOLD
$PAR
Using the results from parts 3 and 4 we see that
$PAR
\(
\displaystyle\frac{\partial^2 u}{\partial x^2} = \cos^2(\theta)\frac{\partial^2 u}{\partial r^2} - \frac{2\sin(\theta)\cos(\theta)}{r}\frac{\partial^2 u}{\partial r \partial\theta}+\frac{\sin^2(\theta)}{r^2}\frac{\partial^2 u}{\partial \theta^2}+\frac{\sin^2(\theta)}{r}\frac{\partial u}{\partial r}+\frac{2\sin(\theta)\cos(\theta)}{r^2}\frac{\partial u}{\partial\theta}
\)
$PAR
Similarly
$PAR
\(
\displaystyle\frac{\partial^2 u}{\partial y^2} = \sin^2(\theta)\frac{\partial^2 u}{\partial r^2} + \frac{2\sin(\theta)\cos(\theta)}{r}\frac{\partial^2 u}{\partial r \partial\theta}+\frac{\cos^2(\theta)}{r^2}\frac{\partial^2 u}{\partial \theta^2}+\frac{\cos^2(\theta)}{r}\frac{\partial u}{\partial r}-\frac{2\sin(\theta)\cos(\theta)}{r^2}\frac{\partial u}{\partial\theta}
\)
$BR
Therefore, making use of some trig identities,
$PAR
\(\displaystyle\nabla^2 u = \frac{\partial^2 u}{\partial x^2}+\frac{\partial^2 u}{\partial y^2}\) = \{ans_rule(70)\}
$PAR
$BBOLD Notation $EBOLD
Here is a list of the notations you will need to input your answers
$PAR
\(\displaystyle \theta = \text{theta} 
  \qquad \frac{\partial \theta}{\partial x} = \text{thetax}
  \qquad \frac{\partial \theta}{\partial y} = \text{thetay} \)
$PAR
\( \displaystyle \theta = \text{theta} 
    \qquad \frac{\partial \theta}{\partial x} = \text{thetax}
    \qquad \frac{\partial \theta}{\partial y} = \text{thetay} \)
$PAR 
\( \displaystyle \frac{\partial r}{\partial x} = \text{rx} 
   \qquad \frac{\partial r}{\partial y} = \text{ry}\)
$PAR
\( \displaystyle  \frac{\partial u}{\partial r} = \text{ur} 
    \qquad \frac{\partial u}{\partial \theta} = \text{utheta} \)
$PAR
\( \displaystyle  \frac{\partial^2 u}{\partial r^2} = \text{urr}
    \qquad \frac{\partial^2 u}{\partial \theta^2} = \text{uthetatheta} \)
$PAR
\(\displaystyle \frac{\partial^2 u}{\partial r\partial\theta} = 
    \frac{\partial^2 u}{\partial\theta\partial r}=\text{urtheta}=\text{uthetar}\)
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

#dr/dx
ANS(Formula("2*x")->cmp());
ANS(Formula("2*r*rx")->cmp());

ANS(Formula($polarcontextnosec,"cos(theta)")->cmp());

#dr/dy
ANS(Formula("2*y")->cmp());
ANS(Formula("2*r*ry")->cmp());
ANS(Formula($polarcontextnosec,"sin(theta)")->cmp());

#dtheta/dx
ANS(Formula("sec^2(theta)*thetax")->cmp());
ANS(Formula("-y/x^2")->cmp());
ANS(Formula($polarcontextnosec,"-sin(theta)/r")->cmp());

#dtheta/dy
ANS(Formula($polarcontext,"sec^2(theta)*thetay")->cmp());
ANS(Formula("1/x")->cmp());
ANS(Formula($polarcontextnosec,"cos(theta)/r")->cmp());

# PART 2
#du/dx
ANS(Formula($polarcontextnosec,"cos(theta)ur-sin(theta)/r*utheta")->cmp());

#du/dy
ANS(Formula($polarcontextnosec,"sin(theta)ur+cos(theta)/r*utheta")->cmp());

# PART 3
ANS(Formula("cos^2(theta)*urr")->cmp());

#### Check for the partial derivatives in either order: urtheta or uthetar
ANS(Formula("sin(theta)*cos(theta)/r^2*utheta-sin(theta)cos(theta)/r*urtheta")->cmp(
checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return (($correct == $student)||(Formula("sin(theta)*cos(theta)/r^2*utheta-sin(theta)cos(theta)/r*uthetar")==$student));
}));

#### Check for the partial derivatives in either order: urtheta or uthetar
ANS(Formula("sin^2(theta)/r*ur-sin(theta)*cos(theta)/r*uthetar")->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return (($correct == $student)||(Formula("sin^2(theta)/r*ur-sin(theta)*cos(theta)/r*urtheta")==$student));
}));

ANS(Formula("sin(theta)*cos(theta)/r^2*utheta+sin^2(theta)/r^2*uthetatheta")->cmp());


# PART 4
ANS(Formula("sin^2(theta)*urr")->cmp());

#### Check for the partial derivatives in either order: urtheta or uthetar
ANS(Formula("-sin(theta)cos(theta)/r^2*utheta+sin(theta)cos(theta)/r*urtheta")->cmp(
checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return (($correct == $student)||(Formula("-sin(theta)cos(theta)/r^2*utheta+sin(theta)cos(theta)/r*uthetar")==$student));
}));

#### Check for the partial derivatives in either order: urtheta or uthetar
ANS(Formula("cos^2(theta)/r*ur+sin(theta)*cos(theta)/r*uthetar")->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return (($correct == $student)||(Formula("cos^2(theta)/r*ur+sin(theta)*cos(theta)/r*urtheta")==$student));
}));

ANS(Formula("-sin(theta)*cos(theta)/r^2*utheta+cos^2(theta)/r^2*uthetatheta")->cmp());

### PART 5

ANS(Formula($polarcontext,"urr+1/r*ur+1/r^2*uthetatheta")->cmp());
ENDDOCUMENT(); 
