## MO(1)
# WeBWorK problem written by Alex Jordan 
# Portland Community College
#
# Template:
# Rationalize the denominator
#
#       2pi*sqrt(L/32)
#
#
# Last updated: Spoddeck 11/10/15
# ENDDESCRIPTION


##############################################


DOCUMENT();

loadMacros(
      "PGstandard.pl", 
      "MathObjects.pl",
      "contextLimitedPowers.pl",
      "contextFraction.pl",
      "PGML.pl",
      "PGcourse.pl",
    );

TEXT(beginproblem());

##############################################

$expression = "2\pi\sqrt{\dfrac{L}{32}}";
$stringans = "pi*(sqrt(2*L))/(4)";


    # code essentially from Davide Cervone 4/25/10
    ###########################
    #
    # Subclass the numeric functions
    #
    package my::Function::numeric;
    our @ISA = ('Parser::Function::numeric');

    #
    # Override sqrt() to return a special value (usually 1) when evaluated
    # effectively eliminating it from the product.
    #
    sub sqrt {
      my $self = shift;
      my $value = $self->context->flag("setSqrt");
      return $value+1 if $value && $_[0] == 1; # force sqrt(1) to be incorrect
      return $value if $value;
      return $self->SUPER::sqrt(@_);
    }

    #
    # end of subclass
    #
    package main;

    ###########################
#    Context("Fraction-NoDecimals");
#    Context()->operators->undefine("+","-");
    Context()->strings->add("cannot be combined"=>{});

    Context()->variables->are(
      L => ["Real", limits => [0,2]],    # only needed if L is used in the square roots
    );
    #
    # make sqrt() use our subclass
    #
    Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
    Context()->flags->set(reduceConstantFunctions=>0);
    #
    # Don't allow fractional powers (avoids 1/2 power)
    # [Could subclass exponentiation to handle that as well]
    #
    LimitedPowers::OnlyPositiveIntegers();

    $ans = Compute($stringans);


##############################################


BEGIN_PGML

The time in seconds that a pendulum swings back and forth can be calculated by the formula below. The pendulum is [`L`] feet long.

Rationalize the denominator of the radical expression and make sure your answer is simplified.

>>[``[$expression]``]<<

If you need to enter a square root symbol, you need to type *sqrt* and use parentheses to enclose what goes inside the radical.  
For example, type *sqrt(3)* to enter [`\sqrt{3}`].

If you need the [`\pi`] symbol, type *pi*.

    [____________________]

END_PGML


##############################################
$showPartialCorrectAnswers = 1;


ANS($ans->cmp(typeMatch=>Formula("1"),
              checker => sub {
      my ($correct,$student,$ans) = @_;
      return 0 if $ans->{isPreview} || $correct != $student;
      #
      #  Get parsed formula for student and correct answers
      #
      $student = $ans->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      #
      # check if equal when sqrt's are replaced by 1
      #
      Context()->flags->set(setSqrt => 1);
      delete $correct->{test_values}, $student->{test_values};
      my $OK = ($correct == $student);
      Context()->flags->set(setSqrt => 0);
      #
      Value::Error("Check to see if your answer is simplified.") unless $OK;
      return $OK;
    },formatStudentAnswer=>"reduced"));


##############################################


BEGIN_PGML_SOLUTION

We will use that:

    [`` \sqrt{x}\sqrt{x} = \sqrt{x^2} = x ``]

This will help us get rid of the radical in the denominator by multiplying it with itself:

    [``
\begin{aligned}
   2\pi\sqrt{\dfrac{L}{32}}
                &= 2\pi\dfrac{\sqrt{L}}{\sqrt{32}} \\
                &= 2\pi\dfrac{\sqrt{L}}{\sqrt{32}}\cdot\dfrac{\sqrt{32}}{\sqrt{32}} \\
                &= 2\pi\dfrac{\sqrt{32L}}{\sqrt{32^2}} \\
                &= 2\pi\dfrac{\sqrt{16*2L}}{32} \\
                &= 4\pi\dfrac{\sqrt{2L}}{16} \\
                &= \pi\dfrac{\sqrt{2L}}{4} \\
\end{aligned}
    ``]

The answer can also be written as [`\dfrac{\pi\sqrt{2L}}{4}`].


END_PGML_SOLUTION


##############################################


ENDDOCUMENT();
