# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('3.2')
## Problem1('1')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Numeric");

$v1 = non_zero_random(-5,5,1);
$v1 =~ s/(-)?1/~~1/;
$v2 = non_zero_random(-5,5,1);
$ans1 = Compute( "$v2 x" );
$ans2 = Compute( "$v2 x + 1" );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Let
\(W =\) {\(\begin{bmatrix} $v1 x\\ y\end{bmatrix} \in \mathbb{R}^2 :
x\in \mathbb{R}\)}.
Give a value of \(y\) so that \(W\) is a subspace.
$BR
\( y = \) \{ $ans1->ans_rule(5) \}
$BR
Give a value of \(y\) so that \(W\) is not a subspace.
$BR
\( y = \) \{ $ans2->ans_rule(5) \}

END_TEXT
Context()->normalStrings;

ANS( $ans1->cmp( checker=>sub {
		   my ( $c, $s, $a ) = @_;
                   return Compute(0) == $s->D()->D() &&
                          Compute(0) == Formula($s)->eval(x=>0); } ) );
ANS( $ans1->cmp( checker=>sub {
		   my ( $c, $s, $a ) = @_;
                   return Compute(0) != $s->D()->D() ||
                          Compute(0) != Formula($s)->eval(x=>0); } ) );

ENDDOCUMENT();