# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('3.2')
## Problem1('2')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Vector");

$v11 = random(2,5,1);
$v12 = random(-5,-2,1);
$v13 = random(1,4,1);
$v21 = random(1,4,1);
$v22 = random(1,4,1);
$v23 = random(-4,-1,1);

$v1 = ColumnVector( $v11, $v12, $v13 );
$v2 = ColumnVector( $v21, $v22, $v23 );
$ans1 = $v1 + $v2;
$ans2 = $ans1 - ColumnVector( 0, 0, 1 );

$d = $v11*$v22 - $v12*$v21;  # guaranteed positive

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Let \(\vec v_1 = $v1\) and \(\vec v_2 = $v2\).
$PAR
Give a vector \(\vec v_3\) so that the vectors
\(\vec v_1\), \(\vec v_2\), \(\vec v_3\)
are ${BBOLD}not$EBOLD linearly independent.
$BR
\(\vec v_3 = \) \{ $ans1->ans_array() \}

$PAR
Give a vector \(\vec v_3\) so that the vectors
\(\vec v_1\), \(\vec v_2\), \(\vec v_3\)
${BBOLD}are$EBOLD linearly independent.
$BR
\(\vec v_3 = \) \{ $ans2->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $ans1->cmp( checker=>sub {
		     my ( $c, $s, $a ) = @_;
		     my ($s1, $s2, $s3) = $s->value;
		     my $k1 = ($s1*$v22 - $s2*$v21)/$d;
		     my $k2 = ($s2*$v11 - $s1*$v12)/$d;
		     return $k1*$v13 + $k2*$v23 == $s3;
		 } ) );
$zeroVec = ColumnVector([0,0,0]);
ANS( $ans2->cmp( checker=>sub {
		     my ( $c, $s, $a ) = @_;
		     my ($s1, $s2, $s3) = $s->value;
		     my $k1 = ($s1*$v22 - $s2*$v21)/$d;
		     my $k2 = ($s2*$v11 - $s1*$v12)/$d;
		     return $zeroVec == $s || $k1*$v13 + $k2*$v23 != $s3;
		 } ) );

ENDDOCUMENT();
