# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('3.1')
## Problem1('1')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Vector2D");

$a1 = non_zero_random(-3,3,1);
$a2 = non_zero_random(-3,3,1);
$m  = list_random(-3,-2,2,3);
$amat = Matrix( [ [ $a1, $a2 ], [ $m*$a1, $m*$a2 ] ] );
$ans1 = ColumnVector( 1, $m );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider the image of the linear transformation
\(T(\vec x) = A\vec x\), where \(A = $amat\).  Complete
the following statement:
$BR
The image of \(T\) consists of all vectors of the form
\( k \) \{ $ans1->ans_array() \}, where \(k\) is some scalar.

END_TEXT
Context()->normalStrings;

ANS( $ans1->cmp( checker=>sub {
		     my ($c,$s,$a) = @_;
		     return $c->isParallel($s) } ) );

ENDDOCUMENT();
