## DESCRIPTION
## First order ODEs: linear ODEs and the integrating factor
## ENDDESCRIPTION

## KEYWORDS('differential equations','first order','integrating factor')

## DBsubject('Differential Equations')
## DBchapter('First Order Differential Equations')
## DBsection('Linear ODEs and the Integrating Factor')
## Date('02/07/2012')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Elementary Differential Equations')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
#"parserAssignment.pl",
"unionLists.pl",
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


#############################################
#  Setup 

Context("Numeric")->variables->are(t=>'Real');

$a = random(2,7,1);
do { $b = random(2,7,1); } until ($a != $b);

$homogeneous = Compute("e^(-$a t)");
$particular = Formula("$b / $a"); # want as a formula for typeMatch reasons
$ivp = Compute("($b/$a) * (1 - e^(-$a t)) ");


#############################################
#  Part 1:
#  Find a single solution to the homogeneous

Context()->texStrings;
BEGIN_TEXT
Consider the intial value problem
$PAR
\[
\frac{dy}{dt} + $a y = $b, 
\ \ \
y(0) = 0.
\]
\{ BeginList("OL",type=>'a') \}
$ITEM
Find a nonzero solution to the associated homogeneous
differential equation.
$BR
\( y = \) \{ ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


ANS( $homogeneous->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $y  = Formula($student);   # get the student answer
    my $yp = Formula($y->D('t')); # take the partial derivative of the student answer

    return ($yp) == (-1 * $a * $y); # is more reliable than ($yp + $a * $y) == Formula(0)
}));



####################################################
#  Part 2
#  Find a particular solution to the non-homogeneous

Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM
Find a particular solution to the nonhomogeneous
differential equation.
$BR
\( y = \) \{ ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


ANS( $particular->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $y  = Formula($student);
    my $yp = Formula($y->D('t'));

    return ($yp + $a * $y) == Formula("$b");
}));


####################################################
#  Part 3
#  Solve the initial value problem

Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM
Find the solution to the initial value problem.
$BR
\( y = \) \{ ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
\{ EndList("OL") \}
END_TEXT
Context()->normalStrings;

ANS( $ivp->cmp() );

COMMENT("MathObject version.");

ENDDOCUMENT();