## DESCRIPTION
## First order ODEs: linear ODEs and the integrating factor
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Integrating factor)
## Date(02/07/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('differential equations','first order','integrating factor')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


#############################################
#  Setup 

Context("Numeric")->variables->are(t=>'Real');

$a = random(2,7,1);
do { $b = random(2,7,1); } until ($a != $b);

$homogeneous = Compute("e^(-$a t)");
$particular = Formula("$b / $a"); # want as a formula for typeMatch reasons
$ivp = Compute("($b/$a) * (1 - e^(-$a t)) ");


#############################################
#  Part 1:
#  Find a single solution to the homogeneous

Context()->texStrings;
BEGIN_TEXT
Consider the intial value problem
$PAR
\[
\frac{dy}{dt} + $a y = $b, 
\ \ \
y(0) = 0.
\]
\{ BeginList("OL",type=>'a') \}
$ITEM
Find a nonzero solution to the associated homogeneous
differential equation.
$BR
\( y = \) \{ ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


ANS( $homogeneous->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $y  = Formula($student);   # get the student answer
    my $yp = Formula($y->D('t')); # take the partial derivative of the student answer

    return ($yp) == (-1 * $a * $y); # is more reliable than ($yp + $a * $y) == Formula(0)
}));



####################################################
#  Part 2
#  Find a particular solution to the non-homogeneous

Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM
Find a particular solution to the nonhomogeneous
differential equation.
$BR
\( y = \) \{ ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


ANS( $particular->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $y  = Formula($student);
    my $yp = Formula($y->D('t'));

    return ($yp + $a * $y) == Formula("$b");
}));


####################################################
#  Part 3
#  Solve the initial value problem

Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM
Find the solution to the initial value problem.
$BR
\( y = \) \{ ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
\{ EndList("OL") \}
END_TEXT
Context()->normalStrings;

ANS( $ivp->cmp() );

;

ENDDOCUMENT();
