## DESCRIPTION
## First order ODEs: linear ODEs and the integrating factor
## ENDDESCRIPTION

## KEYWORDS('differential equations','first order','integrating factor')

## DBsubject('Differential Equations')
## DBchapter('First Order Differential Equations')
## DBsection('Linear ODEs and the Integrating Factor')
## Date('02/07/2012')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Elementary Differential Equations')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"parserAssignment.pl",
);

TEXT(beginproblem());


#############################
#  Setup 

Context("Numeric")->variables->are(
t=>"Real",
r=>"Real",
C=>"Real",
);
Context()->variables->set(
t=>{limits=>[0.5,1.5]},
r=>{limits=>[-1,-1]},
C=>{limits=>[-1,-1]},
);
parser::Assignment->Allow;

$k = random(2,9,1);

$answer = Compute("r = ($k e^t + C ) / (ln(t))");



#############################
#  Main text 

Context()->texStrings;
BEGIN_TEXT
Find the general solution to 
$PAR
\[
t \ln(t) \frac{dr}{dt} + r = $k t e^t.
\]
$BR
Use \( C \) to denote the arbitrary constant in your answer.
$BR
$BR
\{ ans_rule(50) \}
\{ AnswerFormatHelp("equations") \}
END_TEXT
Context()->normalStrings;


##############################
#  Answer evaluation 

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp( checker => sub {
    my ( $correct, $student, $self ) = @_;
    my $stu   = Formula($student->{tree}{rop});

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( Formula($stu->D('C'))==Formula(0) );

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu_t  = Formula($stu->D('t'));
    my $t = Formula(t); # we can't just use 't' below, we need '$t'
    return ($t * ln($t) * $stu_t + $stu) == ($k * $t * e**($t));

    #my $stuDE = Formula( (($t**2+$kk)*$stu_t + 2*$t*$stu - $t**2*($t**2+$kk)) );

    #return ($stuDE==Formula(0));

}));



COMMENT("MathObject version.");

ENDDOCUMENT();
