## DESCRIPTION
## First order ODEs: linear ODEs and the integrating factor
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Integrating factor)
## Date(01/01/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('2.2')
## Problem1('22')
## KEYWORDS('differential equations','first order','integrating factor')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserAssignment.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup 

Context("Numeric")->variables->are(
t=>"Real",
y=>"Real",
C=>"Real",
);
Context()->variables->set(
t=>{limits=>[-4,4]},
y=>{limits=>[-1,-1]},
C=>{limits=>[-1,-1]},
);
parser::Assignment->Allow;

$k = random(2,5,1);
$kk = $k * $k;

$answer = Compute("y = ( t^5/5 + $kk t^3/3 + C ) / (t^2 + $kk)");



#############################
#  Main text 

Context()->texStrings;
BEGIN_TEXT
Find the general solution to 
$PAR
\[
(t^2 + $kk) y^{\,\prime} + 2 t y = t^2 (t^2 + $kk).
\]
$BR
Enter your answer as \(y = \dots\) .
$BR
Use \( C \) to denote the arbitrary constant in your answer.
$BR
$BR
\{ ans_rule(50) \}
\{ AnswerFormatHelp("equations") \}
END_TEXT
Context()->normalStrings;


##############################
#  Answer evaluation 

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp( checker => sub {
    my ( $correct, $student, $self ) = @_;
    my $stu   = Formula($student->{tree}{rop});

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( Formula($stu->D('C'))==Formula(0) );
    return 0 if $stu->isConstant;
    
    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu_t  = Formula($stu->D('t'));
    my $t = Formula(t); # we can't just use 't' below, we need '$t'

    #my $stuDE = Formula( (($t**2+$kk)*$stu_t + 2*$t*$stu - $t**2*($t**2+$kk)) );

    #return ($stuDE==Formula(0));
    return Formula( (($t**2+$kk)*$stu_t + 2*$t*$stu) ) == Formula( ($t**2*($t**2+$kk)) );

}));

ENDDOCUMENT();
