## DBsubject(Abstract algebra)
## DBchapter(Rings)
## DBsection(Ideals and homomorphisms)
## Institution(University of Massachusetts Amherst)
## Author(Daniel Nichols and Siman Wong)
## Level(3)
## MO(1)
## KEYWORDS('maximal ideals')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "algebraMacros.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());
$showPartialCorrectAnswers = 0;


@primesList = ( 2, 3, 5, 7 );
@compositeList = ( 4, 6, 8, 9, 10, 12, 14, 15, 16, 18, 20 );

# pick a composite g = gcd( e, modulus )
$g = $compositeList[ random( 1, @compositeList - 1, 1 ) ];

# pick coprime e', m' which will be multiplied by the modulus
$ePrime = random( 2, 6, 1 );
$mPrime = random( $ePrime + 1, $ePrime + 6, 1 ) until gcd( $ePrime, $mPrime ) == 1;

$modulus = $g * $mPrime;
$e = $g * $ePrime;

# write (e) = ( gcd( e, modulus ) ). maximal ideals containing this
# ideal are generated by the primes that divide gcd( e, modulus )
$ans[ 0 ] = List( grep{ $g % $_ == 0 } @primesList );


	

TEXT(EV3(<<'EOT'));

It is a fact that every ideal of \( \{ cyclic( $modulus ) \} \) is of the form  \( ( b ) \) for some element \( b \) of \( \{ cyclic( $modulus ) \} \). $BR $BR

(a) Determine all maximal ideals of \( \{ cyclic( $modulus ) \} \) containing the ideal \( ( $e ) \). Enter a generator for each of these ideals. That is, if you think \( ( $e ) \) is contained in the maximal ideals \( ( a ) \) and \( ( b ) \), enter \( a, b \). $BR
\{ ans_rule( 15 ) \} $BR

EOT

# special answer checker -- accept any generator for a given ideal
# (p) is the same as (k) for any k such that gcd( k, modulus ) = p
# (a) = (b) for any a = b ( mod modulus )
ANS( $ans[ 0 ]->cmp(
	checker => sub {
		my( $correct, $student, $ansHash ) = @_;
		return ( gcd( $student % $modulus, $modulus ) == $correct % $modulus ? 1 : 0 );
	}
) );



COMMENT( "" );


ENDDOCUMENT(); 
