## MO(1)
# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "pccTables.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl",
);


TEXT(beginproblem());

##############################################


Context("Numeric");
Context()->flags->set(reduceConstants=>0);

$a[0] = random(-3,-1,1);
$a[1] = $a[0] + random(1,6,1);
$a[2] = $a[1] + random(1,6,1);
$a[3] = $a[2] + random(1,6,1);
$a[4] = $a[3] - random(1,6,1);


$multiansg = MultiAnswer(Compute("$a[4]"), Compute("$a[3]"), Compute("$a[2]"), Compute("$a[1]"), Compute("$a[0]"))->with(
  singleResult => 1,checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( @stu ) = @{$student};
      my ( @cor ) = @{$correct};
      my $return = 1;
      for my $i (0..4) {
           if ($stu[$i] != $cor[$i]) {$return = 0;}
           };
      return $return;
      }
);


##############################################

BEGIN_PGML

The first table gives values for the function \(f\). 

END_PGML
BEGIN_TEXT
$PAR
\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(f(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},"\($a[0]\)","\($a[1]\)","\($a[2]\)","\($a[3]\)","\($a[4]\)"]
],
allcellcss=>"padding:6pt;"
)
\}
$PAR
END_TEXT
BEGIN_PGML


Let [`g(x)=f(-x)`].  Fill in as many values in the table below as possible.  

END_PGML
BEGIN_TEXT
$PAR

\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(g(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},$multiansg->ans_rule(5), $multiansg->ans_rule(5), $multiansg->ans_rule(5), $multiansg->ans_rule(5), $multiansg->ans_rule(5)]
],
allcellcss=>"padding:6pt;"
)
\}

END_TEXT


##############################################

$showPartialCorrectAnswers = 1;

ANS( $multiansg->cmp() ); 

##############################################

BEGIN_SOLUTION

\(g\) is the function \(f\) reflected about the \(y\)-axis (which is a horizontal transformation). So the points that are represented in the table for \(f\) need to be refelcted over the \(y\)-axis to get points that are on \(g\)'s graph:

\[\begin{aligned}
(-2,$a[0])&\to(2,$a[0])\\
(-1,$a[1])&\to(-1,$a[1])\\
(0,$a[2])&\to(0,$a[2])\\
(1,$a[3])&\to(-1,$a[3])\\
(2,$a[4])&\to(-2,$a[4])
\end{aligned}\]

And so we find \(g(-2)=$a[4]\), \(g(-2)=$a[3]\), \(g(0)=$a[2]\), \(g(1)=$a[1]\), and \(g(2)=$a[0]\), giving us all five values for the table for \(g\).

END_SOLUTION


ENDDOCUMENT();