# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Given a function, find -f(x) and f(-x).
#
# Last update: Carl Yao 3/24/2018
#
# ENDDESCRIPTION

## DBCCSS('8.F','F-IF')
## DBsubject(Algebra)
## DBchapter(Function)
## DBsection(Domain and Range)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('function','combining functions', 'composition')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = random(2,5,1)*random(-1,1,2);
$b = random(2,5,1)*random(-1,1,2);

Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');
$f = Formula("$a*x+$b")->reduce;

Context()->operators->set(
'*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'u-' => {class => 'bizarro::UOP::minus', isCommand => 1},
);

$ans1 = Formula("-$a*x-$b")->reduce;
$evaluator1 = $ans1->cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      Context()->flags->set(bizarroMul=>1,bizarroNeg=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      my $OK = ($correct == $student);
      Context()->flags->set(bizarroMul=>0,bizarroNeg=>0);
      Value::Error("Your answer is equivalent to the correct answer, but not in the expected form${BR}Maybe it needs to be simplified more") unless $OK;
      return $OK;
});

$ans2 = Formula("-$a*x+$b")->reduce;
$evaluator2 = $ans2->cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      Context()->flags->set(bizarroMul=>1,bizarroNeg=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      my $OK = ($correct == $student);
      Context()->flags->set(bizarroMul=>0,bizarroNeg=>0);
      Value::Error("Your answer is equivalent to the correct answer, but not in the expected form${BR}Maybe it needs to be simplified more") unless $OK;
      return $OK;
});

##############################################


BEGIN_PGML

Let [`f`] be a function given by [`f(x)=[$f]`]. Find and simplify the following: 

a) [`-f(x)={}`][_________________________________________]{$evaluator1}  

a) [`f(-x)={}`][_________________________________________]{$evaluator2}  

    
END_PGML

##############################################

BEGIN_PGML_SOLUTION

a)

[`\begin{aligned}[t]
-f(x)
&=-([$f]) \\
&=[$ans1]
\end{aligned}`]

b)

[`\begin{aligned}[t]
f(-x)
&=[$a](-x)+[$b]\\
&=[$ans2]
\end{aligned}`]

END_PGML_SOLUTION

ENDDOCUMENT();    
