##DESCRIPTION
##Calculus: Limits and Derivatives
##ENDDESCRIPTION

##KEYWORDS('calculus', 'limits', 'derivatives')
##Tagged by YJ

## DBsubject('Calculus')
## DBchapter('Limits and Derivatives')
## DBsection('Derivatives')
## Date('5/26/2005')
## Author('Jeff Holt')
## Institution('UVA')
## TitleText1('Calculus: Early Transcendentals')
## EditionText1('5')
## AuthorText1('Stewart')
## Section1('2.8')
## Problem1('20')

## TitleText2('Calculus: Early Transcendentals')
## EditionText2('6')
## AuthorText2('Stewart')
## Section2('2.7')
## Problem2('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros("PG.pl",
           "PGbasicmacros.pl",
           "parserMultiAnswer.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = non_zero_random(2,3,1);
$b = random(2,8);
$b2= $b**$a;

Context('Numeric');

$mp = MultiAnswer("x^$a", $b)->with(
 checker => sub {
   my ($correct,$student) = @_;
   my $stuformula = Formula($student->[0]);
   my $tmp = $stuformula->substitute(x=>"x-$correct->[1]+$student->[1]");
   my $tmp2 = $correct->[0]->eval(x=>$correct->[1]);
   my $tmp3 = $stuformula->eval(x=>$student->[1]);
   return 1 if ($tmp==$correct->[0] && $tmp2==$tmp3);
   return 0;
  }
);

BEGIN_TEXT
The limit
\[
    \lim_{h\rightarrow 0}\frac{($b+h)^$a - $b2}{h}
\]
represents the derivative of some function \(f(x)\) at some number \(a\).
Find \(f\) and \(a\).
$BR
$BR
\(f(x)\) = \{ $mp->ans_rule(25) \}
$BR
$BR
\(a\) = \{ $mp->ans_rule(25) \}
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();        # This should be the last executable line in the problem.
