# WeBWorK problem written by Carl Yao
# Portland Community College
#
#
# Last updated: Yao 7/3/2018
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','logarithm','equation')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################
DOCUMENT();

loadMacros(
   "PGstandard.pl",
   "MathObjects.pl",
   "parserAssignment.pl",
   "PGML.pl",
   "PCCmacros.pl",
   "answerHints.pl",
   "contextLimitedRadical.pl",
   "PGcourse.pl",
);

##############################################
    
Context("LimitedRadical");
#Context()->flags->set(reduceConstants=>0, reduceConstantFunctions=>0, formatStudentAnswer=>parsed);
                               
Context()->operators->redefine(',',using=>',',from=>'Numeric');
Context()->operators->redefine('or',using=>',',from=>'Numeric');
Context()->operators->set(
      ','=>{string=>' or ',TeX=>'\hbox{ or }'},
      'or'=>{string=>' or ',TeX=>'\hbox{ or }'}
    );
Context()->lists->set(List => {separator => " or "});

# add solution strings to context- this means that if
# students enter these (and they are not correct), then 
# WW will not give a Context warning
Context()->strings->add("no real solutions"=>{},
  "no real solution"=>{alias=>'no real solutions'}, 
  "none"=>{alias=>'no real solutions'}, 
  );

$mymsg = "Are you sure your solution satisfies the equation?";
Context()->{error}{msg}{"Unexpected character '='"} 
  = "$mymsg";
  
$var = RandomVariableName('type' => variable);
Context()->variables->are($var => 'Real');

do {
   $c = random(2,5,1);
   $bMinusA = random(-5,-1,1);
   $b = non_zero_random(-5,5,1);
   $a = $b-$bMinusA;
} until ($a!=0);

$ans = Compute("no real solutions");

$mymsg = "Incorrect";
Context()->{error}{msg}{"Unexpected character '='"} 
  = "$mymsg";

##############################################

TEXT(beginproblem());
BEGIN_PGML
Solve the quadratic equation

     [`` 10^{[$c][$var]}+[$a]=[$b] ``]
     
* If there is no solution, type *no real solutions*.

     [__________________________________]

END_PGML

##############################################


ANS($ans->cmp(checker=>sub {
     my ($correct,$student,$ansHash) = @_;
     if(!Value::String->typeMatch($student))
     {
        Value->Error("$mymsg");
        return 0;
     }
        $check = ($student->value eq "no real solutions");
        Value->Error("$mymsg") unless($check);
        return $check;
    },
));


##############################################

BEGIN_PGML_SOLUTION

    [``
\begin{aligned}
   10^{[$c][$var]}+[$a] &= [$b] \\
   10^{[$c][$var]} &= [$bMinusA]
\end{aligned}
    ``]

Since [` 10^{[$c][$var]} `] cannot have a negative value, this equation has no solution.

END_PGML_SOLUTION

##############################################

ENDDOCUMENT();
