## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Implicit differentiation)
## Institution(John Wiley and Sons)
## Author(Brian Burns)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Anton')
## EditionText1('9')
## Section1('3.1')
## Problem1('8')
## KEYWORDS('derivative')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "littleneck.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);
# allow the student to change the seed for this problem.


rand_button();

TEXT(beginproblem());

###################################
# Setup
$a=random(1,6,1);
$b=random(1,6,1);
Context("Fraction");  
$ab=Compute("$a/$b");
Context("Numeric");  
Context()->variables->add(y=>'Real',dx=>'Real',dy=>'Real'); 

$y=Formula("[x^3-x]/[x^2+1]")->reduce;
$dy=Formula("[x^4+4*x^2-1]/[x^2+1]^2")->reduce;
$f=Formula("x**2")->reduce;
$g=Formula("[$a*x+$b*y]/[$a*x-$b*y]")->reduce;
$F=Formula("$a*x**3-$b*x^2*y")->reduce;
$G=Formula("[$a*x+$b*y]")->reduce;
$df=Formula("3*$a*x**2-2*$b*x*y-$b*x^2*(dy/dx)")->reduce;
$dg=Formula("[$a+$b*(dy/dx)]")->reduce;
$Y=Formula("$b(1+x^2)")->reduce;
$X=Formula("$a*3*x**2-2*$b*x*y-$a")->reduce;
$ans=Compute("$X/$Y");
$ans_2=Compute("$ab*$dy");
## other valid answer is $dy
###################################
# Main text
Context()->texStrings;
BEGIN_TEXT
Find \(\frac{dy}{dx}\) by implicit differentiation when \[$f=$g\]
 
$PAR
\(\frac{dy}{dx}=\) \{ans_rule(50) \}
END_TEXT
###################################
# Answers

$showPartialCorrectAnswers = 0;

ANS($ans->cmp(checker => sub {
       my ($correct,$student) = @_;
       return $ans == $student || $ans_2 == $student;
     }));
#ANS($ans->cmp);

###################################
Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
Cross multiplying gives \[$F=$G\]
Differentiating gives \[$df=$dg\] \[\frac{dy}{dx}($Y)=$X\] \[\frac{dy}{dx}=$ans\]
Note that you can actually solve for \(y\), \[y=$ab$y\] and hence either by substitution or by quotient rule solve for \(\frac{dy}{dx}\) as follows
\[\frac{dy}{dx}=$ab$dy\]
END_SOLUTION
Context()->normalStrings;

###################################
ENDDOCUMENT();
