##DESCRIPTION
##  Algebra: factoring quadratic expressions
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring by grouping)
## Date(09/23/2009)
## Institution(University of Minnesota)
## Author(Jonathan Rogness)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('3.4')
## Problem1('52')
## KEYWORDS('algebra', 'factoring', 'quadratic', 'grouping')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 0;

########################################################################

Context("Numeric");
Context()->variables->add(y=>'Real'); 
Context()->variables->add(a=>'Real'); 
Context()->variables->add(b=>'Real'); 

$c = random(5,7,2);
$d = random(2,4,1);

$poly = Compute("$c*a*x - $c*b*x -$d*a*y + $d*b*y")->reduce;

$factors = List(Formula("a-b"),Formula("$c*x-$d*y"));

########################################################################

Context()->texStrings;
BEGIN_TEXT
Factor by grouping
\[
$poly.
\]
Factors (separate by commas): \{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;

########################################################################

ANS($factors->cmp(
  list_checker => sub {
    my ($cor,$stu,$ansHash,$value) = @_;
    if ( ($stu->[0] == $cor->[0] and $stu->[1] == $cor->[1]) or 
         ($stu->[1] == $cor->[0] and $stu->[0] == $cor->[1]) or
         ($stu->[0] == -($cor->[0]) and $stu->[1] == -($cor->[1])) or
         ($stu->[1] == -($cor->[0]) and $stu->[0] == -($cor->[1]))
    ){ return 2;} 
    if (($stu->[0])*($stu->[1]) != ($cor->[0])*($cor->[1])){
      Value->Error("The product of your factors is wrong.");
    }
  }
));

ENDDOCUMENT();        
