##DESCRIPTION
##  Algebra: Multiplying rational expressions 
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Date(09/30/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## MLT(Mutiply_fractions)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('4.2')
## Problem1('18')
## KEYWORDS('algebra', 'rational expressions', 'multiplying')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");
Context()->variables->add(y=>'Real');

$a = random(3,12,3);
$b = random(2,6,2);
$c = random(25,35,10);
$d = random(3,9,6);

$e = Compute("3*$a/($d)");
$f = Compute("$c/5"); 
$g = Compute("$b+1");

$ans1 = Compute("$e*x");
$ans2 = Compute("$f*y^$g");

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("One of your answers has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Multiply the two rational expressions 
\[
\frac{$a x^{$b}}{$c y^{$b}} \cdot \frac{15 xy}{$d x^{$b}y^2}.
\]
Express your answer in simplest form \(\frac{A}{B}\).$PAR
Answer: \(A =\) \{$mp->ans_rule(20)\} and \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();
