##DESCRIPTION
##  Algebra: Subtracting rational expressions 
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Date(10/21/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## MLT(Subtract_fractions)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('4.4')
## Problem1('4')
## KEYWORDS('algebra', 'rational expressions', 'subtracting')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");

$b = random(2,10,1);
$c = random(2,10,1);

$a = Compute("$c*$b");
$poly1 = Compute("-$a") -> reduce;
$poly2 = Compute("x^2-$b x") -> reduce;
$poly3 = Compute("$c") -> reduce;
$poly4 = Compute("x") -> reduce;

$ans1 = Formula("-$c");
$ans2 = Formula("x-$b");

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("One of your answers has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Subtract the two rational expressions
\[
\frac{$poly1}{$poly2} - \frac{$poly3}{$poly4}.
\]
Express your answer in simplest form \(\frac{A}{B}\). $PAR
Answer: \(A =\) \{$mp->ans_rule(20)\} and \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();
