##DESCRIPTION
##  Algebra: Multiplying rational expressions 
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Date(10/11/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## MLT(Mutiply_fractions)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('4.2')
## Problem1('30')
## KEYWORDS('algebra', 'rational expressions', 'multiplying')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");
Context()->variables->are(n=>'Real');

$a = random(3,7,2);
$b = random(2,4,2);
$c = random(2,5,1);

$poly1 = Compute("10n^2+(5*$a-2)n-$a") -> reduce;
#(5n-1)(2n+$a)
$poly2 = Compute("$c n^2+ (2*$c+1)n +2") -> reduce;
#(n+2)($c n +1)
$poly3 = Compute("$b n^2+ 2*$b n") -> reduce;
#(n+2)($b n)
$poly4 = Compute("2n^2+($a+2)n+$a") -> reduce;
#(n+1)(2n+$a)

$ans1 = Compute("5*$b n^2 - $b n");
$ans2 = Compute("$c n^2 + (1+$c)n + 1");

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("One of your answers has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Multiply the two rational expressions
\[
\frac{$poly1}{$poly2} \cdot \frac{$poly3}{$poly4}.
\]
Express your answer in simplest form \(\frac{A}{B}\).$PAR
Answer: \(A =\) \{$mp->ans_rule(20)\} and \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();
