##DESCRIPTION
##  Algebra: Multiplying and dividing rational expressions 
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Date(10/13/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('4.2')
## Problem1('48')
## KEYWORDS('algebra', 'rational expressions', 'dividing', 'multiplying')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");
Context()->variables->add(y=>'Real'); 
Context()->variables->add(z=>'Real'); 

$a = random(3,7,2);
$b = random(2,8,2);
$c = random(11,13,2);
$d = random(1,4,1);
$e = random(1,4,1);
$f = random(1,2,1);

$poly1 = Compute("$b x y^($d)") -> reduce;
$poly2 = Compute("$a x^($f)") -> reduce;
$poly3 = Compute("2*$a x y^($d+2)") -> reduce;
$poly4 = Compute("2*$c y") -> reduce;
$poly5 = Compute("$a y^2 z^($f)") -> reduce;
$poly6 = Compute("$c x^($e)") -> reduce;

if ($a == 3 && $b == 6) {
$ans1 = Formula("($b/3) x^($e-$f+2)*y^(2*$d-1)");
$ans2 = Formula("($a)/3 z^($f)");
} else{
$ans1 = Formula("$b x^($e-$f+2)*y^(2*$d-1)");
$ans2 = Formula("$a z^($f)");
}

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("One of your answers has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Perform the indicated operations involving the three rational expressions 
\[
\frac{$poly1}{$poly2} \cdot \frac{$poly3}{$poly4} \div \frac{$poly5}{$poly6}.
\]
Express your answer in simplest form \(\frac{A}{B}\).$PAR
Answer: \(A =\) \{$mp->ans_rule(20)\} and \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();
