##DESCRIPTION
##  Algebra: Summing rational expressions 
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Date(10/19/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## MLT(Add_fractions)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('4.3')
## Problem1('18')
## KEYWORDS('algebra', 'rational expressions', 'adding')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");

$a = random(1,5,1);
$b = random(2,5,1);
$c = random(3,5,2);


$poly1 = Compute("$b x-1") -> reduce;
$poly2 = Compute("$a x^2") -> reduce;
$poly3 = Compute("$c(x-2)") -> reduce;
$poly4 = Compute("$a x^2") -> reduce;

$ans1 = Compute("($b+$c)x-2*$c-1");
$ans2 = Compute("$a x^2");

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("One of your answers has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );


########################################################################

Context()->texStrings;
BEGIN_TEXT
Add the two rational expressions
\[
\frac{$poly1}{$poly2} + \frac{$poly3}{$poly4}.
\]
Express your answer in simplest form \(\frac{A}{B}\). $PAR
Answer: \(A =\) \{$mp->ans_rule(20)\} and \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();
