##DESCRIPTION
##  Algebra: Dividing rational expressions 
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Date(09/30/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## MLT(Divide_fractions)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('4.2')
## Problem1('22')
## KEYWORDS('algebra', 'rational expressions', 'dividing')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");
Context()->variables->are(a=>'Real');
Context()->variables->add(b=>'Real');
Context()->variables->add(c=>'Real');

$d = random(24,30,6);
$e = random(4,7,1);
$g = random(3,9,6);
$h = random(2,3,1);
$i = random(3,7,1);

$f = Compute("7");
$j = Compute("$e+$h");
$k = Compute("$i-2");
$l = Compute("$d/3");
$m = Compute("$f*$g/3");

$ans1 = Compute("$l*b^$j*c^$k");
$ans2 = Compute("$m*a");

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("One of your answers has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Divide the two rational expressions 
\[
- \frac{$d a b^{$e}}{$f c} \div - \frac{$g a^2 c}{b^{$h}c^{$i}}.
\]
Express your answer in simplest form \(\frac{A}{B}\).$PAR
Answer: \(A =\) \{$mp->ans_rule(20)\} and \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();
