##DESCRIPTION
##  Algebra: exponents and radicals 
##ENDDESCRIPTION


## DBsubject(Arithmetic)
## DBchapter(Irrational numbers)
## DBsection(Simplify radical numbers)
## Date(10/27/2009)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## MLT(rewrite_a_radical)
## Level(2)
## MO(1)
## TitleText1('Algebra for College Students')
## AuthorText1('Kaufmann, Schwitters')
## EditionText1('8')
## Section1('5.2')
## Problem1('46')
## KEYWORDS('algebra', 'radicals', 'fraction', 'exponents')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


########################################################################

install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 0;

TEXT(beginproblem());

########################################################################

Context("Numeric");

$a = random(1,3,1);

$b = Compute("$a+1");
$c = Compute("$a+2");
$d = Compute("$c^2*5");

$numerator = Real("$a*$c");
$denominator = Real("$b");
$radical = Real("5");

########################################################################

$mp = MultiAnswer($numerator,$denominator)->with(
         singleResult => 0,
         separator => "/",
         tex_separator => "/",
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator (and maybe more) is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("Your fraction has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Change the radical
\[
\frac{$a}{$b}\sqrt{$d}
\]
into simplest radical form \(\frac{A}{B} \sqrt{C}\), where \(A\), \(B\), and \(C\) are all integers.
$PAR
Answer: \(A =\) \{$mp->ans_rule(15)\}, \(B =\) \{$mp->ans_rule(15)\}, and \(C =\) \{ ans_rule(15) \}
$HR
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );
ANS( $radical->cmp );

ENDDOCUMENT();        
