# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions and their Graphs')
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1()
## Problem1()
## Author('Alex Jordan')
## Institution('PCC')



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PCCgraphMacros.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

##############################################



Context("Numeric");

$r[0] = random(-4,-1,1);
for my $i (1..3)
{
  $r[$i] = $r[$i-1] + random(1,5,1);
}

$p = Compute("-(x-$r[0])/((x-$r[1])*(x-$r[2]))")->reduce;

for my $i ($r[0]..$r[3]-1)
{$vals[$i-$r[0]] = abs($p->eval(x=>$i+0.5));}

$ymax = 1.3*max(@vals);

$xmax = max($r[3]+1,5);
$gr = init_graph(-5,-$ymax,$xmax,$ymax,axes=>[0,0],grid=>[$xmax+5,10],size=>[400,400]);

$zoommin = $r[0]-0.5;
$zoommax = $r[0]+0.5;
$zoomheight = max(abs($p->eval(x=>($zoommin))), abs($p->eval(x=>($zoommax))));
$gr2 = init_graph($zoommin,-$zoomheight,$zoommax,$zoomheight,axes=>[0,0],size=>[200,200]);

$gr -> lb('reset');
$gr2 -> lb('reset');



foreach my $i (1..$xmax) {
  $gr -> lb(new Label ( $i,-0.8, $i,'black','center','middle'));
  $gr -> lb(new Label (-$i,-0.8,-$i,'black','center','middle'));
}

foreach my $i (1..5) {
  $label = $r[0]+($i-3)*0.2;
  $gr2 -> lb(new Label ( $label,-0.08*$zoomheight, $label,'black','center','middle'));
}

$eps = 0.01;
$d[0] = $r[1]-$eps;
$d[1] = $r[1]+$eps;
$d[2] = $r[2]-$eps;
$d[3] = $r[2]+$eps;


add_functions($gr, "$p for x in <-5,$d[0]> using color:blue and weight:1");
add_functions($gr, "$p for x in <$d[1],$d[2]> using color:blue and weight:1");
add_functions($gr, "$p for x in <$d[3],$xmax> using color:blue and weight:1");

add_functions($gr2, "$p for x in <$zoommin,$zoommax> using color:blue and weight:1");

$gr->stamps( closed_circle($r[0],0,'black') );
$gr2->stamps( closed_circle($r[0],0,'black') );

##############################################

$refreshCachedImages = 1;
BEGIN_PGML

Find a formula for a rational function [`f`] that might have this graph. Do not over-complicate things: assume that all degrees are as small as can be compatible with this graph.

[@EnlargeImageStatementPGML()@]**


>>[@image(insertGraph($gr), width=>400, height=>400, tex_size=>800)@]* <<

>> Zoomed in near [`x=[@$r[0]@]`]: [@image(insertGraph($gr2), width=>200, height=>200, tex_size=>400)@]* <<


    [`f(x)={}`][________________________________]

END_PGML

##############################################

$showPartialCorrectAnswers = 1;

ANS( $p->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"abs($c0) * $p");
      return $correct == $student;
    }
) );

##############################################

BEGIN_PGML_SOLUTION

We need to see from the graph that there is a zero at [`[@$r[0]@]`]. Based on the zoomed-in graph, we can see that this zero has degree [`1`].

We also need to see that there are vertical asymptotes at [`[@$r[1]@]`] and [`[@$r[2]@]`]. Since the arms of the function go in opposite directions (one up, one down) as [`x`] approaches either [`[@$r[1]@]`] or [`[@$r[2]@]`], then the degrees of these vertical asymptotes must be odd: [`1`], [`3`], [`5`], etc. 

Since this is a rational function, having a zero at [`[@$r[0]@]`] means that [`(x-[@$r[0]@])`] is a factor of the numerator. Since it is a degree-1 zero, there is no exponent on this factor (other than [`1`]).

Having vertical asymptotes at [`[@$r[1]@]`] and [`[@$r[2]@]`] means that [`(x-[@$r[1]@])`] and [`(x-[@$r[2]@])`] are factors of the denominator. The instructions say to keep the degrees as small as is compatible with the graph, and taking the degrees to be [`1`] for both factors is compatible with the graph.

So far, this tells us that [`f(x)=c\frac{(x-[@$r[0]@])}{(x-[@$r[1]@])(x-[@$r[2]@])}`] for some constant [`c`]. The long term-behavior (horizontal asymptote) is the line [`y=0`] since the total degree of the numerator is less than the total degree of the denominator, and this agrees with what we see in the graph. 

For a really large [`x`]-value, the quantity [`\frac{(x-[@$r[0]@])}{(x-[@$r[1]@])(x-[@$r[2]@])}`] gives a positive output since all factors will be positive. And so [`c\frac{(x-[@$r[0]@])}{(x-[@$r[1]@])(x-[@$r[2]@])}`] will have the same sign as [`c`]. We can see in the graph that when [`x`] is large, that [`f(x)`] is negative, and therefore [`c`] needs to be negative. You can enter any formula that has this pattern to get credit for this problem, including just [`[$p]`] (where [`c=-1`]).

END_PGML_SOLUTION

##############################################


ENDDOCUMENT();