# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions and their Graphs')
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1()
## Problem1()
## Author('Alex Jordan')
## Institution('PCC')



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",

);
TEXT(beginproblem());



##############################################


Context("Numeric");

$r[0] = random(-4,-1,1);
for my $i (1..2)
{
  $r[$i] = $r[$i-1] + random(1,5,1);
}

$p = Compute("(x-$r[0])*(x-$r[1])*(x-$r[2])")->reduce;

##############################################

BEGIN_PGML

Find a formula for a degree three polynomial function [`f`] whose zeros are [`[@$r[0]@]`], [`[@$r[1]@]`], and [`[@$r[2]@]`]. Technically, there are many correct answers for this question.

    [`f(x) = {}`][________________________________________]

END_PGML


##############################################

ANS( $p->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $p");
      return $correct == $student;
    }
) );

##############################################

BEGIN_PGML_SOLUTION

Since the zeros are [`[@$r[0]@]`], [`[@$r[1]@]`], and [`[@$r[2]@]`], there are corresponding factors [`(x-[@$r[0]@])`], [`(x-[@$r[1]@])`], and [`(x-[@$r[2]@])`]. Since the polynomial is degree [`3`], there cannot be any more factors that use [`x`]. But there could be a constant factor. So any answer of the form [`c(x-[@$r[0]@])(x-[@$r[1]@])(x-[@$r[2]@])`] with [`c\neq0`] works. In particular, [`(x-[@$r[0]@])(x-[@$r[1]@])(x-[@$r[2]@])`] is one correct answer.



END_PGML_SOLUTION

ENDDOCUMENT();
