## DESCRIPTION
## Choosing and implementing the correct series test
## ENDDESCRIPTION

## KEYWORDS('calculus', 'series')

## DBsubject(Calculus - single variable)
## DBchapter(Infinite Sequences and Series)
## DBsection(Strategy for testing series)
## Date('02/01/2018')
## Author('Eric Errthum')
## Institution('Winona State University')
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('7')
## Section1('11.7')
## Problem1('6')

############################################################################
## Copyright (C) 2018 Eric Errthum, Ph.D.
##
## This work is licensed under the Creative Commons
## Attribution-NonCommercial-ShareAlike 3.0 Unported License
## <https://creativecommons.org/licenses/by-nc-sa/3.0/>
##

#######################
## Initialization

DOCUMENT();
loadMacros(
	"PGbasicmacros.pl",
	"PGanswermacros.pl",
	"parserPopUp.pl",
	"scaffold.pl",
	"parserOneOf.pl",
	"parserMultiAnswer.pl",
	);


TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

Context("Numeric");
Context( )->variables->add("n"=>"Real");

###################################
# Setup

$a=random(5,10,1);
$b=random(2,10,1);


$terms="\frac{1}{$a n+$b}";
$series="\sum_{n=1}^\infty  $terms";

$compfun1=Compute("1/n");
$compfun2=Compute("1/($a*n)");
$compfun=OneOf($compfun1,$compfun2);
$compterms="\frac{1}{n}";

$test= PopUp(["Select", "Divergence Test", "Geometric Test", "p-Series Test", "Integral Test", "Comparison Test", "Alternating Series Test", "Ratio Test", "Root Test"], "Comparison Test");

$popupComp=PopUp(["Select","Direct","Limit"],"Limit");
$popupComp2=PopUp(["Select","Direct","Limit"],"Limit");
$popupIneq=PopUp(["Select","<",">"],"<");
$popupReq=PopUp(["Select", "Greater than or equal to 1","Less than or equal to 1","Not 0 or Infinity"],"Not 0 or Infinity");
$popupConv1=PopUp(["Select","Converges","Diverges"],"Diverges");
$popupConv2=PopUp(["Select","Converges","Diverges"],"Diverges");

$limans=1/$a;

$divcon= PopUp(["Select","converges absolutely", "converges conditionally", "diverges"], "converges absolutely");

###########################################
# Problem Statement
BEGIN_TEXT
Test the series for convergence or divergence.
\[ $series .\]

END_TEXT

Scaffold::Begin(is_open  => "correct_or_first_incorrect");

Section::Begin("Part 1: Determine the Test");

BEGIN_TEXT
Which test is the best test to use for this series? $BR
\{ $test->menu\}.

END_TEXT

## Answer Evaluation

ANS($test->cmp);
Section::End();

###########################################
Section::Begin("Part 2: Perform the Test");

Scaffold::Begin(      can_open => "when_previous_correct",
      is_open  => "correct_or_first_incorrect");  

#only the active section is open, but students can open previous sections if they want. 

###########################################
Section::Begin("Part a: Identify the Comparable Function and Proper Comparison.");

BEGIN_TEXT
\{ beginproblem() \}

$PAR
The comparable series with known behavior is \(\displaystyle \sum_{n=1}^\infty\) \{ ans_rule(15)  \}.$BR

The comparison test we can use is \{$popupComp->menu\}.

END_TEXT

## Answer Evaluation
## Allows comparison to 1/n or 1/($a*n)

Context($compfun1->context);
$comp=MultiAnswer($compfun)->with(checker => sub{
      my ( $correct, $student, $self ) = @_;
      my ( $compstu) = @{$student};
      my ( $compfun ) = @{$correct};
      if ( $compstu == $compfun2) {$limans=1; $compterms="\frac{1}{$a n}"; return 1} elsif ($compstu == $compfun1) {return 1} else {return 0};
});
ANS($comp->cmp); 
ANS($popupComp->cmp);

Section::End();

###########################################
Section::Begin("Part b: Perform the Test");

BEGIN_TEXT
\{ beginproblem() \}

$PAR
Performing the Limit Comparison Test gives \(\displaystyle \lim_{n\to\infty}\frac{\frac{1}{$a n+$b}}{$compterms} = \) \{ ans_rule(15)  \} which is \{$popupReq->menu\}. Since the comparable series \{$popupConv1->menu\}, by \{$popupComp2->menu\} comparison, the original series \{$popupConv2->menu\}.

END_TEXT

## Answer Evaluation

ANS(num_cmp($limans));
ANS($popupReq->cmp);
ANS($popupConv1->cmp);
ANS($popupComp2->cmp);
ANS($popupConv2->cmp);


Section::End();
#########################################
Scaffold::End();

Section::End();


Scaffold::End();

######################################

ENDDOCUMENT();