##DESCRIPTION
##  Solving Linear Equations
##ENDDESCRIPTION

## KEYWORDS('linear', 'equations')

## DBsubject('Algebra') 
## BookTitle('Algebra: Form and Function')
## DBchapter('Equations and Inequalities') 
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.5')
## Problem1('')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");


Context("Numeric");

Context()->strings->add("no solutions"=>{}, "infinitely many"=>{}, "no solution"=>{alias=>"no solutions"}, "none"=>{alias=>"no solutions"}, "all reals"=>{alias=>"infinitely many"});
@vars = ("x","y","z","p","q","k","u","v","h","s");
@caps = ("C","N","M","R");
$c1 = random(0,scalar(@caps)-1,1);
$c2 = random(0,scalar(@vars)-1,1);
do {$c3 = random(0,scalar(@vars)-1,1)} until ($c3 != $c2);
do {$c4 = random(0,scalar(@vars)-1,1)} until ($c4 != $c2 && $c4 != $c3);
$var1 = $caps[$c1];
$var2 = $vars[$c2];
$var3 = $vars[$c3];
$var4 = $vars[$c4];
Context()->variables->are(lc($var1)=>"Real", lc($var2)=>"Real", lc($var3)=>"Real", lc($var4)=>"Real", uc($var1)=>"Real", uc($var2)=>"Real", uc($var3)=>"Real", uc($var4)=>"Real", );
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(2,10,1);
$b = random(-1,1,2)*random(2,9,1);


$frac = "$var1";
$case = random(0,1,1);
if ($case = 0) {
  $dispexpn = "$var2 = $a $var1 $var3 - $b $var1 $var4";
  $num = Formula("$var2")->reduce;
  $den = Formula("$a $var3 - $b $var4")->reduce;
} else {
  $dispexpn = "$var2 + $b $var1 $var3 $var4 = $a $var1"; 
  $num = Formula("$var2")->reduce;
  $den = Formula("$a - $b $var3 $var4")->reduce;
}




$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Solve the equation for \($var1\):
$PAR
\[
$dispexpn
\]
$PAR
$displayfrac
END_TEXT


Context()->normalStrings;



###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp(studentsMustReduceFractions=>1));
Context()->texStrings;
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle $var1 = \frac{$num}{$den}\)
  $PAR
  END_TEXT
} 
else  {
    BEGIN_TEXT
    $PAR
    $BBOLD Note: $EBOLD
    Write your answer as a single fraction, and type the numerator and denominator separately in the answer blanks provided.
    $PAR
    END_TEXT
} 
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
