##DESCRIPTION
##  Solving Linear Equations
##ENDDESCRIPTION

## KEYWORDS('linear', 'equations')

## DBsubject('Algebra') 
## BookTitle('Algebra: Form and Function')
## DBchapter('Equations and Inequalities') 
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.5')
## Problem1('')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");


Context("Numeric");

Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real",
L=>"Real",
F=>"Real",
Y=>"Real",
T=>"Real",
S=>"Real",
U=>"Real",
V=>"Real",
P=>"Real",
Q=>"Real",
R=>"Real",
W=>"Real",
Z=>"Real",
l=>"Real",
f=>"Real",
X=>"Real",);

$var = list_random('x','y','z','w','p','q','r','t','u','v');
$var1 = list_random('x','y','z',,'w','p','q');
$var2 = list_random('r','s','t','u','v');
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(2,10,1);
$b = random(-1,1,2)*random(2,9,1);

$frac = "y";

$num = Formula("L - $a x")->reduce;
$den = Formula("$b F")->reduce;
$texanswer = "\($var = \dfrac{L - $a x}{$b F} \)";
if ($b < 0) {
    $texanswer = "\($var = -\dfrac{L - $a x}{".(-$b)." F} \)";
}

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Solve the equation for \(y\):
$PAR
\[
$a x + $b y F = L.
\]
$PAR
$displayfrac
END_TEXT

if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    $texanswer
    $PAR
    END_TEXT
} 
else  {
    BEGIN_TEXT
    $PAR
    $BBOLD Note: $EBOLD
    Write your answer as a single fraction, and type the numerator and denominator separately in the answer blanks provided.
    $PAR
    END_TEXT
} 
Context()->normalStrings;

PeriodicStatus(); 



###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp(studentsMustReduceFractions=>1));

COMMENT('MathObject version', 'Features PeriodicRerandomization');
ENDDOCUMENT();
