## DESCRIPTION
## The geometry of vectors in R^n
## ENDDESCRIPTION

## Edited by Carl Yao

## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Vectors and vector arithmetic)
## Date(01/01/2013)
## Institution(Hope College)
## Author(Paul Pearson)
## MLT(GiveTwoPointsForVec)
## MLTleader(1)
## Level(1)
## TitleText1('Multivariable Mathematics')
## AuthorText1('Darin Stephenson')
## EditionText1('Version 3.0')
## Section1('1.1')
## Problem1('11')
## KEYWORDS('vectors')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "parserMultiAnswer.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "specialTrigValues.pl",
  "PGcourse.pl",
);

TEXT(beginproblem());


#############################
#  Setup

Context("Vector2D");
Context()->flags->set( ijk=>0 );

$A = non_zero_point2D(-9,9,1);
do { $B = non_zero_point2D(-9,9,1); } until ($A != $B);
$AB = Vector($B - $A);

($a,$b) = $AB->value;

$multians = MultiAnswer($A, $B)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $Astu, $Bstu ) = @{$student};
      my ( $Acor, $Bcor ) = @{$correct};
      if ( $Astu == Point(0,0) || $Bstu == Point(0,0) ) { 
           return 0;
      }
      if ( Vector($Bstu - $Astu) == $AB ) {
          return 1;
      } else {
           return 0;
      }
  }
);


######################
#  Main text

BEGIN_PGML

Find a representation of the vector [` \overrightarrow{AB} = [$AB] `] by giving appropriate values for the points [`A`] and [`B`] such that neither [`A`] nor [`B`] is the origin.

[@KeyboardInstructions(
"\n".'* Enter points, like [`(0,1)`]'
)@]**

    [`A =`][__________________]{$multians} 

    [`B =`][__________________]{$multians} 

END_PGML


########################

BEGIN_PGML_SOLUTION

If [`A`] is at [`(0,0)`], then B would be at [`([$a],[$b])`].

Since we cannot use [`(0,0)`] for [`A`], we can simply add the same number to either [`x`] or [`y`] values of those two points. For example:

    [``\begin{aligned}
  &A(0,0+1)\text{ and }B([$a],[$b]+1) \\
  &A(0+1,0)\text{ and }B([$a]+1,[$b]) \\
  &A(0+1,0+1)\text{ and }B([$a]+1,[$b]+1) \\
  ...
    \end{aligned}``]

END_PGML_SOLUTION

ENDDOCUMENT();
