## DESCRIPTION
## Precalculus: Trigonometric Identities
## ENDDESCRIPTION

## Tagged by cmd6a 5/11/06

## DBsubject(Trigonometry)
## DBchapter(Analytic trigonometry)
## DBsection(Product-to-sum & sum-to-product formulas)
## Institution(ASU)
## Level(2)
## MO(1)
## KEYWORDS('trigonometry','trigonometric identities','precalculus','identities')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric");

$b=random(2,8,1);
$a=random($b+1,$b+8,1);

#evaluate answers

$mp = MultiAnswer(Compute("$a-$b"), Compute("$a+$b"))->with(
  checker => sub {
    my ($correct, $student,$self) = @_;
    if($student->[0] == $correct->[0]) {
      return( ($student->[1] == $correct->[1]) ? [1,1] : [1,0] );
    } elsif($student->[0] == $correct->[1]) {
      return( ($student->[1] == $correct->[0]) ? [1,1] : [1,0] );
    } elsif($student->[1] == $correct->[1] or $student->[1] == $correct->[0]) {
      return [0,1];
    }
    return 0;
  }
);

Context()->texStrings;
BEGIN_TEXT
Use the product to sum formula to fill in the blanks in the identity below:

$BR
\( \sin($a x)\cos($b x) = \frac{1}{2}(\sin ( \) 
\{$mp->ans_rule(5) \} \( x)+\sin( \)  \{$mp->ans_rule(5) \} \( x)) \)
END_TEXT
Context()->normalStrings;

ANS($mp->cmp());



ENDDOCUMENT();        # This should be the last executable line in the problem.

