# DESCRIPTION
# Sigma notation
# ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Finite sequences and series)
## DBsection(Notation)
## Date(11/01/2012)
## Institution(Hope College)
## Author(Paul Pearson)
## MLT(write_out_terms)
## Level(2)
## KEYWORDS('calculus', 'sigma notation')

################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


################################
#  Setup

Context("ArbitraryString");

$n = random(5,8,1);

$sum_string = "1/2";
foreach my $i (3..$n) {
  $sum_string = $sum_string . "+1/$i";
}


$answer = String($sum_string);


################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Write out the following sum, one term for each value of \( k \).
Simplify each term as much as possible.  
For example, enter \( 1+4+9 \) instead of \( 1^2 + 2^2 + 3^2 \) or \( 14 \).
$PAR
\( \displaystyle \sum_{k=2}^{$n} \frac{1}{k} = \)
\{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;


#################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp( 
checker => sub {
  my ( $correct, $student, $ansHash ) = @_;
  #warn $student;
  #warn $correct;
  $student =~ s/ //g; # compress all blank space
  $student =~ s/~~+~~-/~~-/g; # convert +- to -
  $student =~ s/~~-~~+/~~-/g; # convert -+ to -
  $student =~ s/~~+~~+/~~+/g; # convert ++ to +
  $student =~ s/~~-~~-/~~+/g; # convert -- to +
  return $correct == String($student);
}
));

COMMENT('MathObject version.');

ENDDOCUMENT();
