# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Given the slope and a point, find a line's equation in slope-intercept form.
# Application problem. Slope is negative.
# 
# Last updated: Jordan 8/31/13; Carl Yao, 7/9/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Lines')
## KEYWORDS('linear','equation','slope','slope-intercept','application')
## DBCCSS('F-LE.2','8.F.3','A-CED','S-ID.7')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserAssignment.pl",
  "contextCurrency.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
Context()->variables->add(y=>'Real');
Context()->noreduce('(-x)-y','(-x)+y');
Context()->flags->set(showExtraParens=>0);

$m= -random(20000,50000,1000);
$mx = Formula("$m x")->reduce;
$b= random(600000,1000000,1000);

$x1=random(2,4,1);
$y1=$m*$x1+$b;
$x2=random(5,7,1);
$y2=$m*$x2+$b;
$x3=random(9,11,1);
$y3=$m*$x3+$b;
$y4=0;

$x4 = ($y4-$b)/$m;
if ($x4==int($x4)) {
   $out1 = "$x4 &= x";
   $yx4 = 2000 + $x4;
} else {
   $x4 = sprintf("%.2f", $x4);
   $out1 = "$x4 & \approx x";
   $yx4 = 2000 + int($x4);
}

$yx1 = 2000+$x1;
$yx2 = 2000+$x2;
$yx3 = 2000+$x3;

parser::Assignment->Allow;
$ansSI=Formula("y=$m*x+$b")->reduce;

Context("Currency");
Context()->flags->set(trimTrailingZeros=>1);
$cy1 = Currency($y1);
$cy2 = Currency($y2);
$cy3 = Currency($y3);
$cy4 = Currency($y4);
$cb = Currency($b);
$cm = Currency(-$m);
$cz = Currency("0");

Context("Numeric");
parser::Assignment->Allow;
Context()->variables->add(y=>'Real');
Context()->flags->set(tolType=>'absolute');
Context()->variables->set(x=>{limits=>[$x1,$x2]});
Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
);

$ansPS = Formula("y=$m*(x-$x1)+$y1");
$ansPS_evaluator = $ansPS -> cmp(
      checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is a line equation for the right line, but it is not in the expected point-slope form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
      return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

A company set aside a certain amount of money in the year 2000. The company spent exactly [`[$cm]`] from that fund each year on perks for its employees. In [` [$yx1]`], there was still [`[$cy1]`] left in the fund. 

Let [`x`] be the number of years since 2000, and let [`y`] be the amount of money, in dollars, left in the fund that year. Use a linear equation to model the amount of money left in the fund after so many years. 

a) A point-slope equation to model this is [______________]{$ansPS_evaluator}.

b) In the year [`[$yx3]`], there was [___________]{$cy3} left in the fund.

c) In the year [___________]{$yx4}, the fund will be empty.

END_PGML



##############################################

$s1=$y2-$y1;
$s2=$x2-$x1;
$s3=$m*$x1;
$s4=-$s3;
$s5=$m*$x3;
$s6=$y4-$b;
$s7=sprintf("%.2f", $s6/$m);

$outputy1 = $y1<0 ? "($y1)" : $y1;
$outputx1 = $x1<0 ? "($x1)" : $x1;

BEGIN_PGML_SOLUTION

a)    A line's equation in point-slope form looks like [` y=m(x-x_0)+y_0 `], where [`m`] is the slope, and [`(x_0,y_0)`] is some point we know to be on the line. We first need to find the line's slope. We have been told that the account decreases by [`[$cm]`] each year, so the slope of the linear model is [`[$m]`] (dollars per year).

    So now we have that [` y=[$m](x-x_0)+y_0 `]. We can use [`(x_0,y_0)=([$x1],[$y1])`] to get [`[$ansPS]`].

b)    In the year [`[$yx3]`], [`x`] was [`[$x3]`]. To find the amount remaining in the fund, we substitute [`[$x3]`] in for [`x`] in the equation [` y=[$m](x-[$x1])+[$y1] `], and we have:

    [``
\begin{aligned}
  y &= [$m]([$x3]-[$x1]) +[$y1] \\
  y &= [$y3]
\end{aligned}
``]

    So in the year [`[$yx3]`], the fund had [`[$cy3]`] remaining in it.

c)    The fund will be empty when it has [`[$cz]`] left in it. That is, [`y`] will equal [`0`]. To find how many years until this happens, we substitute [`[$y4]`] in for [`y`] in the equation [` y=[$m](x-[$x1])+[$y1] `], and we have:

    [``
\begin{aligned}
  y &= [$m](x-[$x1]) +[$y1] \\
  [$y4] &= [$m](x-[$x1]) +[$y1] \\
  [@$y4-$y1@] &= [$m](x-[$x1])\\
  \frac{[@$y4-$y1@]}{[$m]} &= x - [$x1]\\
  [@($y4-$y1)/$m@] &= x - [$x1]\\
  [@($y4-$y1)/$m+$x1@] &= x\\
  [$x4] &= x
\end{aligned}
``]

    Approximately [`[$x4]`] years after 2000, in the year [$yx4], the fund will be depleted.

END_PGML_SOLUTION


ENDDOCUMENT();
