# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/20/17
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','application','subtract','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextCurrency.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(p=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$original = random(50,300,10);
$pNum = random(5,8,0.1);
$p = $pNum/100;
$withTax = $original*(1+$p);

$ansString = "p+$p*p=$withTax";
$ans = Compute("$ansString");

Context("Currency");
$withTaxC = Currency($withTax);


##############################################

TEXT(beginproblem());

BEGIN_PGML

A bicycle for sale costs [`[$withTaxC]`], which includes [`[$pNum]%`] sales tax. What was the cost before sales tax?

Assume the bicycle's price before sales tax is [`p`] dollars. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

Assume the bicycle's original price before tax is [`p`] dollars. With a tax rate of [`[$pNum]%`],  the tax is [`[$p]p`] dollars. The bicycle's new price is its original price plus tax, so we can write the equation:

    [``p+[$p]p=[$withTax]``]

END_PGML_SOLUTION

ENDDOCUMENT();
