# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/19/17
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','application','subtract','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextCurrency.pl",
  "PCCmacros.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(s=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$lastYear = random(30000,60000,1000);
$pNum = random(1,3,0.1);
$p = $pNum/100;
$thisYear = $lastYear*(1+$p);

$ansString = "s+$p*s=$thisYear";
$ans = Compute("$ansString");

Context("Currency");
$thisYearC = Currency($thisYear);

$sex = list_random('male','female');
$pronoun = 'he' if ($sex eq 'male');
$pronoun = 'she' if ($sex eq 'female');
$object = 'him' if ($sex eq 'male');
$object = 'her' if ($sex eq 'female');
$possessive = 'his' if ($sex eq 'male');
$possessive = 'her' if ($sex eq 'female');
$Possessive = 'His' if ($sex eq 'male');
$Possessive = 'Her' if ($sex eq 'female');
$name = RandomName(sex=>$sex);

##############################################

TEXT(beginproblem());

BEGIN_PGML

[$name]'s annual salary as a radiography technician is [`[$thisYearC]`]. [$Possessive] salary increased by [`[$pNum]%`] from last year. What was [$possessive] salary last year?

Assume [$possessive] salary last year was [`s`] dollars. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

Assume [$possessive] salary last year was [`s`] dollars. Since [$possesive] salary increased by [`[$pNum]%`] from last year, the increase was [`[$p]s`] dollars. This year's salary is the sum of last year's salary and the increase, so we can write the equation:

    [``s+[$p]s=[$thisYear]``]

END_PGML_SOLUTION

ENDDOCUMENT();
