# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/22/09


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Partial derivatives)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(2varapp)
## MLTleader(1)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('14.1')
## Problem1('13')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('6')
## Section2('14.1')
## Problem2('14')
## TitleText3('Applied Calculus')
## AuthorText3('Hughes-Hallett')
## EditionText3('4')
## Section3('9.Review')
## Problem3('17')
## KEYWORDS('partial', 'derivative', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

Context("Numeric");
$showPartialCorrectAnswers = 1;

@wt = ( [ 31, 25, 19, 13,  7,   1,  -5, -11 ],
        [ 27, 21, 15,  9,  3,  -4, -10, -16 ],
        [ 25, 19, 13,  6,  0,  -7, -13, -19 ],
        [ 24, 17, 11,  4, -2,  -9, -15, -22 ],
        [ 23, 16,  9,  3, -4, -11, -17, -24 ] );

$t0ind = random(0,7,1);
$w0ind = random(0,4,1);
$w0 = 5*($w0ind + 1);
$t0 = 5*(7-$t0ind);
$t1ind = random(0,7,1);
$w1ind = random(0,4,1);
$w1 = 5*($w1ind + 1);
$t1 = 5*(7-$t1ind);

if ( $w0ind == 0 ) {
    $fw = Compute( ($wt[1]->[$t0ind] - $wt[0]->[$t0ind])/5 );
    $fwnum = "($wt[1]->[$t0ind] - $wt[0]->[$t0ind])";
} elsif ( $w0ind == 4 ) {
    $fw = Compute( ($wt[4]->[$t0ind] - $wt[3]->[$t0ind])/5 );
    $fwnum = "($wt[4]->[$t0ind] - $wt[3]->[$t0ind])";
} else {
    $fw = Compute( ($wt[$w0ind+1]->[$t0ind] - $wt[$w0ind-1]->[$t0ind])/10 );
    $fw1 = Compute( ($wt[$w0ind+1]->[$t0ind] - $wt[$w0ind]->[$t0ind])/5 );
    $fw2 = Compute( ($wt[$w0ind]->[$t0ind] - $wt[$w0ind-1]->[$t0ind])/5 );
    $fwnum = "$wt[$w0ind+1]->[$t0ind] - $wt[$w0ind]->[$t0ind]";
}

if ( $t1ind == 0 ) {
    $ft = Compute( ($wt[$w1ind]->[0] - $wt[$w1ind]->[1])/5 );
    $ftnum = "($wt[$w1ind]->[0] - $wt[$w1ind]->[1])";
} elsif ( $t1ind == 7 ) {
    $ft = Compute( ($wt[$w1ind]->[6] - $wt[$w1ind]->[7])/5 );
    $ftnum = "($wt[$w1ind]->[6] - $wt[$w1ind]->[7])";
} else {
    $ft = Compute( ($wt[$w1ind]->[$t1ind-1] - $wt[$w1ind]->[$t1ind+1])/10 );
    $ft1 = Compute( ($wt[$w1ind]->[$t1ind-1] - $wt[$w1ind]->[$t1ind])/5 );
    $ft2 = Compute( ($wt[$w1ind]->[$t1ind] - $wt[$w1ind]->[$t1ind+1])/5 );
    $ftnum = "($wt[$w1ind]->[$t1ind-1] - $wt[$w1ind]->[$t1ind])";
}

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

The following questions refer to the table shown below, giving
the temperature adjusted for wind-chill, \( C \), in
\({}^\circ\) F, as a function \( f(w,T) \) of the wind
speed, \( w \), in mph, and the temperature, \( T \), in
\({}^\circ\) F. The temperature adjusted for wind-chill tells
you how cold it feels, as a result of the combination
of wind and temperature.

$PAR
$BCENTER
\{ begintable(9) \}
\{ row( "", "\(T=35\)", "\(T=30\)", "\(T=25\)", "\(T=20\)", "\(T=15\)",
        "\(T=10\)", "\(T=5\)", "\(T=0\)" ) \}
\{ row( "\(w=5\)",  @{$wt[0]} ) \}
\{ row( "\(w=10\)", @{$wt[1]} ) \}
\{ row( "\(w=15\)", @{$wt[2]} ) \}
\{ row( "\(w=20\)", @{$wt[3]} ) \}
\{ row( "\(w=25\)", @{$wt[4]} ) \}
\{ endtable() \}
$ECENTER
$PAR

${BBOLD}(a)$EBOLD
Estimate \( f_w($w0, $t0) \):$BR
\( f_w($w0, $t0) \approx \) \{ ans_rule(35) \}
$BR
${BITALIC}Be sure you can state in practical terms what your answer
means.$EITALIC

$PAR
${BBOLD}(b)$EBOLD
Estimate \( f_T($w1, $t1) \):$BR
\( f_T($w1, $t1) \approx \) \{ ans_rule(35) \}
$BR
${BITALIC}Be sure you can state in practical terms what your answer
means.$EITALIC

END_TEXT
Context()->normalStrings;

if ( $w0ind == 0 || $w0ind == 4 ) {
    ANS( $fw->cmp() );
    if ( $w0ind == 4 ) {
	$fwmsg = "In this case, however, we have no data for values of " .
	  "\(w > $w0\), and so may instead take the difference quotient " .
	    "\[\frac{f(w_0,t_0) - f(w_0-h,t_0)}{h}.\]";
    } else {
	$fwmsg = "";
    }
    $fwval = $fw;
    $fwaltans = "";
} else {
    ANS( $fw1->cmp( checker=>sub {
		       my ( $correct, $student, $ansHash ) = @_;
		       return $correct == $student || $fw == $student ||
			 $fw2 == $student;
		   } ) );
    $fwmsg = "";
    $fwval = $fw1;
    $fwaltans = "Note that we could also approximate the partial " .
        "derivative by taking a difference quotient with the point to " .
	"the left of $w0 (that is, taking \(f_w($w0,$t0)\approx " .
	"\frac{f($w0,$t0) - f($w0-5),$t0}{5}\), or by taking a " .
	"central difference, \(f_w($w0,$t0)\approx \frac{f($w0+5,$t0) - " .
	"f($w0-5,$t0)}{10}\).";
}
if ( $t1ind == 0 || $t1ind == 7 ) {
    ANS( $ft->cmp() );
    if ( $t1ind == 7 ) {
	$ftmsg = "In this case, however, we have no data for values of " .
	  "\(t > $t1\), and so may instead take the difference quotient " .
	    "\[\frac{f(t_0,t_0) - f(w_0,t_0-h)}{h}.\]";
    } else {
	$ftmsg = "";
    }
    $ftval = $ft;
    $ftaltans = "";
} else {
    ANS( $ft1->cmp( checker=>sub {
		       my ( $correct, $student, $ansHash ) = @_;
		       return $correct == $student || $ft == $student ||
			 $ft2 == $student;
		   } ) );
    $ftmsg = "";
    $ftval = $ft1;
    $ftaltans = "Note that we could also approximate the partial " .
        "derivative by taking a difference quotient with the point to " .
	"the left of $t1 (that is, taking \(f_t($w1,$t1)\approx " .
	"\frac{f($w1,$t1) - f($w1,$t1-5)}{5}\), or by taking a " .
	"central difference, \(f_t($w1,$t1)\approx \frac{f($w1,$t1+5) - " .
	"f($w1,$t1-5)}{10}\).";
}

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
For \( f_w($w0, $t0) \)
we know
\[ f_w(w_0, t_0) \approx \frac{f(w_0 + h,t_0) - f(w_0,t_0)}{h}. \]
$fwmsg
Thus, we have
\[ f_w($w0, $t0) \approx \frac{$fwnum}{5} = $fwval. \]
$PAR
$fwaltans

This means that when the wind speed is $w0 mph
and the true temperature is $t0 \({}^\circ\) F,
as the wind speed increases from $w0 mph by 1 mph
we feel an approximately \($fwval^\circ\) F
drop in temperature.  This rate is negative because the temperature
you feel drops as the wind speed increases.

$PAR

${BBOLD}(b)$EBOLD
For \( f_t($w1, $t1) \)
we know
\[ f_t(w_0, t_0) \approx \frac{f(w_0,t_0 + h) - f(w_0,t_0)}{h}. \]
$ftmsg
Thus, we have
\[ f_t($w1, $t1) \approx \frac{$ftnum}{5} = $ftval. \]
$PAR
$ftaltans

This means that when the wind speed is $w1 mph
and the true temperature is $t1 \({}^\circ\) F,
the apparent temperature increases by approximately
\($ftval^\circ\) F for an increase of \(1^\circ\) F
in the true temperature.  This rate is positive because
the temperature you feel increases as the true temperature
increases.

END_SOLUTION
Context()->normalStrings;



ENDDOCUMENT();
