# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/22/09


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Partial derivatives)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(2varapp)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('14.1')
## Section2('14.1')
## Section3('14.1')
## Problem1('19')
## Problem2('21')
## Problem3('29')
## KEYWORDS('partial', 'derivative', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

Context("Numeric");
$showPartialCorrectAnswers = 1;

@ftab = ( [ "", "\(t=14\)", "\(t=16\)", "\(t=18\)", "\(t=20\)", "\(t=22\)", 
	    "\(t=24\)" ],
	  [ "\(c=0\)",  100, 100, 100,  99,  97,  95 ],
	  [ "\(c=2\)",  100,  99,  98,  97,  95,  92 ],
	  [ "\(c=6\)",   96,  95,  93,  90,  86,  80 ],
	  [ "\(c=15\)",  96,  93,  82,  70,  58,  36 ] );
$whicht = random(2,5,1);
$whichc = random(2,3,1);
@t0l = ("", 14, 16, 18, 20, 22, 24 );
@c0l = ("", 0, 2, 6, 15 );
$t0 = $t0l[$whicht];
$c0 = $c0l[$whichc];
$f0   = $ftab[$whichc]->[$whicht];
$f0pt = $ftab[$whichc]->[$whicht+1];  $t0p = $t0l[$whicht+1];
$f0mt = $ftab[$whichc]->[$whicht-1];  $t0m = $t0l[$whicht-1];
$f0pc = $ftab[$whichc+1]->[$whicht];  $c0p = $c0l[$whichc+1];
$f0mc = $ftab[$whichc-1]->[$whicht];  $c0m = $c0l[$whichc-1];
$ft = Compute( "($f0pt - $f0mt)/($t0p-$t0m)" );
# $ft = Compute( "($f0pt - $f0mt)/(4)" );
$fc = Compute( "($f0pc - $f0mc)/($c0p-$c0m)" );
$ftp = Compute( "($f0pt - $f0)/($t0p-$t0)" );
$ftm = Compute( "($f0 - $f0mt)/($t0-$t0m)" );
# $ftm = Compute( "($f0 - $f0mt)/(4)" );
$fcp = Compute( "($f0pc - $f0)/($c0p-$c0)" );
$fcm = Compute( "($f0 - $f0mc)/($c0-$c0m)" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

An experiment to measure the toxicity of formaldehyde yielded the data
in the table below.  The values show the percent, \( P=f(t,c) \),
of rats surviving an exposure to formaldehyde at a concentration of
\( c \) (in parts per million, ppm) after \( t \) months.
$PAR
$BCENTER
\{ begintable(7) \}
\{ row( @{$ftab[0]} ) \}
\{ row( @{$ftab[1]} ) \}
\{ row( @{$ftab[2]} ) \}
\{ row( @{$ftab[3]} ) \}
\{ row( @{$ftab[4]} ) \}
\{ endtable() \}
$ECENTER
$PAR

${BBOLD}(a)$EBOLD
Estimate \( f_t($t0,$c0) \): $BR
\( f_t($t0, $c0)\approx \) \{ ans_rule(35) \}

$PAR
${BBOLD}(b)$EBOLD
Estimate \( f_c($t0,$c0) \): $BR
\( f_c($t0, $c0)\approx \) \{ ans_rule(35) \}

$PAR
${BITALIC}(Be sure that you can give the practical meaning of these
two values in terms of formaldehyde toxicity.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS($ft->cmp( checker=>sub { 
		   my ($correct, $student, $ansHash) = @_;
		   return $student == $correct ||
		       $student == $ftp || $student == $ftm } ) );
ANS($fc->cmp( checker=>sub { 
		   my ($correct, $student, $ansHash) = @_;
		   return $student == $correct ||
		       $student == $fcp || $student == $fcm } ) );

$ftround = round(100*$ft)/100;
$fcround = abs(round(100*$fc)/100);

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

We have \(f($t0,$c0) = $f0\).

$PAR
${BBOLD}(a)$EBOLD
Thus, to estimate \( f_t($t0,$c0) \), we can take
\[ f_t($t0,$c0) \approx \frac{f($t0+2,$c0) - f($t0-2,$c0)}{4} =
    \frac{$f0pt - $f0mt}{4} = $ft. \]
Alternately, we could take either of the two one-sided difference
quotients,
\[ f_t($t0,$c0) \approx \frac{f($t0+2,$c0) - f($t0,$c0)}{2} = $ftp \]
or
\[ f_t($t0,$c0) \approx \frac{f($t0,$c0) - f($t0-2,$c0)}{2} = $ftm. \]
In each case, the units of \(f_t\) are percent per month.  This says
$t0 months after rats are exposed to a formaldehyde concentration
of $c0 ppm, the percent of rats surviving is decreasing at a rate of
about $ftround per month. In other words, during the $t0 month,
around an additional $ftround percent of the rats will die.

$PAR
${BBOLD}(b)$EBOLD
Proceeding similarly to find \(f_c($t0,$c0)\), we have
\[ f_c($t0,$c0) \approx \frac{f($t0,$c0+2) - f($t0,$c0-2)}{4} =
    \frac{$f0pc - $f0mc}{4} = $fc, \]
or, with one-sided difference quotients,
\[ f_c($t0,$c0) \approx \frac{f($t0,$c0+2) - f($t0,$c0)}{2} = $fcp \]
or
\[ f_c($t0,$c0) \approx \frac{f($t0,$c0) - f($t0,$c0-2)}{2} = $fcm. \]
These indicate that if the original concentration increases from $c0
by 1 ppm, the percent surviving after $t0 months decreases by about
$fcround.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
