# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/22/09


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Partial derivatives)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('14.1')
## Section2('14.1')
## Section3('14.1')
## Problem1('26')
## Problem2('31')
## Problem3('37')
## KEYWORDS('partial', 'derivative', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

Context("Numeric");
$showPartialCorrectAnswers = 1;

$gr = init_graph( 0, 0, 40, 600, axes=>[0,0], size=>[350,350], grid=>[16,24] );
$gr->lb('reset');

foreach my $i ( ( 10, 20, 30 ) ) {
    $gr->lb( new Label( $i, 0, $i, "black", "center", "bottom" ) );
}
$gr->lb( new Label( 39, 10, "T", "black", "right", "bottom" ) );
foreach my $i ( (100, 200, 300, 400, 500 ) ) {
    $gr->lb( new Label( 0, $i, $i, "black", "left", "middle" ) );
}
$gr->lb( new Label( 1, 590, "H", "black", "left", "top" ) );

$gr->lb( new Label( 5, 100, "H(T,0.1)", "black", "left", "top" ) );
$gr->lb( new Label( 5, 225, "H(T,0.2)", "black", "left", "top" ) );
$gr->lb( new Label( 5, 375, "H(T,0.3)", "black", "left", "bottom" ) );
$gr->lb( new Label( 5, 525, "H(T,0.4)", "black", "left", "bottom" ) );

add_functions( $gr,
	       "585 - 15.75*x + 0.2*x^2 for x in <0,40> using color:blue " .
	           "and weight:2",
	       "430 - 11.25*x + 0.1375*x^2 for x in <0,40> using color:blue " .
	           "and weight:2",
	       "290 - 6*x + 0.05*x^2 for x in <0,40> using color:blue and " .
	           "weight:2",
	       "145 - 3.125*x + 0.03125*x^2 for x in <0,40> using " .
	           "color:blue and weight:2" );

## the values on these functions are
%hval = ( 0  => { "0.1" => 145,     "0.2" => 290, "0.3" => 430,
		  "0.4" => 585 },
	  10 => { "0.1" => 116.875, "0.2" => 235, "0.3" => 331.25,
		  "0.4" => 447.5 },
	  20 => { "0.1" => 95,      "0.2" => 190, "0.3" => 260,
		  "0.4" => 350 },
	  30 => { "0.1" => 79.375,  "0.2" => 155, "0.3" => 216.25,
		  "0.4" => 292.5 },
	  40 => { "0.1" => 70,      "0.2" => 130, "0.3" => 200,
		  "0.4" => 275 } );
## reasonable approximations are
%happ = ( 0  => { "0.1" => 150, "0.2" => 290, "0.3" => 430, "0.4" => 590 },
	  10 => { "0.1" => 120, "0.2" => 240, "0.3" => 330, "0.4" => 450 },
	  20 => { "0.1" => 100, "0.2" => 190, "0.3" => 260, "0.4" => 350 },
	  30 => { "0.1" => 80,  "0.2" => 160, "0.3" => 220, "0.4" => 300 },
	  40 => { "0.1" => 70,  "0.2" => 130, "0.3" => 200, "0.4" => 280 } );

## some values at which we ask for approximations
$w0 = random(0.2,0.3,0.1);
## the first derivative is
$w0a = $w0 + 0.1;
$ht0 = Compute( "($happ{20}->{$w0} - $happ{10}->{$w0})/10" );
$ht0a = Compute( "($happ{10}->{$w0} - $happ{0}->{$w0})/10" );
$ht0b = Compute( "($happ{20}->{$w0} - $happ{0}->{$w0})/20" );


$t1 = random(10,20,10);         $t1a = $t1 + 10;    $t1b = $t1 - 10;
$t2 = random($t1+10,30,10);     $t2a = $t2 + 10;    $t2b = $t2 - 10;
$w1 = random(0.2,0.2,0.1);      $w1a = $w1 + 0.1;   $w1b = $w1 - 0.1;
$w2 = random($w1+0.1,0.3,0.1);  $w2a = $w2 + 0.1;   $w2b = $w2 + 0.1;

$ht11 = Compute( "($happ{$t1a}->{$w1} - $happ{$t1}->{$w1})/10" );
$ht11a = Compute( "($happ{$t1}->{$w1} - $happ{$t1b}->{$w1})/10" );
$ht11b = Compute( "($happ{$t1a}->{$w1} - $happ{$t1b}->{$w1})/20" );
$ht12 = Compute( "($happ{$t1a}->{$w2} - $happ{$t1}->{$w2})/10" );
$ht12a = Compute( "($happ{$t1}->{$w2} - $happ{$t1b}->{$w2})/10" );
$ht12b = Compute( "($happ{$t1a}->{$w2} - $happ{$t1b}->{$w2})/20" );
$ht21 = Compute( "($happ{$t2a}->{$w1} - $happ{$t2}->{$w1})/10" );
$ht21a = Compute( "($happ{$t2}->{$w1} - $happ{$t2b}->{$w1})/10" );
$ht21b = Compute( "($happ{$t2a}->{$w1} - $happ{$t2b}->{$w1})/20" );
$ht22 = Compute( "($happ{$t2a}->{$w2} - $happ{$t2}->{$w2})/10" );
$ht22a = Compute( "($happ{$t2a}->{$w2} - $happ{$t2}->{$w2})/10" );
$ht22b = Compute( "($happ{$t2a}->{$w2} - $happ{$t2b}->{$w2})/20" );

$hw11 = Compute( "($happ{$t1}->{$w1a} - $happ{$t1}->{$w1})/0.1" );
$hw11a = Compute( "($happ{$t1}->{$w1} - $happ{$t1}->{$w1b})/0.1" );
$hw11b = Compute( "($happ{$t1}->{$w1a} - $happ{$t1}->{$w1b})/0.2" );
$hw12 = Compute( "($happ{$t1}->{$w2a} - $happ{$t1}->{$w2})/0.1" );
$hw12a = Compute( "($happ{$t1}->{$w2} - $happ{$t1}->{$w2b})/0.1" );
$hw12b = Compute( "($happ{$t1}->{$w2a} - $happ{$t1}->{$w2b})/0.2" );
$hw21 = Compute( "($happ{$t2}->{$w1a} - $happ{$t2}->{$w1})/0.1" );
$hw21a = Compute( "($happ{$t2}->{$w1} - $happ{$t2}->{$w1b})/0.1" );
$hw21b = Compute( "($happ{$t2}->{$w1a} - $happ{$t2}->{$w1b})/0.2" );
$hw22 = Compute( "($happ{$t2}->{$w2a} - $happ{$t2}->{$w2})/0.1" );
$hw22a = Compute( "($happ{$t2}->{$w2} - $happ{$t2}->{$w2b})/0.1" );
$hw22b = Compute( "($happ{$t2}->{$w2a} - $happ{$t2}->{$w2b})/0.2" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

An airport can be cleared of fog by heating the air.  The amount of
heat required depends on the air temperature and the wetness of the
fog.  The figure below shows the heat \( H(T,w) \) required (in
calories per cubic meter of fog) as a function of the temperature
\( T \) (in degrees Celsius) and the water content \( w \) (in grams
per cubic meter of fog).  Note that this figure is not a contour
diagram, but shows cross-sections of \( H \) with \( w \) fixed at
\( 0.1 \), \( 0.2 \), \( 0.3 \), and \( 0.4 \).
$PAR
$BCENTER
\{ image( insertGraph( $gr ), tex_size=>250, height=>350, width=>350,
	  extra_html_tags=>'alt="graph showing H(T,0.1), H(T,0.2), ' .
	  'H(T,0.3) and H(T,0.4), which are decresaing functions ' .
	  'of T, each at a higher value of H."' ) \}
$ECENTER
$PAR
${BBOLD}(a)$EBOLD
Estimate \( H_T(10, $w0) \):$BR
\( H_T(10,$w0) \approx \) \{ ans_rule(35) \}
$BR
${BITALIC}(Be sure you can interpret this
partial derivative in practical terms.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
Make a table of values for \(H(T,w)\) from the figure, and use it
to estimate \(H_T(T,w)\) for each of the following:
$BR
\(T = $t1, w = $w1\) : \(H_T(T,w) \approx \) \{ ans_rule(35) \} $BR
\(T = $t2, w = $w1\) : \(H_T(T,w) \approx \) \{ ans_rule(35) \} $BR
\(T = $t1, w = $w2\) : \(H_T(T,w) \approx \) \{ ans_rule(35) \} $BR
\(T = $t2, w = $w2\) : \(H_T(T,w) \approx \) \{ ans_rule(35) \}

$PAR
${BBOLD}(c)$EBOLD
Repeat ${BBOLD}(b)$EBOLD to find \(H_w(T,w)\) for each of the following:
$BR
\(T = $t1, w = $w1\) : \(H_w(T,w) \approx \) \{ ans_rule(35) \} $BR
\(T = $t2, w = $w1\) : \(H_w(T,w) \approx \) \{ ans_rule(35) \} $BR
\(T = $t1, w = $w2\) : \(H_w(T,w) \approx \) \{ ans_rule(35) \} $BR
\(T = $t2, w = $w2\) : \(H_w(T,w) \approx \) \{ ans_rule(35) \}
$BR
${BITALIC}(Be sure you can interpret this
partial derivative in practical terms.)$EITALIC

END_TEXT
Context()->normalStrings;

Context()->flags->set( tolerance=>.1 );
ANS($ht0->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $ht0a || $student == $ht0b;
} ) );
ANS($ht11->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $ht11a || $student == $ht11b;
} ) );
ANS($ht21->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $ht21a || $student == $ht21b;
} ) );
ANS($ht12->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $ht12a || $student == $ht12b;
} ) );
ANS($ht22->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $ht22a || $student == $ht22b;
} ) );
# Context()->flags->set( tolType=>'absolute', tolerance=>50 );
ANS($hw11->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $hw11a || $student == $hw11b;
} ) );
ANS($hw21->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $hw21a || $student == $hw21b;
} ) );
ANS($hw12->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $hw12a || $student == $hw12b;
} ) );
ANS($hw22->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student == $correct || $student == $hw22a || $student == $hw22b;
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The quantity \(H_T(10,$w0)\) is approximated by an
difference quotient.
The first partial derivative with respect to \(T\) is approximated  by
\[ H_T(10,$w0) \approx { H(10+\Delta T, $w0) - H(10, $w0)\over \Delta T},\]
for small \(\Delta T\).  We are free to choose \(\Delta T\) to find
values from the figure we are given.  If we take \(\Delta T = 10\),
we get the approximation
\[
H_T(10,$w0) \approx { H(20, $w0) - H(10, $w0)\over 10} \approx
{$happ{20}->{$w0} - $happ{10}->{$w0}\over 10} = $ht0.
\]

$PAR
(Note that you may get a different answer if you read different values
from the graph.)  The geometric meaning of the partial derivative
\(H_T(10,$w0)\) that we just approximated is the slope of the
curve shown in the figure in the text corresponding to \(w = $w0\) at the point
where \(T= 10\).
$PAR
In practical terms, we have found that for fog at \(10^\circ\) C containing
0.1 g water per m\({}^3\) of fog, a \(1^\circ\) C increase in temperature
will reduce the heat requirement for dissipating the fog by about
1~calories per cubic meter of fog.

$PAR
${BBOLD}(b)$EBOLD
We can approximate values of \(H(T,w)\) from the figure, giving the
table below:
$PAR
$BCENTER
\{ begintable(5) \}
\{ row( ("\(w =\)", "0.1", "0.2", "0.3", "0.4") ) \}
\{ row( ("\(T = 10\)", $happ{10}->{0.1}, $happ{10}->{0.2}, $happ{10}->{0.3},
	 $happ{10}->{0.4}) ) \}
\{ row( ("\(T = 20\)", $happ{20}->{0.1}, $happ{20}->{0.2}, $happ{20}->{0.3},
	 $happ{20}->{0.4}) ) \}
\{ row( ("\(T = 30\)", $happ{30}->{0.1}, $happ{30}->{0.2}, $happ{30}->{0.3},
	 $happ{30}->{0.4}) ) \}
\{ row( ("\(T = 40\)", $happ{40}->{0.1}, $happ{40}->{0.2}, $happ{40}->{0.3},
	 $happ{40}->{0.4}) ) \}
\{ endtable() \}
$ECENTER
$PAR
This lets us approximate, using a difference quotient, each of the
derivatives indicated.  For \( H_T($t1,$w1) \), we have
\[ H_T($t1,$w1) \approx \frac{($happ{$t1a}->{$w1} - $happ{$t1}->{$w1})}{10}
    = $ht11. \]
(We could, of course, use a left or central difference instead.)
Similarly, we have:
\[ H_T($t2,$w1) \approx $ht21, H_T($t1,w2) \approx $ht12,
   H_T($t2,$w2) \approx $ht22. \]

$PAR
${BBOLD}(c)$EBOLD
Using our table of values, we can similarly calculate the partial derivatives
with respect to \(w\).  For \( H_w($t1,$w1) \), we have
\[ H_w($t1,$w1) \approx \frac{($happ{$t1}->{$w1a} - $happ{$t1}->{$w1})}{0.1}
    = $hw11. \]
(We could, of course, use a left or central difference instead.)
Similarly, we have:
\[ H_w($t2,$w1) \approx $hw21, H_w($t1,w2) \approx $hw12,
   H_w($t2,$w2) \approx $hw22. \]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
