# Original: /Library/ASU-topics/setCalculus/stef/stef4_3p3.pg

## DESCRIPTION
## Calculus: How derivatives affect the shape of a graph
## ENDDESCRIPTION

## KEYWORDS('calculus', 'derivatives' 'curve sketching')

## DBsubject('Calculus')
## DBchapter('Applications of Differentiation')
## DBsection('How Derivatives Affect the Shape of a Graph')
## Date('9/5/2012')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('AUCI')
## EditionText1('1')
## AuthorText1('Petrillo')
## Section1('4.1')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "contextLimitedPolynomial.pl",
  "contextPolynomialFactors.pl",
  "contextLimitedPowers.pl",
  "PGgraphmacros.pl",
  "unionTables.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
$refreshCachedImages=1;

Context("PolynomialFactors-Strict");
Context()->flags->set(singleFactors=>0);
LimitedPowers::OnlyIntegers(
minPower => 0, maxPower => 4,
message => "the numerator is a quartic",
);
$contextp = Context()->copy;

Context("Numeric");
Context()->variables->add(u=>"Real");
Context()->reduction->set('(-x)+y'=>0,'(-x)-y'=>0);

# If you keep the two x-intercepts to the left of the two vertical asymptotes
# then you get two spots where the tangent line is horizontal.
$a = Compute(random(-9,-5));
$b = $a + Compute(random(1,3));
$c = $b + Compute(random(1,3));
if ($c == 0) {$c++;};
$d = $c + Compute(random(1,3));
if ($d == 0) {$d++;};

$f = Formula("(x^2-($a+$b)*x+$a*$b)/(x^2-($c+$d)*x+$c*$d)")->reduce;
$derf = $f->D('x');

$yintercept = $f->eval(x=>"0");
$xintercepts = List($a,$b);
$verticalasymptotes = List($c,$d);

$coeff1 = Real("$a+$b-$c-$d");
$coeff2 = Real("2($c*$d-$a*$b)");
$coeff3 = Real("$a*$b*($c+$d)-$c*$d*($a+$b)");
$derivativenumerator = Formula($contextp,"$coeff1*x^2+$coeff2*x+$coeff3");


$root1 = Compute("(-$coeff2-sqrt(($coeff2)^2-4*$coeff1*$coeff3))/(2*$coeff1)");
$root2 = Compute("(-$coeff2+sqrt(($coeff2)^2-4*$coeff1*$coeff3))/(2*$coeff1)");

$localextrema1 = $f->eval(x=>"$root1");
$localextrema2 = $f->eval(x=>"$root2");


##
## The numerator of the derivative is a quartic equation that can be reduced 
## to a quadratic with the following coefficients
##

$horizontaltangent = List($root1,$root2) ;


################################################################################
##  Generate the graphs
##
($xmin,$xmax) = (-10,10);
$ymax = 20;
if(abs($localextrema1)+5 > $ymax) {$ymax = int(abs($localextrema1)+6)}; 
if(abs($localextrema2)+5 > $ymax) {$ymax = int(abs($localextrema2)+6)}; 
$ymin = -$ymax;

$g[0] = Formula("((x-$a)(x-$b))/((x-$c)(x-$d))");
$gr[0] = init_graph($xmin,$ymin,$xmax,$ymax,'axes'=>[0,0],grid=>[10,10],);
$f1_graph[0] = new Fun($g[0]->perlFunction,$gr[0]);
$f1_graph[0]->steps(25); 
$f1_graph[0]->color("red"); 
$f1_graph[0]->weight(2);
$f1_graph[0]->domain(-10,$c-0.01);
$f2_graph[0] = new Fun($g[0]->perlFunction,$gr[0]);
$f2_graph[0]->steps(25); 
$f2_graph[0]->color("red"); 
$f2_graph[0]->weight(2);
$f2_graph[0]->domain($c+0.01,$d-0.01);
$f3_graph[0] = new Fun($g[0]->perlFunction,$gr[0]);
$f3_graph[0]->steps(25); 
$f3_graph[0]->color("red"); 
$f3_graph[0]->weight(2);
$f3_graph[0]->domain($d+0.01,10);

$g[1] = Formula("((x-$a)(x-$c))/((x-$d)(x-$b))");
$gr[1] = init_graph($xmin,$ymin,$xmax,$ymax,'axes'=>[0,0],grid=>[10,10],);
$f1_graph[1] = new Fun($g[1]->perlFunction,$gr[1]);
$f1_graph[1]->steps(25); 
$f1_graph[1]->color("red"); 
$f1_graph[1]->weight(2);
$f1_graph[1]->domain(-10,$b-0.01);
$f2_graph[1] = new Fun($g[1]->perlFunction,$gr[1]);
$f2_graph[1]->steps(25); 
$f2_graph[1]->color("red"); 
$f2_graph[1]->weight(2);
$f2_graph[1]->domain($b+0.01,$d-0.01);
$f3_graph[1] = new Fun($g[1]->perlFunction,$gr[1]);
$f3_graph[1]->steps(25); 
$f3_graph[1]->color("red"); 
$f3_graph[1]->weight(2);
$f3_graph[1]->domain($d+0.01,10);

$g[2] = Formula("((x-$a)(x-$d))/((x-$b)(x-$c))");
$gr[2] = init_graph($xmin,$ymin,$xmax,$ymax,'axes'=>[0,0],grid=>[10,10],);
$f1_graph[2] = new Fun($g[2]->perlFunction,$gr[2]);
$f1_graph[2]->steps(25); 
$f1_graph[2]->color("red"); 
$f1_graph[2]->weight(2);
$f1_graph[2]->domain(-10,$b-0.01);
$f2_graph[2] = new Fun($g[2]->perlFunction,$gr[2]);
$f2_graph[2]->steps(25); 
$f2_graph[2]->color("red"); 
$f2_graph[2]->weight(2);
$f2_graph[2]->domain($b+0.01,$c-0.01);
$f3_graph[2] = new Fun($g[2]->perlFunction,$gr[2]);
$f3_graph[2]->steps(25); 
$f3_graph[2]->color("red"); 
$f3_graph[2]->weight(2);
$f3_graph[2]->domain($c+0.01,10);

$g[3] = Formula("((x-$c)(x-$d))/((x-$a)(x-$b))");
$gr[3] = init_graph($xmin,$ymin,$xmax,$ymax,'axes'=>[0,0],grid=>[10,10],);
$f1_graph[3] = new Fun($g[3]->perlFunction,$gr[3]);
$f1_graph[3]->steps(25); 
$f1_graph[3]->color("red"); 
$f1_graph[3]->weight(2);
$f1_graph[3]->domain(-10,$a-0.01);
$f2_graph[3] = new Fun($g[3]->perlFunction,$gr[3]);
$f2_graph[3]->steps(25); 
$f2_graph[3]->color("red"); 
$f2_graph[3]->weight(2);
$f2_graph[3]->domain($a+0.01,$b-0.01);
$f3_graph[3] = new Fun($g[3]->perlFunction,$gr[3]);
$f3_graph[3]->steps(25); 
$f3_graph[3]->color("red"); 
$f3_graph[3]->weight(2);
$f3_graph[3]->domain($b+0.01,10);

for $i (0..3) {
 # plot_functions( $gr[$i], $f[$i]); 
  $fig[$i]=image(insertGraph($gr[$i]),
  width=>200,height=>200,tex_size=>450);
}

@perm = shuffle(4);
@fig = @fig[@perm];
@inv = invert(@perm);

@letter = ("A", "B", "C", "D");

$popup = PopUp(["?","A","B","C","D"], $letter[$inv[0]]);


Context()->texStrings;

BEGIN_TEXT
$BR
Let \(\displaystyle f(x) = $f\). 
$BR
$BR
(a) The ${BITALIC}y$EITALIC-intercept is ${BITALIC}y$EITALIC = \{ans_rule\}.
$BR
$BR
(b) The ${BITALIC}x$EITALIC-intercept(s) is/are ${BITALIC}x$EITALIC =  \{ans_rule\}
$BR
$BR
(c)  The vertical asymptote(s) is/are ${BITALIC}x$EITALIC = \{ans_rule\}
$BR
$BR
(d) To find the local extrema of ${BITALIC}f$EITALIC, we set \( f^\prime(x) = 0 \) and solve for ${BITALIC}x$EITALIC.  In this case,
$BR$BR
$SPACE $SPACE $SPACE \( f^\prime(x) = \)  \{ans_rule(40)\}, 
$BR
$BR
and we must find the zeros of the numerator.  In reduced form (expanded without parentheses), we must ultimately solve the equation 
$BR$BR
$SPACE $SPACE $SPACE  \{ans_rule(40)\} \( = 0\).
$BR$BR
Therefore, the local extrema are at ${BITALIC}x$EITALIC = \{ans_rule(40)\}. 
END_TEXT
BEGIN_TEXT
$BR
$BR
$BR
(e)  Use the information from parts (a)-(d) (and without a calculator!) to choose the correct graph of ${BITALIC}f$EITALIC. 
\{ $popup->menu() \}
$BR
$BR
$BCENTER
\{
BeginTable().
  AlignedRow([$fig[0],$fig[1]]).
  TableSpace(5,0).
  AlignedRow(["A","B"]).
  TableSpace(25,6).
  AlignedRow([$fig[2],$fig[3]]).
  TableSpace(5,0).
  AlignedRow(["C","D"]).
EndTable();
\}
$BR
(Click on a graph to enlarge it.)
$ECENTER
END_TEXT
Context()->normalStrings;

ANS($yintercept->cmp() );
ANS($xintercepts->cmp() );
ANS($verticalasymptotes->cmp() );
ANS($derf->cmp() );

ANS($derivativenumerator->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $derivativenumerator");
      return $correct == $student;
    }
) );

ANS($horizontaltangent->cmp() );
ANS( $popup->cmp() );
ENDDOCUMENT();
