# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 05/06/09


## DBsubject(Calculus - multivariable)
## DBchapter(Concepts for multivariable functions)
## DBsection(Traces, contours, and level sets)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('12.2')
## Section2('12.2')
## Section3('12.2')
## Problem1('23')
## Problem2('28')
## Problem3('38')
## KEYWORDS('graphing', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Functions of Several Variables')

Context("Numeric");
Context()->variables->add( y => 'Real', z => 'Real' );
$showPartialCorrectAnswers = 1;

## pick values for a and b
$b = random(2,3,1);
$a = random(2,3,1)*$b;
$signa = list_random(-1,1);
$sa = $signa*$a;

## and make b have the opposite sign
$signb = $signa*-1;
$sb = $signb*$b;

$c = random(1,5,1);

$v1 = ( $signa == 1 ) ? Compute("y") : Compute("x");
$v2 = ( $signa == 1 ) ? Compute("x") : Compute("y");
$v3 = Compute("z");

$eqn = Compute( "$sa*x^2 + $sb*y^2 + $c" )->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

By setting one variable constant, find a plane that intersects the
graph of \( z = $eqn \) in a:
$PAR
${BBOLD}(a)$EBOLD
Parabola opening upward: the plane \{ ans_rule(5) \} = \{ ans_rule(5) \}
$BR
${BITALIC}(Give your answer by specifying the variable in the
first answer blank and a value for it in the second.)$EITALIC
$PAR
${BBOLD}(b)$EBOLD
Parabola opening downward: the plane \{ ans_rule(5) \} = \{ ans_rule(5) \}
$BR
${BITALIC}(Give your answer by specifying the variable in the
first answer blank and a value for it in the second.)$EITALIC
$PAR
${BBOLD}(c)$EBOLD
Pair of intersecting straight lines: the plane
\{ ans_rule(5) \} = \{ ans_rule(20) \}
$BR
${BITALIC}(Give your answer by specifying the variable in the
first answer blank and a value for it in the second.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $v1->cmp() );  ANS( Compute(0)->cmp( checker=>sub { return 1; } ) );
ANS( $v2->cmp() );  ANS( Compute(0)->cmp( checker=>sub { return 1; } ) );
ANS( $v3->cmp() );  ANS( Compute($c)->cmp() );

$k1 = ( $signa == 1 ) ? $a : $b;
$k2 = ( $signa == 1 ) ? $b : $a;
$k3 = $a/$b;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
If \($v1\) is constant, we get \(z = $k1 $v1^2 + a\), for some
value of \(a\), so that the cross-sections are upward-opening
parabolas.

$PAR
${BBOLD}(b)$EBOLD
If \($v2\) is constant, we get \(z = -$k2 $v2^2 + a\), for some
value of \(a\), so that the cross-sections are downward-opening
parabolas.

$PAR
${BBOLD}(c)$EBOLD
If \(z\) is constant and equal to $c, we get \(y = \pm \sqrt{$k3} x\),
which are two intersecting straight lines.

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
