# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('graphing', 'multivariable', 'functions')
## Tagged by glr 05/06/09


## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Graphs and Level Surfaces')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('12.2')
## Problem1('23')

## Textbook tags
## HHChapter1('Functions of Several Variables')
## HHSection1('Graphs of Functions of Two Variables')

Context("Numeric");
Context()->variables->add( y => 'Real', z => 'Real' );
$showPartialCorrectAnswers = 1;

## pick values for a and b
$b = random(2,3,1);
$a = random(2,3,1)*$b;
$signa = list_random(-1,1);
$sa = $signa*$a;

## and make b have the opposite sign
$signb = $signa*-1;
$sb = $signb*$b;

$c = random(1,5,1);

$v1 = ( $signa == 1 ) ? Compute("y") : Compute("x");
$v2 = ( $signa == 1 ) ? Compute("x") : Compute("y");
$v3 = Compute("z");

$eqn = Compute( "$sa*x^2 + $sb*y^2 + $c" )->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

By setting one variable constant, find a plane that intersects the
graph of \( z = $eqn \) in a:
$PAR
${BBOLD}(a)$EBOLD
Parabola opening upward: the plane \{ ans_rule(5) \} = \{ ans_rule(5) \}
$BR
${BITALIC}(Give your answer by specifying the variable in the
first answer blank and a value for it in the second.)$EITALIC
$PAR
${BBOLD}(b)$EBOLD
Parabola opening downward: the plane \{ ans_rule(5) \} = \{ ans_rule(5) \}
$BR
${BITALIC}(Give your answer by specifying the variable in the
first answer blank and a value for it in the second.)$EITALIC
$PAR
${BBOLD}(c)$EBOLD
Pair of intersecting straight lines: the plane
\{ ans_rule(5) \} = \{ ans_rule(20) \}
$BR
${BITALIC}(Give your answer by specifying the variable in the
first answer blank and a value for it in the second.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $v1->cmp() );  ANS( Compute(0)->cmp( checker=>sub { return 1; } ) );
ANS( $v2->cmp() );  ANS( Compute(0)->cmp( checker=>sub { return 1; } ) );
ANS( $v3->cmp() );  ANS( Compute($c)->cmp() );

$k1 = ( $signa == 1 ) ? $a : $b;
$k2 = ( $signa == 1 ) ? $b : $a;
$k3 = $a/$b;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
If \($v1\) is constant, we get \(z = $k1 $v1^2 + a\), for some
value of \(a\), so that the cross-sections are upward-opening
parabolas.

$PAR
${BBOLD}(b)$EBOLD
If \($v2\) is constant, we get \(z = -$k2 $v2^2 + a\), for some
value of \(a\), so that the cross-sections are downward-opening
parabolas.

$PAR
${BBOLD}(c)$EBOLD
If \(z\) is constant and equal to $c, we get \(y = \pm \sqrt{$k3} x\),
which are two intersecting straight lines.

END_SOLUTION
Context()->normalStrings;

COMMENT('Uses MathObjects');

ENDDOCUMENT();
