# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al.,
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 12/07/09

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Applications - Newton's law of cooling)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(FO-AN-DE)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('11.5')
## Section2('11.5')
## Section3('11.5')
## Problem1('25')
## Problem2('27')
## Problem3('40')
## KEYWORDS('differential equation', 'calculus', 'exponential growth', 'exponential decay')

## Textbook tags
## HHChapter1('Differential Equations')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real', T=>'Real', k=>'Real' );
$showPartialCorrectAnswers = 1;

$hour1 = random(1,3,1);
$temp1 = random(66,72,2);
$hour2 = random(9,11,1);
$delta = $hour2 - $hour1;
$temp2 = random(53,60,1);
$outtemp = random(5,15,1);
$hour3 = random(6,10,1);
$del2 = 12 - $hour1 + $hour3;

$de = Compute( "k*($outtemp - T)" );

## solving this, we end up with
$k = Compute( "ln(($temp1 - $outtemp)/($temp2 - $outtemp))/$delta" );
$endtemp = Compute( "$outtemp + ($temp1 - $outtemp)*e^(-1*($k)*($del2))" );

$worry = PopUp( [ "?", "yes", "no" ], $endtemp < 32 ? "yes" : "no" );
$revision = PopUp( [ "?", "up", "down" ], "down" );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

At $hour1:00 pm one winter afternoon, there is a power
failure at your house in Wisconsin, and  your heat does not work
without electricity.  When the power goes out, it is
\( $temp1^{\circ}\)F in your
house.  At $hour2:00 pm, it is \( $temp2^{\circ} \)F
in the house, and you notice
that it is \( $outtemp^{\circ} \)F outside.

$PAR
${BBOLD}(a)$EBOLD
Assuming that the temperature, \( T \), in your home obeys Newton's Law
of Cooling, write the differential equation satisfied by \( T \).
$BR
\( \frac{dT}{dt} = \) \{ ans_rule(35) \}
$BR
${BITALIC}(Use ${BBOLD}k$EBOLD for any constant of proportionality in
your equation; your equation may involve ${BBOLD}T$EBOLD and the values
in the problem.$EITALIC

$PAR
${BBOLD}(b)$EBOLD
Solve the differential equation to estimate the temperature in the
house when you get up at $hour3:00 am the next morning.
$BR
Temperature = \{ ans_rule(35) \}
$BR
Should you worry about your water pipes freezing?
\{ $worry->menu() \}

$PAR
${BBOLD}(c)$EBOLD
Think about your equation in (a): what assumption did you make
about the temperature outside?  Given this (probably incorrect)
assumption, would you revise your estimate up or down?  (And why?)
$BR
Revise \{ $revision->menu() \}

END_TEXT
Context()->normalStrings;

ANS( $de->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    return $correct == $student || $correct == -1*$student;
} ) );

ANS( $endtemp->cmp() );
ANS( $worry->cmp() );
ANS( $revision->cmp() );

$t1mo = $temp1 - $outtemp;
$t2mo = $temp2 - $outtemp;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The differential equation is
\[ {dT\over dt} = -k(T - $outtemp),\]
where \($outtemp^{\circ}\)F is the outside temperature.

$PAR
${BBOLD}(b)$EBOLD
Integrating both sides yields
\[ \int{dT\over T - $outtemp} =-\int k\, dt.\]
Then \(\ln |T - $outtemp| = -kt +C\), so \(T= $outtemp +Be^{-kt}\).
We know that initially \(T=$temp1^{\circ}\)F. Thus,
letting \(t=0\) correspond to $hour1 pm, we get
\[ $temp1 = $outtemp C e^{0}, \quad\mbox{ so }\quad
   C = $t1mo. \]
Thus
\[ T= $outtemp + $t1mo e^{-kt}.\]
Since $hour2:00 pm corresponds to \(t=$delta\),
\[ $temp2 = $outtemp + $t1mo e^{-$delta k},\]
so
\[ $t2mo = $t1mo e^{-$delta k}, \]
and dividing by $t1mo and taking the natural log of both sides
we have
\[ k = -\frac{1}{$delta}\, \ln(\frac{$t2mo}{$t1mo}) \approx $k. \]

$PAR
At $hour3:00 the next morning (\(t=$del2\)) we have
\[
T \approx $outtemp + $t1mo e^{$del2(-k)} = $endtemp.
\]
so the pipes won't freeze.

$PAR
${BBOLD}(c)$EBOLD
We assumed that the temperature outside the house stayed
constant at \($outtemp^\circ\)F.
This is probably incorrect because the
temperature was most likely warmer during the day
and colder overnight.  Thus, when the
temperature in the house dropped from
\($temp1^\circ\)F to \($temp2^{\circ}\)F,
the outside temperature was probably
higher than \($outtemp^{\circ}\)F,
which changes our calculation of the value of
the constant \(k\).  The house temperature will most certainly be lower
than \($endtemp^{\circ}\)F
at $hour3 am, but not by much---not enough to freeze.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
