## DESCRIPTION
## Higher order ODEs: constant coefficient second order nonhomogeneous linear ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Variation of parameters)
## Date(09/01/2015)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Multivariable Mathematics')
## AuthorText1('Darin Stephenson')
## EditionText1('Version 3.0')
## Section1('7.2')
## Problem1('2')
## KEYWORDS('differential equations','second order linear ODE','constant coefficients','nonhomogeneous')

##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "answerHints.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

###########################################
#  The diff eq

$k = random(2,4,1);
$k2 = 2 * $k;
$kk = $k * $k;

# char poly r^2 + $k^2

$diffeq = "y^{\,\prime\prime} + $k2 y^{\,\prime} + $kk y = $k x"; # tex



###########################################
# Part 1
# Get rid of adaptive parameters which might mess with answer checkers.

Context("Numeric");
Context()->variables->are(
  'x' => ["Real", limits => [-1,1] ],
  'c1' => ["Real", limits => [2,4] ],
  'c2' => ["Real", limits => [2,4] ],
);
Context()->flags->set(
  tolType => 'absolute',
  tolerance => 0.0001,
);

$yhomogeneous = Compute("c1 e^(-$k x) + c2 x e^(-$k x)");


Context()->texStrings;
BEGIN_TEXT
Solve the following differential equation by variation of parameters.  Fully evaluate all integrals.
\[
$diffeq.
\]

\{ BeginList('OL',type=>'a') \}

$ITEMSEP
$ITEM Find the most general solution to the 
associated homogeneous differential equation.  
Use \( c_1 \) and \( c_2 \) in your answer to denote 
arbitrary constants, and enter them as ${BTT}c1${ETT}
and ${BTT}c2${ETT}.
$BR
$BR
\( y_h = \)
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

ANS( $yhomogeneous->cmp( 
  checker => sub {
    my ( $correct, $stu, $answerHash ) = @_;

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      $stu->{variables}{'x'} != 1 ||
      $stu->{variables}{'c1'} != 1 ||
      $stu->{variables}{'c2'} != 1
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(random(0.2,0.3,0.01));

    my $a11 = $stu->eval('c1'=>1,'c2'=>0,x=>$x);
    my $a12 = $stu->eval('c1'=>0,'c2'=>1,x=>$x);
 
    my $a21 = $stu->D('x')->eval('c1'=>1,'c2'=>0,x=>$x);
    my $a22 = $stu->D('x')->eval('c1'=>0,'c2'=>1,x=>$x);

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($a11 * $a22 == $a21 * $a12);

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu1 = Formula($stu->D('x'));
    my $stu2 = Formula($stu->D('x','x'));
    return ($stu2 + $kk * $stu) == (-1 * $k2 * $stu1); 

}));





#############################
#  Part 2

#Context("Numeric");
#Context()->variables->are(
#  'x' => ["Real", limits => [-5,5] ],
#  'a' => "Parameter",
#  'b' => "Parameter",
#);
#Context()->flags->set(
#  tolType => 'absolute',
#  tolerance => 0.0001,
#);

#$yparticular = Compute("a e^(-$k x) + b x e^(-$k x) + x/$k - 2/$kk");

$yparticular = Compute("x/$k - 2/$kk");

Context()->texStrings;
BEGIN_TEXT

$ITEMSEP
$ITEM Find a particular solution to the 
nonhomogeneous differential equation \( $diffeq \).
$BR
$BR
\( y_p = \) 
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;

ANS( $yparticular->cmp(
 checker => sub {
   my ($correct, $student, $ansHash) = @_;
     $dstu = $student->D('x');
     $ddstu = $dstu->D('x');
     return ($ddstu + $k2*$dstu + $kk*$student)==Compute("$k*x");
 }
) );


###########################################
# Part 3
#
# Get rid of adaptive parameters which might mess with answer checkers.

Context("Numeric");
Context()->variables->are(
  'x' => ["Real", limits => [-1,1] ],
  'c1' => ["Real", limits => [2,4] ],
  'c2' => ["Real", limits => [2,4] ],
);
Context()->flags->set(
  tolType => 'absolute',
  tolerance => 0.0001,
);

$ygeneral = Compute("c1 e^(-$k x) + c2 x e^(-$k x) + x/$k - 2/$kk");


BEGIN_TEXT
$ITEMSEP
$ITEM Find the most general solution to the 
original nonhomogeneous differential equation.  
Use \( c_1 \) and \( c_2 \) in your answer to denote 
arbitrary constants.
$BR
$BR
\( y = \)
\{ ans_box(2,60) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL') \}

END_TEXT

ANS($ygeneral->cmp( 
  checker => sub {
    my ( $correct, $stu, $answerHash ) = @_;

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      $stu->{variables}{'x'} != 1 ||
      $stu->{variables}{'c1'} != 1 ||
      $stu->{variables}{'c2'} != 1
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(random(0.2,0.3,0.01));

    my $a11 = $stu->eval('c1'=>1,'c2'=>0,x=>$x);
    my $a12 = $stu->eval('c1'=>0,'c2'=>1,x=>$x);
 
    my $a21 = $stu->D('x')->eval('c1'=>1,'c2'=>0,x=>$x);
    my $a22 = $stu->D('x')->eval('c1'=>0,'c2'=>1,x=>$x);

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($a11 * $a22 == $a21 * $a12);

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu1 = Formula($stu->D('x'));
    my $stu2 = Formula($stu->D('x','x'));
    my $g = Formula("$k x");

    return $g == ($stu2 + $k2 * $stu1 + $kk * $stu); 

}));

COMMENT("MathObject version.  Characteristic polynomial (r + k)^2.");

ENDDOCUMENT();
