# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('multivariable', 'functions', 'level surface')
## Tagged by glr 05/14/09


## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('12.5')
## Problem1('21')

## Textbook tags
## HHChapter1('Functions of Several Variables')
## HHSection1('Functions of Three Variables')

Context("Numeric");
Context()->variables->add( y=>'Real', z=>'Real' );
Context::Value::Function->{cmpDefaults} = {upToConstant => 1};
$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(2,8,1);
$c = random(2,8,1);
$p = Compute( "$a*x + $b*y - $c" )->reduce();
$p2 = Compute( "$a*x + $b*y - z" )->reduce();
$g1 = Compute( "z - $a*x - $b*y" );
$g2 = Compute( "$a*x + $b*y - z" );

# $f1 = Compute("x");

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a formula for a function \( g(x,y,z) \) whose level surfaces are
planes parallel to the plane \( z = $p \).
$PAR
\( g(x,y,z) = \) \{ ans_rule(45) \}

END_TEXT
Context()->normalStrings;

ANS($g1->cmp( checker=>sub { my ( $c, $s, $self ) = @_;
			      my $context = Context()->copy();
			      $context->flags->set( no_parameters=>0 );
			      $context->variables->add( 'C0'=>'Parameter' );
			      my $c0 = Formula($context,'C0');
			      $s = Formula($context, $s);
			      return $c0 + $g1 == $s || $c0 + $g2 == $s; } ) );
## for some reason checking $s == $c0 + $g1 || $s == $c0 + $g2 fails

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The equation of any plane with normal
\($a\,\overrightarrow{i}+$b\,\overrightarrow{j}-\overrightarrow{k}\) is
\($p2 = c\), for some constant \(c\).  Thus we could take
\(g(x,y,z)=$p2\).  Other answers are possible.


END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
