# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

## Tagged by glr 05/14/09


## DBsubject(Calculus - multivariable)
## DBchapter(Concepts for multivariable functions)
## DBsection(Traces, contours, and level sets)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('12.5')
## Section2('12.5')
## Section3('12.5')
## Problem1('23')
## Problem2('26')
## Problem3('31')
## KEYWORDS('multivariable', 'functions', 'level surface')

## Textbook tags
## HHChapter1('Functions of Several Variables')

Context("Numeric");
Context()->variables->add( y=>'Real', z=>'Real' );
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 1;

$a = random(1,5,1);
$asq = $a*$a;
$xy = list_random('x','y');
$yx = ( $xy eq 'x' ) ? 'y' : 'x';

$f = Compute( "sqrt($asq - $xy^2)" );
$g1 = Compute( "$f - z" );
$g2 = Compute( "z - ($f)" );

$surf = MultiAnswer( $g1, Compute(0) )->with(
    singleResult => 0,
    checker => sub {
	my ( $cor, $stu, $ansHash ) = @_;
	my ( $corF, $corC ) = @$cor;
	my ( $stuF, $stuC ) = @$stu;

	if ( $corF == ( $stuF - $stuC ) ||
	     $g2 == ( $stuF - $stuC ) ) {
	    return [ 1, 1 ];
	} else {
	    return [ 0, 0 ];
	}
    }
);

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

The surface \( S \) is the graph of \( f(x,y) = $f \).
$PAR
Be sure you can explain why \( S \) is the upper half of a circular
cylinder of radius $a, centered along the  \( $yx \)-axis.  Write three
or four sentences on a sheet of paper that clearly demonstrate this.
$PAR
Find a level surface \( g(x,y,z)=c \) representing \( S \).
$PAR
\( g(x,y,z) = \) \{ $surf->ans_rule(45) \},
$BR
with \(c = \) \{ $surf->ans_rule(10) \}
$BR
${BITALIC}(Note that because your answers for \(g\) and \(c\) are
interdependent, you cannot get partial credit for this problem.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $surf->cmp() );
# ANS( $g1->cmp( checker=>sub { my ( $cor, $stu, $self ) = @_;
# 			      return $stu == $cor || $stu == $g2; } ) );
# ANS( Compute(0)->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR


The graph of \(f(x,y)\) is obtained by plotting points \((x,y,z)\),
where \(z=f(x,y)\). Since the square root function is never negative, we
have \(z \ge 0\).  Setting \(z = \sqrt{$asq - $xy^2}\) and squaring
both sides leads to \($xy^2 + z^2 = $asq\), which is the equation for
a circular cylinder of radius $a lying along the \($yx\)-axis (since
\($yx\) is missing from the equation). The graph of the function includes
only those points where \(z \ge 0\), that is, the upper half of the
cylinder.
$PAR
If we take \(g(x,y,z) = f(x,y) - z = $f - z\), then the level surface
\(g(x,y,z) = 0\) is the surface \(S\).
$PAR
Clearly, we could also take \(g(x,y,z) = z - f(x,y) = z - $f\) as well.


END_SOLUTION
Context()->normalStrings;

;
ENDDOCUMENT();
