## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION

## KEYWORDS ('linear algebra','matrix','eigenvalue','eigenvector','systems','systems of differential equations')

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Systems of Differential Equations')
## Date('01/31/2012')
## Author('Paul Pearson')
## Institution('Rochester and Fort Lewis College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"parserRadioButtons.pl",
"unionLists.pl",
"weightedGrader.pl",
);

install_weighted_grader();

TEXT(beginproblem());



#####################################
#  Setup 1

Context("Complex");

# roots -a +/- bi

$a = random(2,6,1);
do { $b = random(2,6,1); } until ($b != $a);

$bb = $b * $b;


$A11 = -1*($a);  $A12 = -1*($bb);
$A21 = 1;        $A22 = -1*($a);

$lambda1 =  Complex("-$a + $b i");
$v11 = Complex("-$b");
$v21 = Complex("i");

$lambda2 = Complex("-$a -$b i");
$v12 = Complex("-$b");
$v22 = Complex("-i");


$multians1 = MultiAnswer($lambda1, $v11, $v21, $lambda2, $v12, $v22)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='0'>
    <tr>
      <td> %s, </td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td>&nbsp;&nbsp; %s,</td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  
    "%s, \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack; \ \ 
     %s, \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};

        if (   $c[1]->typeMatch($s[1]) 
            && $c[2]->typeMatch($s[2]) 
            && $c[4]->typeMatch($s[4]) 
            && $c[5]->typeMatch($s[5]) 
            && ($s[1] * $s[5] == $s[2] * $s[4])) {
              Value::Error("Your vectors are not linearly independent");
        }

        if (   $c[0]->typeMatch($s[0]) 
            && $c[1]->typeMatch($s[1]) 
            && $c[2]->typeMatch($s[2]) 
            && ($s[1] != 0 || $s[2] != 0)) {
          if ( ($A11 * $s[1] + $A12 * $s[2] == $s[0] * $s[1]) &&
               ($A21 * $s[1] + $A22 * $s[2] == $s[0] * $s[2]) 
          ){
                $flag1=1;
          }
        }

        if (   $c[3]->typeMatch($s[3]) 
            && $c[4]->typeMatch($s[4]) 
            && $c[5]->typeMatch($s[5]) 
            && ($s[4] != 0 || $s[5] != 0)) {
          if ( ($A11 * $s[4] + $A12 * $s[5] == $s[3] * $s[4]) &&
               ($A21 * $s[4] + $A22 * $s[5] == $s[3] * $s[5]) 
          ){
                $flag2=1;
          }
        }

        if ($flag1==1 && $flag2 ==1) {
           return 1;
        } elsif ($flag1 == 1) {
            Value::Error("The eigenpair { lambda1, v1 } is correct");
        } elsif ($flag2 == 1) {
            Value::Error("The eigenpair { lambda2, v2 } is correct");
        } else {
            Value::Error("Neither of the eigenpairs is correct");
        }
       
    }
);






#####################################
#  Main text 1

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL', type=>'a') \}
$ITEM
Consider the linear system
$PAR
\[ 
\displaystyle
\boldsymbol{ \vec{y}^{\,\prime} }
=
\left\lbrack
\begin{array}{rr}
$A11 & $A12 \\
$A21 & $A22
\end{array}
\right\rbrack
\boldsymbol{\vec{y}}.
\]
$BR
$BR
Find the eigenvalues and eigenvectors for the coefficient matrix.
$BCENTER
\{ 
mbox(
[
'\( \lambda_1 = \) ',
$multians1->ans_rule(5),
', \( \boldsymbol{\vec{v}_1} =  \) ', 
display_matrix([[$multians1->ans_rule(5)],[$multians1->ans_rule(5)]]),
', and \( \lambda_2 = \) ',
$multians1->ans_rule(5),
', \( \boldsymbol{\vec{v}_2} =  \) ', 
display_matrix([[$multians1->ans_rule(5)],[$multians1->ans_rule(5)]]),
],
allowbreaks=>'yes'
) 
\}
$ECENTER
END_TEXT
Context()->normalStrings;


##################################
#  Answer evaluation 1

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 50 );



###############################################################
###############################################################


########################
#  Setup 2

Context("Numeric");
Context()->variables->add(
't'=>'Real',
'c1'=>'Real',
'c2'=>'Real',
);

$x1 = Compute("e^(-$a t) * ( $b c1 cos($b t) + $b c2 sin($b t) )");
$x2 = Compute("e^(-$a t) * ( c1 sin($b t) - c2 cos($b t) )");

$gensolution = MultiAnswer($x1, $x2)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='0'>
    <tr><td>x<sub>1</sub>(t) = %s</td></tr>
    <tr><td>x<sub>2</sub>(t) = %s</td></tr>
    </table>",
    tex_format =>  
    "\begin{array}{l} 
    x_1(t) = %s \\ 
    x_2(t) = %s 
    \end{array}",

    checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};

      if ( !$c[0]->typeMatch($s[0]) || !$c[1]->typeMatch($s[1]) ) {
           Value::Error("<b>You must enter both x<sub>1</sub>(t) and x<sub>2</sub>(t)</b>");
      }

      if (    $c[0]->typeMatch($s[0]) && $s[0] != Formula(0) && $s[0]->isFormula
           && $c[1]->typeMatch($s[1]) && $s[1] != Formula(0) && $s[1]->isFormula
           && (    $s[0]->D('c1') == Formula(0) || $s[1]->D('c1') == Formula(0)
                || $s[0]->D('c2') == Formula(0) || $s[1]->D('c2') == Formula(0)
              )
         ) 
      { Value::Error("Don't forget the constants c1 and c2"); } 

 
      if (    $c[0]->typeMatch($s[0]) && $s[0] != Formula(0) && $s[0]->isFormula
           && $c[1]->typeMatch($s[1]) && $s[1] != Formula(0) && $s[1]->isFormula
           && $s[0]->D('c1') != Formula(0) 
           && $s[1]->D('c1') != Formula(0)
           && $s[0]->D('c2') != Formula(0)
           && $s[1]->D('c2') != Formula(0)
           && $A11 * $s[0] + $A12 * $s[1] == $s[0]->D('t')
           && $A21 * $s[0] + $A22 * $s[1] == $s[1]->D('t')
         ) 
      { return 1; } else { return 0; }

    }
);


$radio = RadioButtons(
  ["source / unstable node","sink / stable node","saddle","center point / ellipses","spiral source","spiral sink","none of these"],
  "spiral sink", # correct answer
  last => ["none of these"], # can be a list
);



Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM
Find the most general solution to the linear system of 
differential equations \( \boldsymbol{\vec{x}^{\, \prime}} = P \boldsymbol{\vec{x}} \). 
Use \( c_1 \) and \( c_2 \) to denote arbitrary constants, and enter them as ${BTT}c1${ETT}
and ${BTT}c2${ETT}.
$BR
$BR
\( x_1(t) = \) \{ $gensolution->ans_rule(40) \}
$BR
\( x_2(t) = \) \{ $gensolution->ans_rule(40) \}

$ITEMSEP
$ITEM The behavior of this system is best described as
$BR
\{ $radio->buttons() \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;



WEIGHTED_ANS( $gensolution->cmp(), 45 );
WEIGHTED_ANS( $radio->cmp(), 5 );



COMMENT('MathObject version.  Complex eigenvalues -a +/- bi and complex eigenvectors.');

ENDDOCUMENT(); 