## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Phase planes)
## Date(03/19/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-PP-DRE)
## Level(4)
## MO(1)


##################################
#  Initialization


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "MatrixUnimodular.pl",
  "unionLists.pl",
  "parserRadioButtons.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#ProblemRandomize(when=>"Always");

install_weighted_grader();


##################################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");


$lambda1 = random(1,3,1);
do { $lambda2 = random(1,3,1); } until ($lambda1 != $lambda2);

$radio = RadioButtons(
  ["source / unstable node","sink / stable node","saddle","center point / ellipses","spiral source","spiral sink","none of these"],
  "source / unstable node", # correct answer
  last => ["none of these"], # can be a list
);

@M = small_unimodular_diagonalization_GL2Z($lambda1,$lambda2);

# the matrix to be analyzed
$A11 = $M[0];  $A12 = $M[2];
$A21 = $M[1];  $A22 = $M[3];

$v1top = Formula("$M[4] e^($lambda1 t)");
$v1bot = Formula("$M[5] e^($lambda1 t)");

$v2top = Formula("$M[6] e^($lambda2 t)");
$v2bot = Formula("$M[7] e^($lambda2 t)");


$multians = MultiAnswer($v1top, $v1bot, $v2top, $v2bot)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='0'>
    <tr>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td>&nbsp;&nbsp;</td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  
    "\left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack, \ \ 
     \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};

        $flag1 = 0;
        $flag2 = 0;
        $flag  = 0;

        Value::Error("Your solutions are not linearly independent") if
        $c[0]->typeMatch($s[0]) &&
        $c[1]->typeMatch($s[1]) &&
        $c[2]->typeMatch($s[2]) &&
        $c[3]->typeMatch($s[3]) &&
        ($s[0] * $s[3] == $s[1] * $s[2]);

        if ( $c[0]->typeMatch($s[0]) && 
             $c[1]->typeMatch($s[1]) && 
             ($s[0] != 0 || $s[1] != 0) ) 
        {
          if ( ($A11 * $s[0] + $A12 * $s[1] == $s[0]->D('t')) &&
               ($A21 * $s[0] + $A22 * $s[1] == $s[1]->D('t')) 
          ){
                $flag1=1;
          }
        }

        if ( $c[2]->typeMatch($s[2]) && 
             $c[3]->typeMatch($s[3]) && 
             ($s[2] != 0 || $s[3] != 0) ) 
        {
          if ( ($A11 * $s[2] + $A12 * $s[3] == $s[2]->D('t')) &&
               ($A21 * $s[2] + $A22 * $s[3] == $s[3]->D('t')) 
          ){
                $flag2=1;
          }
        }

        if($flag1==1 && $flag2 ==1){
           $flag=1;
        }else{
           if ($flag1 == 1) {
              Value::Error("The solution on the left is correct");
           } elsif ($flag2 == 1) {
              Value::Error("The solution on the right is correct");
           }else{
              Value::Error("Neither of the solutions is correct");
           }
        }
        return $flag; 
       
    }
);


#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

$ITEM
Find the most general real-valued solution to the linear system 
of differential equations
\(
\displaystyle
\boldsymbol{\vec{x}^{\, \prime}}
=
\left\lbrack
\begin{array}{rr}
$M[0] & $M[2] \\
$M[1] & $M[3]
\end{array} 
\right\rbrack
\boldsymbol{\vec{x}}.
\)
$BR
$BR
\{ 
mbox(
[
display_matrix([['\( x_1(t) \)'],['\( x_2(t) \)']]),
'\( = c_1 \)', 
display_matrix([[$multians->ans_rule(10)],[$multians->ans_rule(10)]]),
'\( + \ c_2 \)', 
display_matrix([[$multians->ans_rule(10)],[$multians->ans_rule(10)]]),
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM In the phase plane, this system is best described as a
$BR
\{ $radio->buttons() \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians->cmp(), 95 );
WEIGHTED_ANS( $radio->cmp(), 5 );

COMMENT('MathObject version. Distinct positive integer eigenvalues.');

ENDDOCUMENT();
