## DESCRIPTION
## Linear systems: Nonhomogeneous linear systems
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Applications)
## Date(04/12/2011)
## Institution(Fort Lewis College)
## Author(Veronika Furst)
## MLT(SDE-A-mixture)
## MLTleader(1)
## Level(5)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.8')
## Problem1('33')
## KEYWORDS('Nonhomogeneous linear systems')


##################################
#  Initialization


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##################################
#  Setup

Context("Matrix");
Context()->variables->are(t=>"Real");

$r = 10;
$V = 1000;
$C1 = random(0.2,0.8,0.1);
$C2 = 0;

$init1 = random(0,30,10);
$init2 = random(0,30,10);

# matrix A listed by rows (not columns)
@A = ();
$A[0] = Compute("-$r / $V * 2"); # A11
$A[1] = Compute("$r / $V");  # A12
$A[2] = Compute("$r / $V");  # A21
$A[3] = Compute("-$r / $V * 2"); # A22


$multians1 = MultiAnswer($A[0], $A[1], $A[2], $A[3])->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\\\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);




@p = ();
$p[0] = Compute("$C1 * $r"); # p1
$p[1] = Compute("$C2 * $r"); # p2


$multians2 = MultiAnswer($p[0], $p[1])->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$coeff1 = Compute("1/2*($init2 - $init1 + 100/3*$p[0])");
$coeff2 = Compute("1/2*($init1 + $init2 - 100*$p[0])");

$Q1 = Compute("$coeff1*(-1)*e^(-3/100*t) + $coeff2*e^(-1/100*t) + 200/3*$p[0]");
$Q2 = Compute("$coeff1*e^(-3/100*t) + $coeff2*e^(-1/100*t) + 100/3*$p[0]");

$Q1->{limits} = [90,110];
$Q2->{limits} = [90,110];


$Q1inf = Compute("200/3*$p[0]");
$Q2inf = Compute("100/3*$p[0]");


#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{
ColumnTable(
"The flow system shown in the figure is activated at time 
\( t=0 \).  Let \( Q_i(t) \) denote the amount of solute 
present in the \( i \)th tank at time \( t \).  Assume that 
all the flow rates are a constant \( $r \ \mathrm{L/min} \).  
It follows that the volume of solution in each tank remains 
constant;  assume this volume to be \( $V \ \mathrm{L} \).  
The concentration of solute in the inflow to Tank 1 (from 
a source other than Tank 2) is \( $C1 \ \mathrm{kg / L} \), 
and the concentration of solute in the inflow to Tank 2 
(from a source other than Tank 1) is \( $C2 \ \mathrm{kg / L} \).  
Assume each tank is mixed perfectly."
,
$BCENTER.
image("KJ-4-8-33.png", width=>302, height=>103, tex_size=>600).
$ECENTER
,
indent => 0, separation => 30, valign => "TOP"
);
\}

$BR

\{ BeginList('OL',type=>'a') \}

$ITEM Set up a system of first-order differential equations that models this
situation.
$BR
$BR
\{
mbox(
[
display_matrix([['\(Q_1^{\,\prime}\)'],['\(Q_2^{\,\prime}\)']]),
' = ', 
display_matrix([
[$multians1->ans_rule(10),$multians1->ans_rule(10)],
[$multians1->ans_rule(10),$multians1->ans_rule(10)]
]),
display_matrix([['\(Q_1\)'],['\(Q_2\)']]),
' + ',
display_matrix([[$multians2->ans_rule(10)],[$multians2->ans_rule(10)]])
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM
If \( Q_1(0) = $init1 \ \mathrm{kg} \) and \( Q_2(0) = $init2 \
\mathrm{kg} \), find the amount of solute in each tank 
after \( t \) minutes.
$BR
\( Q_1(t) = \) \{ ans_rule(60) \} kg
$BR
\( Q_2(t) = \) \{ ans_rule(60) \} kg


$ITEMSEP
$ITEM
As \( t \to \infty \), how much solute is in each tank?
$BR
In the long run, Tank 1 will have \{ ans_rule(10) \} kg of solute.
$BR
In the long run, Tank 2 will have \{ ans_rule(10) \} kg of solute.
$BR
(Reread the question and think about why this answer makes sense.)
\{ EndList('OL') \}

END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

ANS( $multians2->cmp() );

#foreach my $i (0..3) {
#  ANS( $A[$i]->cmp() );
#}

#foreach my $j (0..1) {
#  ANS( $p[$j]->cmp() );
#}

ANS( $Q1->cmp() );
ANS( $Q2->cmp() );
ANS( $Q1inf->cmp() );
ANS( $Q2inf->cmp() );


;

ENDDOCUMENT();
