##DESCRIPTION
##  superposition principle.
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Introductory concepts)
## DBsection(Verification of solutions)
## Date(2/13/2010)
## Institution(Alfred University)
## Author(Darwyn Cook)
## MLT(Darwyn-DEFM)
## Level(4)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('3')
## Section1('4')
## Problem1('')
## KEYWORDS('superposition')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Interval");
$domain=Compute("(0,inf)");

$frac=Context("Fraction");
$frac->variables->add(a=>"Real",b=>"Real");
$p = Compute(random(2,9));
$p2 = Compute($p**2);
$c = Compute(random(1,9));
$d = Compute(random(1,9));
$homogeneousprint = Formula($frac,"a/x+b/x^(1/$p)");

Context("Numeric");
Context()->variables->add(a=>"Real",b=>"Real");
Context()->variables->set(x=>{limits=>[1,10]});
Context()->flags->set(tolerance=>0.00001,tolType=>"absolute");

$homogeneous = Formula("a/x+b/x^(1/$p)");
$non = Formula("$c+$d*x")->reduce;
$general = Formula($homogeneous+$non);
$rhs = Formula("(2*$p+2)*$d*x+$c")->reduce;
$wronskian = Formula("($p-1)/($p*x^(2+1/$p))");
$negwronskian = -$wronskian;

Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add("y"=>"Real","y'"=>"Real","y''"=>"Real");
$diffeqn = Formula("$p*x^2*y''+(2*$p+1)*x*y'+y")->reduce;



################################################
##
## IVP setup
##

$y0 = Compute(random(2,9));
$y1 = Compute(random(2,9));
$a = Compute("(-$y0-$y1*$p+$d*$p+$c+$d)/($p-1)");
$b = Compute("-$p*(-$y1-$y0+$c+2*$d)/($p-1)");
$ivp = $general->substitute(a=>$a,b=>$b,x=>x);


##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
One of the following is a general solution of the homogeneous differential equation \($diffeqn=0\)
$BR
\[
\begin{aligned}
y&=\frac{a}{x^2}+ \frac{b}{x^{\frac{1}{$p2}}} \\
y&=$homogeneousprint \\
\end{aligned}
\]
and one of the following is a solution to the nonhomogeneous equation \($diffeqn=$rhs\)
\[
\begin{aligned}
y&=x^2-x \\
y&=$non \\
y&=e^x \\
\end{aligned}
\]
By superposition, the general solution of the equation \($diffeqn=$rhs\) is
$BR
\(y = \) \{$general->ans_rule\}
$BR
Find the solution with
$BR
\[
\begin{aligned}
y(1)&=$y0 \\
y^\prime(1)&=$y1 \\
\end{aligned}
\]
\(y = \)\{$ivp->ans_rule\}
$BR
$BR
Fortunately, since I already had it computed, my professor also asked me to find the Wronskian (using only the solutions to the homogeneous equation, without the coefficients a and b) \{$wronskian->ans_rule\}. The fundamental theorem shows that this solution is the unique solution to the IVP on the interval \{$domain->ans_rule\}
END_TEXT

Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($general->cmp);
ANS($ivp->cmp);
ANS($wronskian->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
               return (($correct == $student)||($student == -$correct));}));
ANS($domain->cmp);
ENDDOCUMENT();        
