##DESCRIPTION
##  superposition principle.
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear independence)
## Date(2/13/2010)
## Institution(Alfred University)
## Author(Darwyn Cook)
## MLT(Darwyn-DEFM)
## MLTleader(1)
## Level(4)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('3')
## Section1('4')
## Problem1('')
## KEYWORDS('superposition')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "answerHints.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup the differential equation and general solution to the homogeneous equation
#
#
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add(a=>"Real",b=>"Real","y"=>"Real","y'"=>"Real","y''"=>"Real");

$m1 = Compute(non_zero_random(-3,3));
$m2 = Compute($m1+non_zero_random(-3,3));
if ($m1 == 1){$m1++;};
if ($m2 == 1){$m2++}; # If either of these is 1, $d will be zero which is problematic.
if ($m1 == 0){$m1=3;};
if ($m2 == 0){$m2=3;}; # If either of these is 1, $d will be zero which is problematic.
if ($m2 == $m1){$m2 = 5}; # If these are equal you get a zero Wronskian.
$bogus2 = -$m2;
if ($bogus2 == $m1){$bogus2=-5;};
$diffeqn = Formula("y''-($m1+$m2)y'+($m1*$m2)y")->reduce;

$homogeneous = Formula("ae^($m1*x)+be^($m2*x)");

$wronskian = Formula("($m1-$m2)e^(($m1+$m2)x)");# If you change this you have to change the answer checker

###########################################################################
##
## Now setup the nonhomogeneous equation and solution
##
$c = Compute(non_zero_random(-3,3));
$d = Compute("(1-($m1+$m2)+$m1*$m2)*$c");
$rhs = Formula("$d*e^x");
$nonsoln = Formula("$c e^x")->reduce;
$general = Compute($homogeneous+$nonsoln);

# randomize the lists of choices
@gen = (Compute("a e^{$m1 x}+b e^{$bogus2 x}"),Compute("a e^{$bogus1 x}+b e^{$m2 x}"),$homogeneous,Compute("ae^{$bogus1 x}+be^{$bogus2 x}"));
@genselected = @gen[ NchooseK(4,4) ];

@part = (Compute("(2*$c)e^{x}")->reduce,Compute("(3*$c)e^{x}")->reduce,$nonsoln,Compute("(4*$c) e^{x}")->reduce);
@partselected = @part[ NchooseK(4,4)];

################################################
##
## IVP setup
##

$y0 = Compute(random(2,9));
$y1 = Compute(random(2,9));
$a = Compute("($y1-$c-$y0*$m2+$c*$m2)/(-$m2+$m1)");
$b = Compute("(-$m1*$y0+$y1-$c+$c*$m1)/($m2-$m1)");
$ivp = $general->substitute(a=>$a,b=>$b,x=>x);

Context("Interval");
$domain = Compute("(-inf,inf)");
##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
$BR
One of the following is a general solution of the homogeneous differential equation \($diffeqn=0\)
$BR
\[
\begin{aligned}
y&=$genselected[0] \\
y&=$genselected[1] \\
y&=$genselected[2] \\
y&=$genselected[3] \\
\end{aligned}
\]
and one of the following is a solution to the nonhomogeneous equation \($diffeqn=$rhs\)
\[
\begin{aligned}
y&=$partselected[0] \\
y&=$partselected[1] \\
y&=$partselected[2] \\
y&=$partselected[3] \\
\end{aligned}
\]
By superposition, the general solution of the equation \($diffeqn=$rhs\) is
$BR
\(y = \) \{$general->ans_rule(30)\}
$BR
Find the solution with
$BR
\[
\begin{aligned}
y(0)&=$y0 \\
y^\prime(0)&=$y1 \\
\end{aligned}
\]
\(y = \)\{$ivp->ans_rule(30)\}
$BR
$BR
Being the highly trained differential equations fighting machine that I am, I checked the Wronskian (using the solutions to the homogeneous equation without the coefficients a and b) for good measure, and found it to be \{$wronskian->ans_rule\}. The fundamental theorem tells me that  that this is the unique solution to the IVP on the interval \{$domain->ans_rule\}.
END_TEXT

Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($general->cmp);
ANS($ivp->cmp);
ANS($wronskian->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                return (($correct == $student)||(-$correct == $student));}));
ANS($domain->cmp);
ENDDOCUMENT();        
