##DESCRIPTION
#
# File Created: 3/27/2009
# Last Modified: 3/27/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear independence)
## Date(3/27/2009)
## Institution(Alfred University)
## Author(Juan Marin)
## Level(2)
## TitleText1('Differential Equations with Boundary-Value Problems')
## AuthorText1('Zill and Cullen')
## EditionText1('7')
## Section1('4.7')
## Problem1('')


DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraders.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);
 
$showPartialCorrectAnswers = 1;

Context("Numeric");
Context()->operators->remove('+','-','/');
Context()->functions->disable("All");
Context()->functions->enable("exp");
$a = random(2,5,1);
$b = $a*$a;
$c = 2*$a;
$w = Compute("e^($c*x)");

$interval = Context("Interval");
$int = Compute("(-inf,inf)");

TEXT(beginproblem());

Context()->texStrings;
BEGIN_TEXT
$PAR
It can be shown that \(y_1=e^{$a x}\)  and \(y_2=xe^{$a x}\)  are 
solutions to the differential equation 
\(\frac{d^2y}{dx^2} - $c\frac{dy}{dx} + $b y =0\).

$PAR 
\(W(y_1,y_2)\) =  \{$w->ans_rule()\}.
$BR
\(c_1y_1+c_2y_2\) is the general solution to the equation on the interval \{$int->ans_rule\}
END_TEXT

ANS($w->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0(e^($c*x))");
                $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));
ANS($int->cmp());
ENDDOCUMENT()
