## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','spring system')

## Date('April 2014')
## Author('Benjamin Walter')
## Institution('METU NCC')

######################################
#  Initialization

DOCUMENT() ;        
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserVectorUtils.pl",
"PGcourse.pl",
"parserPopUp.pl",
"problemRandomize.pl"
) ;

TEXT(beginproblem());

#######################################
#  Setup

Context('Numeric');

####
# Problem statement
#
$freq = random(2,5,1);
$l = $freq**2;

@V = (non_zero_random(-5,5,1),
      non_zero_random(-5,5,1),
      non_zero_random(-5,5,1));

######
# Part A
#
$mA1 = random(1,3,1);
do { $mA2 = random(1,3,1); } until ( $mA2 != $mA1 );

$Aamp1 = random(2,10,1);  
$Aamp2 = abs($Aamp1 * $V[$mA2-1] / $V[$mA1-1]);

#####
# Part B
#
$mB1 = random(1,3,1);
$tB1 = random(1,5,1);
$tB2 = $tB1 + (2 * pi) / $freq;

######
# Part C
#
$mC1 = random(1,3,1);
do { $mC2 = random(1,3,1); } until ( $mC2 != $mC1 );
$popup = PopUp(["?", "same", "opposite"],"same");
if ( $V[$mC1-1]*$V[$mC2-1] < 0 ) {
  $popup = PopUp(["?", "same", "opposite"],"opposite");
}

####
# Part D
#
$mD1 = random(1,3,1);
$pos = random(-4,4,1);
do { $vel = random(-4,4,1); } until (($pos != 0 ) || ( $vel != 0)); 

Context()->variables->add(t=>'Real');

$displ = Formula("$pos * cos($freq * t) + $vel / $freq * sin($freq * t)");


Context('Vector');

$v = ColumnVector( @V );


#######################################
#  Main text

Context()->texStrings;

BEGIN_TEXT

Suppose a spring system with three masses and many springs oscillates at 
a fundamental mode with 
$BR
$BCENTER
eigenvalue \(\lambda = $l \qquad \) and \(\qquad\) eigenvector
\(\mathbf{v} = $v\).
$ECENTER

$PAR
$HR
$PAR

${BBOLD}(A)${EBOLD}
If mass $mA1 oscillates with an amplitude of $Aamp1, then what is the 
amplitude of oscillation for mass $mA2?
$BR \(\quad\)
Amplitude = \{ ans_rule(5) \}

$PAR

${BBOLD}(B)${EBOLD}
If mass $mB1 is at maximum height at time \(t = $tB1\), when will it
next be at maximum height?
$BR \(\quad\)
\(t = \) \{ ans_rule(5) \}

$PAR

${BBOLD}(C)${EBOLD}
Mass $mC1 oscillates in the \{ $popup->menu() \} direction 
as mass $mC2.

$PAR

${BBOLD}(D)${EBOLD}
If mass $mD1 begins at displacement \(u_{$mD1}(0) = $pos\) and velocity 
\(u_{$mD1}'(0) = $vel\), then what is its 
dispacement function?
$BR
\(\quad u_{$mD1}(t) = \) \{ ans_rule(20) \}

END_TEXT

Context()->normalStrings;

#######################################
#  Answer evaluation

ANS( Real($Aamp2)->cmp(
  checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  Value->Error("Amplitudes must be postivie!") if ($student == -$correct);
  return cos($correct) == cos($student);
} ) );


ANS( Real($tB2)->cmp() );

ANS( $popup->cmp() );

ANS( $displ->cmp() );

COMMENT('Integer Eigenvectors; nonzero initial pos and velocity');

ENDDOCUMENT();
