# DESCRIPTION
# Find a 2x2 matrix zero divisor.
# ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Matrix algebra)
## Institution(Rochester)
## Level(2)
## MO(1)
## KEYWORDS('linear algebra','matrix','linear system')
## Tagged by cmd6a 5/3/06


###########################################################################
# initialization 
###########################################################################
DOCUMENT();        
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Matrix");
$a = non_zero_random(-9,9,1);
do {$b = non_zero_random(-9,9,1)} until ($b != $a);
$c = random(-5,-2,1);

$Left = Matrix([[$a,$b],[$c*$a, $c*$b]]);
$Right = Matrix([[-$b, -$b],[$a,$a]]); 
$Zero = Matrix([[0,0],[0,0]]); 


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Find a non-zero, \(2\times 2\) matrix such that:$BR
\($Left\)\(\ast\)\{$Right->ans_array()\} \(= $Zero\). 
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################

ANS($Right->cmp( 
  checker => sub {
    my ($corr, $stu, $ansHash) = @_; 
    if ( !$stu->isZero && ($Left*$stu)->isZero ){
      return 1;
    }
    else {return 0;} 
  }
));


###########################################################################
# solution
###########################################################################
Context("Fraction-NoDecimals");
$c1 = (abs($a) != 1) ? $a : (($a == 1) ? "" : "-");
$pm2 = ($b < 0) ? "-" : "+";
$c2 = abs($b);
$c3 = $c*$a;
$pm4 = ($b < 0) ? "+" : "-";
$c4 = abs($c*$b);
$na = -$a;
$frac = Fraction(-$b,$a);
Context()->texStrings;
BEGIN_SOLUTION
$PAR $BBOLD SOLUTION: $EBOLD
First let 
\[
X = \left[\begin{array}{cc} x_1 & x_2 \\ x_3 & x_4 \end{array}\right]
\] 
represent the desired matrix. Expanding the LHS of the matrix equation above yields the system of equations
\[ \left\{'\{'\} \begin{array}{rcrcr}
  $c1 x_1 & $pm2 & $c2 x_3 & = & 0\\
  $c1 x_2 & $pm2 & $c2 x_4 & = & 0\\
  $c3 x_1 & $pm4 & $c4 x_3 & = & 0\\
  $c3 x_2 & $pm4 & $c4 x_4 & = & 0
\end{array} \right. \]
in which the 3rd and 4th equations are equivalent to the the 1st and 2nd, respectively. Therefore, there are infinitely many solutions
given by \(x_1 = $frac x_3 \) and \(x_2 = $frac x_4 \), where \(x_3,x_4 \in \mathbb{R}\) can be chosen arbitrarily.  To ensure that \(X\) is not the zero matrix choose \(x_3 \neq 0\) or \(x_4 \neq 0\).  In particular one may take \(x_3 = $na\) and \(x_4 = 0\) which gives \(x_1 = $b\) and \(x_2 = 0\) so 
\[
X = \left[\begin{array}{rc} $b & 0 \\ $na & 0 \end{array}\right]
\]
END_SOLUTION
Context()->normalStrings;

COMMENT('Edited and updated in 2012/2013 by Rick Lynch @ Mizzou. Originally taken from NPL.
$BR
Desc: Find a 2x2 matrix zero divisor.');
ENDDOCUMENT();
