# DESCRIPTION
# 
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter()
## DBsection()
## KEYWORDS()
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "contextLimitedPowers.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;

#
# Subclass the numeric functions
#
package my::Function::numeric;
our @ISA = ('Parser::Function::numeric');

#
# Override sqrt() to return a special value (usually 1) when evaluated
# effectively eliminating it from the product.
#
sub sqrt {
  my $self = shift;
  my $value = $self->context->flag("setSqrt");
  return $value+1 if $value && $_[0] == 1; # force sqrt(1) to be incorrect
  return $value if $value;
  return $self->SUPER::sqrt(@_);
}

#
# end of subclass
#
package main;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
Context()->flags->set(reduceConstantFunctions=>0);
Context()->strings->add("not real"=>{}, "not a real number"=>{alias=>"not real"});
LimitedPowers::OnlyPositiveIntegers();

$r1 = random(2,5);
do {$r2 = random(2,5)} until ($r2 != $r1);
if ($r1 == 2) {
  $a1 = list_random(3,4,5,6,7,8,9,10,11,12,13,14,15,20);
} elsif ($r1 == 3) {
  $a1 = random(2,5);
} else {
  $a1 = random(2,3);
}
$a1r = ($a1)**$r1;

if ($r2 == 2) {
  $a2 = list_random(3,4,5,6,7,8,9,10,11,12,13,14,15,20);
} elsif ($r2 == 3) {
  $a2 = random(2,5);
} else {
  $a2 = random(2,3);
}
$a2r = ($a2)**$r2;

for ($i = 0; $i < random(20,50); $i++) { #more random??
  $case = random(0,1,1); # 0 - neg on inside of (a), 1 - neg on inside of (b)
}

if ($r1 == 2) {
  $dispa = ($case == 0) ? "\(\sqrt{-$a1r} = \)" : "\(-\sqrt{$a1r} = \)";
} else {
  $dispa = ($case == 0) ? "\(\sqrt[$r1]{-$a1r} = \)" : "\(-\sqrt[$r1]{$a1r} = \)";
}

if ($r2 == 2) {
  $dispb = ($case == 1) ? "\(\sqrt{-$a2r} = \)" : "\(-\sqrt{$a2r} = \)";
} else {
  $dispb = ($case == 1) ? "\(\sqrt[$r2]{-$a2r} = \)" : "\(-\sqrt[$r2]{$a2r} = \)";
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Simplify.
$PAR
(a) $dispa \{ans_rule(20)\}
$PAR
(b) $dispb \{ans_rule(20)\}
$PAR
Write your answer using radical notation if necessary or write ${BBOLD}not real${EBOLD} if the answer is not a real number.
END_TEXT


###########################################################################
# check the answer  
###########################################################################
if (($case == 0 && $r1 % 2 != 0) || $case == 1) {
$ans1 = Formula("-$a1");
$dans1 = "\($ans1\)";
ANS( $ans1-> cmp( checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 if $ans->{isPreview} || $correct != $student;
  #
  #  Get parsed formula for student and correct answers
  #
  $student = $ans->{student_formula};
  $correct = $correct->{original_formula} if defined $correct->{original_formula};
  #
  # check if equal when sqrt's are replaced by 1
  #
  Context()->flags->set(setSqrt => 1);
  delete $correct->{test_values};
  delete $student->{test_values};
  my $OK = ($correct == $student);
  Context()->flags->set(setSqrt => 0);
  #
  Value::Error("Check to see if your answer is simplified.") unless $OK;
  return $OK;
  }, formatStudentAnswer=>"reduced"
  )
);
} else {
  $dans1 = "not real";
  ANS(String("not real")->cmp());
}

if (($case == 1 && $r2 % 2 != 0) || $case == 0) {
$ans2 = Formula("-$a2");
$dans2 = "\($ans2\)";
ANS( $ans2-> cmp( checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 if $ans->{isPreview} || $correct != $student;
  #
  #  Get parsed formula for student and correct answers
  #
  $student = $ans->{student_formula};
  $correct = $correct->{original_formula} if defined $correct->{original_formula};
  #
  # check if equal when sqrt's are replaced by 1
  #
  Context()->flags->set(setSqrt => 1);
  delete $correct->{test_values};
  delete $student->{test_values};
  my $OK = ($correct == $student);
  Context()->flags->set(setSqrt => 0);
  #
  Value::Error("Check to see if your answer is simplified.") unless $OK;
  return $OK;
  }, formatStudentAnswer=>"reduced"
  )
);
} else {
  $dans2 = "not real";
  ANS(String("not real")->cmp());
}


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} (a) $dans1, $SPACE (b) $dans2
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Help:${EBOLD} 
  \{ AnswerFormatHelp("exponents", "Click here for help writing exponents and square roots.") \}
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
