##DESCRIPTION
## Parametrizing curves in space
##ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(pc_circles_center_radius)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('17.1')
## Problem1('18')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('17.1')
## Problem2('18')
## KEYWORDS('Vector','Parametric','Geometry')


DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT( beginproblem() );

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[-5,5]});

$r = random(3,8,1);
$x = Formula("$r * cos(t)") -> reduce;
$y = Formula("-($r) * sin(t)") -> reduce;

$multians = MultiAnswer($x, $y)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $xstu, $ystu ) = @{$student};
      if ( 
           ( (($xstu)**2+($ystu)**2) == $r**2 ) &&
           ( ($ystu->eval(t=>"0.1")) < 0 ) &&
           ( ($xstu->eval(t=>"0")) == $r ) &&
           ( ($ystu->eval(t=>"0")) == 0  ) 
         ) {
            return 1;
      } else {
            return 0;
      }
  }
);

Context()->texStrings;
BEGIN_TEXT
Find a parametrization of the circle of radius
\( $r \) in the xy-plane, centered at the origin,
oriented clockwise.  The point \( ($r,0) \) 
should correspond to \( t = 0 \).  Use \( t \) as the 
parameter for all of your answers.
$BR
$BR
\( x(t) = \) \{$multians->ans_rule(30)\}
$BR
\( y(t) = \) \{$multians->ans_rule(30)\}
END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());


;
ENDDOCUMENT();
