##DESCRIPTION
## Parametrizing curves in space
##ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(pc_ellipse)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('17.1')
## Problem1('22')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('17.1')
## Problem2('22')
## KEYWORDS('Vector','Parametric','Geometry')


DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT( beginproblem() );

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[-5,5]});

$rx = random(4,8,1);
$ry = $rx - 2;
$xd = 2 * $rx;
$yd = 2 * $ry;
$x = Formula("$rx * cos(t)") -> reduce;
$y = Formula("$ry * sin(t)") -> reduce;

$multians = MultiAnswer($x, $y)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $xstu, $ystu ) = @{$student};
      if ( 
           ( (($xstu/$rx)**2+($ystu/$ry)**2) == 1 ) &&
           ( ($ystu->eval(t=>"0.1")) > 0 ) &&
           ( ($xstu->eval(t=>"0")) == $rx ) &&
           ( ($ystu->eval(t=>"0")) == 0  ) 
         ) {
            return 1;
      } else {
            return 0;
      }
  }
);

Context()->texStrings;
BEGIN_TEXT
Find a parametrization of the ellipse centered at
the origin in the xy-plane that has major diameter 
\( $xd \) along the x-axis, minor diameter \( $yd \) 
along the y-axis, and is oriented counter-clockwise.  
Your parametrization should make the point \( ($rx,0) \) 
correspond to \( t = 0 \).  Use \( t \) as the 
parameter for all of your answers.
$BR
\( x(t) = \) \{$multians->ans_rule(30)\}
$BR
\( y(t) = \) \{$multians->ans_rule(30)\}
END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());


;
ENDDOCUMENT();
