##DESCRIPTION
## Parametrizing curves in space
##ENDDESCRIPTION

##KEYWORDS('Vector','Parametric','Geometry')

## DBsubject('Calculus')
## DBchapter('Vector Functions')
## DBsection('Vector Functions and Space Curves')
## Date('10/31/09')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('17.1')
## Section2('17.1')
## Problem1('16')
## Problem2('18')


DOCUMENT();      

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"PGcourse.pl",
);

TEXT( beginproblem() );

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[-5,5]});

$r = random(3,8,1);
$x = Formula("$r * cos(t)") -> reduce;
$y = Formula("-($r) * sin(t)") -> reduce;

$multians = MultiAnswer($x, $y)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $xstu, $ystu ) = @{$student};
      if ( 
           ( (($xstu)**2+($ystu)**2) == $r**2 ) &&
           ( ($ystu->eval(t=>"0.1")) < 0 ) &&
           ( ($xstu->eval(t=>"0")) == $r ) &&
           ( ($ystu->eval(t=>"0")) == 0  ) 
         ) {
            return 1;
      } else {
            return 0;
      }
  }
);

Context()->texStrings;
BEGIN_TEXT
Find a parametrization of the circle of radius
\( $r \) in the xy-plane, centered at the origin,
oriented clockwise.  The point \( ($r,0) \) 
should correspond to \( t = 0 \).  Use \( t \) as the 
parameter for all of your answers.
$BR
$BR
\( x(t) = \) \{$multians->ans_rule(30)\}
$BR
\( y(t) = \) \{$multians->ans_rule(30)\}
END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());


COMMENT('MathObject version');
ENDDOCUMENT();
