##DESCRIPTION
## Parametrizing curves in space
##ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(pc_cartesian_to_parametric)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('17.1')
## Problem1('26')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('17.1')
## Problem2('26')
## KEYWORDS('Vector','Parametric','Geometry')


DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT( beginproblem() );

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[-5,5]});

$a = random(-5,-2,1);
$x = Formula("$a * t**2");
$y = Formula("0");
$z = Formula("t");

$multians = MultiAnswer($x, $y, $z)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $xstu, $ystu, $zstu ) = @{$student};
      if ( 
           ($xstu == $a*$zstu**2) && ($ystu == 0) 
         ) {
            return 1;
      } else {
            return 0;
      }
  }
);

Context()->texStrings;
BEGIN_TEXT

Find a parametrization of the curve \( x = $a z^2 \) 
in the xz-plane.  Use \( t \) as the 
parameter for all of your answers.
$BR
$BR
\( x(t) = \) \{$multians->ans_rule(30)\}
$BR
\( y(t) = \) \{$multians->ans_rule(30)\}
$BR
\( z(t) = \) \{$multians->ans_rule(30)\}

END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 0;

ANS($multians->cmp());


;
ENDDOCUMENT();
