##DESCRIPTION
## Parametrizing curves in space
##ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(pc_ellipse_3d)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('17.1')
## Problem1('38')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('17.1')
## Problem2('38')
## KEYWORDS('Vector','Parametric','Geometry')


DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT( beginproblem() );

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[-5,5]});

$a = random(2,5,1);
$b = random(-5,-2,1);
while (abs($a) == abs($b)) { $a = random(2,5,1); }
$c = random(6,9,1);
$x = Formula("$a * cos(t)");
$y = Formula("$b * sin(t)");
$z = Formula("$c");

$multians = MultiAnswer($x, $y, $z)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $xstu, $ystu, $zstu ) = @{$student};
      if ( 
           ( (($xstu/$a)**2+($ystu/$b)**2) == 1 ) && 
           ( ($xstu->eval(t=>"0")) == $a ) &&
           ( ($xstu->eval(t=>"pi/2")) == 0 ) &&
           ( ($ystu->eval(t=>"0")) == 0 ) &&
           ( ($ystu->eval(t=>"pi/2")) == $b ) &&
           ( $zstu == $c ) 
         ) {
            return 1;
      } else {
            return 0;
      }
  }
);

Context()->texStrings;
BEGIN_TEXT

Find parametric equations for the quarter-ellipse 
from \( ($a,0,$c) \) to \( (0,$b,$c) \) centered
at \( (0,0,$c) \) in the plane \( z = $c \).  Use 
the interval \( 0 \leq t \leq \pi/2 \).
$BR
$BR
\( x(t) = \) \{$multians->ans_rule(30)\}
$BR
\( y(t) = \) \{$multians->ans_rule(30)\}
$BR
\( z(t) = \) \{$multians->ans_rule(30)\}

END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 0;

ANS($multians->cmp());


;
ENDDOCUMENT();
