# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Write an one-variable linear inequality to solve word problems.
#
# Last updated: Jordan 7/27/13; Kling, 7/24/13; Hughes 7/15/2013, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Inequalities')
## KEYWORDS('inequality','linear','interval notation','subtract','divide')
## DBCCSS('7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "parserNumberWithUnits.pl",
  "contextCurrency.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
);

##############################################

package MultiAnswer;
our $separator = ' ';
package main;

$a=random(20,40,5);
$b=random(2.5,3.5,0.1);
$m=random(30,70,10);
$m=Compute($m);
$mmi = NumberWithUnits($m, 'mi');
$c=$a+$b*$m;

Context("Currency");
Context()->currency->addSymbol("dollars","dollar");
Context()->flags->set(trimTrailingZeros=>1);
$ac = Currency("$a");
$bc = Currency("$b");
$cc = Currency("$c");


Context("Interval");
$interval = Interval("[0,$m]");

$left = Formula("$a+$b x");
$right = Formula("$c");
Context()->strings->add("<"=>{}, ">"=>{}, "="=>{}, "<="=>{}, ">="=>{}, "!="=>{});

$comparitor = String("<=");
$opposite = String(">=");


$multians = MultiAnswer($left, $comparitor, $right)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stuLeft, $stuComp, $stuRight ) = @{$student};
      my ( $corLeft, $corComp, $corRight ) = @{$correct};
      if ( ($corLeft == $stuLeft && $corComp == $stuComp && $corRight == $stuRight) ||
           ($corLeft == $stuRight && $opposite == $stuComp && $corRight == $stuLeft) ) {
          return 1;
      } 
  }
);



##############################################

TEXT(beginproblem());

BEGIN_PGML

You are riding in a taxi and can only pay with cash. You have to pay a flat fee of [`[$ac]`], and then pay [`[$bc]`] per mile. You have a total of [`[$cc]`] in your pocket.

Let [`x`] be the number of miles the taxi will drive you. You want to know how many miles you can afford. Write an inequality to represent this situation in terms of how many miles you can afford: 


[@KeyboardInstructions(q?Your inequality has three parts, a left side, a right side, and a comparison operator in the middle: [`<`], [`>`], [`\leq`], [`\geq`], [`=`], or [`\neq`]. Enter these as [|<|], [|>|]*, [|<=|]*, [|>=|]*, [|=|]*, and [|!=|]*.?)@]**

>> [__]{$multians}{width=>10} [__]{$multians}{width=>2} [__]{$multians}{width=>10} <<

Solve this inequality. At most how many miles can you afford? 

    You can afford at most [___]{$m} miles.

Use interval notation to express the number of miles you can afford.

    [________]{$interval}

END_PGML

##############################################

$showPartialCorrectAnswers=1;



$s1=$c-$a;

BEGIN_PGML_SOLUTION

Let [`x`] be the number of miles the taxi will drive.

    [`
\begin{aligned}
   [$b]x+[$a] & \leq  [$c] \\
   [$b]x+[$a]\mathbf{{}-[$a]} & \leq  [$c]\mathbf{{}-[$a]} \\
   [$b]x      & \leq  [$s1] \\
  \frac{[$b]x}{[$b]} & \leq  \frac{[$s1]}{[$b]} \\
   x          & \leq  [$m]
\end{aligned}
`]

You can afford at most [`[$m]`] miles.

The interval that represents the number of miles that you can afford is [`[$interval]`].

END_PGML_SOLUTION

ENDDOCUMENT();

