# WeBWorK problem written by Carl Yao, 2013
# Portland Community College
#
# Write an one-variable linear inequality to solve word problems.
#
# Last updated: Jordan 8/29/13;  Kling, 7/24/13; Hughes 7/15/2013, Yao, 6/26/13
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Inequalities')
## KEYWORDS('inequality','linear','set notation','interval notation','multiply','subtract')
## DBCCSS('7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
,
  "PGcourse.pl",
);

##############################################
package MultiAnswer;
our $separator = ' ';
package main;

Context("Numeric");

$a=random(70,75,1);
$b=random(85,90,1);
$m=random(76,84,1);
$ans2=Compute($m*3-$a-$b);
while ($ans2==($a+$b)/2) {
   $m=random(76,84,1);
   $ans2=Compute($m*3-$a-$b);
}

Context("Interval");
$interval = Interval("[$ans2,100]");

$left = Formula("($a+$b+x)/3");
$right = Formula("$m");

Context()->strings->add("<"=>{}, ">"=>{}, "="=>{}, "<="=>{}, ">="=>{}, "!="=>{});

$comparitor = String(">=");
$opposite = String("<=");


$multians = MultiAnswer($left, $comparitor, $right)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stuLeft, $stuComp, $stuRight ) = @{$student};
      my ( $corLeft, $corComp, $corRight ) = @{$correct};
      if ( ($corLeft == $stuLeft && $corComp == $stuComp && $corRight == $stuRight) ||
           ($corLeft == $stuRight && $opposite == $stuComp && $corRight == $stuLeft) ) {
          return 1;
      } 
  }
);

##############################################

TEXT(beginproblem());

BEGIN_PGML

[%Not using percentages to avoid confusion%]
Your grade in a class is determined by the average of three test scores. You scored [`[$a]`] and [`[$b]`] on the first two tests. To earn at least [`[$m]`] for this course, how much do you have to score on the third test?

Let [`x`] be the score you will earn on the third test. Write an inequality to represent this situation. 

Your inequality has three parts, a left side, a right side, and a comparison operator in the middle: [`<`], [`>`], [`\leq`], [`\geq`], [`=`], or [`\neq`]. Enter these as *<*, *>*, *<=*, *>=*, *=*, and *!=*.

END_PGML
BEGIN_TEXT
$PAR
$BCENTER 
\{$multians->ans_rule(10)\} \{$multians->ans_rule(2)\} \{$multians->ans_rule(10)\} 
$ECENTER
$PAR
END_TEXT
ANS( $multians->cmp() ); 
BEGIN_PGML

Solve this inequality. What is the minimum that you have to earn on the third test in order to earn a [`[$m]`] for the course? 

    [___________]{$ans2}

You cannot score over [`100`] on the third test. Use interval notation to represent the range of scores you can earn on the third test in order to earn at least [`[$m]`] for this course.

    [______________]{$interval} 

END_PGML

##############################################

$showPartialCorrectAnswers=1;


$s1=3*$m;
$s2=$a+$b;

BEGIN_PGML_SOLUTION

Let [`x`] be the score you will earn on the third test. 

    [`
\begin{aligned}
   \frac{[$a]+[$b]+x}{3} & \geq  [$m] \\
   3 \cdot \frac{[$a]+[$b]+x}{3} & \geq  3 \cdot [$m] \\
   [$a]+[$b]+x & \geq  [$s1] \\
   [$s2] +x & \geq  [$s1] \\
   [$s2] +x \mathbf{{}-[$s2]}& \geq  [$s1]\mathbf{{}-[$s2]} \\
          x  & \geq  [$ans2]
\end{aligned}
`]

You have to score at least [`[$ans2]`] on the third test in order to earn [`[$m]`] in this course.

To earn at least [`[$m]`] in this course, the interval for your score on the third test is [`[$interval]`].

END_PGML_SOLUTION

ENDDOCUMENT();

