## DESCRIPTION
## First order ODEs: separable differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Substitutions)
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('1.3')
## Problem1('4')
## KEYWORDS('differential equations','first order','separable differential equations')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserAssignment.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric")->variables->add(
y=>"Real", k=>"Real"
);
parser::Assignment->Allow;

$a = random(2,9,1);
do { $b = random(2,5,1); } until ($b != $a);

$ab = $a * $b;

$answer = Compute("y = k * e^(x^2/2 + $b x) - $a");


#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Find a solution to 
\( \displaystyle \frac{dy}{dx} = xy + $a x + $b y + $ab \).
If necessary, use \( k \) to denote an arbitrary constant.
$BR
$BR
\{ ans_rule(30) \}
\{ AnswerFormatHelp("equations") \}
END_TEXT
Context()->normalStrings;



##############################
#  Answer evaluation 

$showHint = 2;
BEGIN_HINT
Factor the right side of the equation.
END_HINT

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp( checker => sub {
    my ( $correct, $student, $self ) = @_;
    if ($self->{_filter_name} ne 'produce_equivalence_message') {
      my $stu = Formula($student->{tree}{rop});
      if ($stu->isConstant) {
         Value::Error('Your answer should not be constant');
         return 0;
      }
      my $stu_x = $stu->D('x');
      return $stu_x == Formula("x*$stu + $a*x + $b*$stu + $ab");
    }
  })
);

;

ENDDOCUMENT();
