## DESCRIPTION
## First order ODEs: separable differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('1.3')
## Problem1('1')
## KEYWORDS('differential equations','first order','separable differential equations')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGunion.pl",
  "parserImplicitPlane.pl",
  "parserAssignment.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup 1

Context("ImplicitPlane");
Context()->variables->are(
x=>"Real",y=>"Real"
);

$a = random(2,7,1);
$aa = $a * $a;

$s = random(-1,1,2);
$sa = $s * $a;

$b = random(-7,-2,1);
$bb = $b * $b;

$c = -($b);
$cc = $bb;

$answera = Formula("y^2 - (x/$a)^2"); 

$answerb = ImplicitPlane("$sa y = x");


#############################
#  Main text 1

Context()->texStrings;
BEGIN_TEXT
Solve the differential equation
\( \displaystyle \frac{dy}{dx} = \frac{x}{$aa y} \).

\{ BeginList('OL',type=>'a') \}

$ITEM Find an implicit solution and put your answer 
in the following form:
$BR
\{ ans_rule(20) \} = constant. $SPACE
(Help: \{helpLink("formulas")\}.)

$ITEMSEP
$ITEM Find the equation of the solution through
the point \( (x,y) = ($sa,1) \).  Your equation must describe a single curve \(y=f(x)\) with the domain of \(f\) as large as possible. $SPACE(Help: \{helpLink("equations")\}.)
$BR
\{ ans_rule(20) \}
END_TEXT
Context()->normalStrings;

##############################
#  Answer evaluation 1

$showPartialCorrectAnswers = 1;

ANS( $answera->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $stu_x = $student->D('x');
      my $stu_y = $student->D('y');
      return ( Formula("$aa*y")*$stu_x + Formula("x")*$stu_y == Formula(0) && $stu_x != Formula(0));
    }
) );


ANS( $answerb->cmp() );



####################################
#  Setup 2

Context("Numeric");
Context()->variables->are(
x=>"Real",y=>"Real"
);
parser::Assignment->Allow;

$answerc = Compute("y = -sqrt( (x/$a)^2 + $bb )");


#############################
#  Main text 2

Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM Find the equation of the solution through
the point \( (x,y) = (0,$b) \).  Your answer should 
be of the form \( y = f(x) \). $SPACE(Help: \{helpLink("equations")\}.)
$BR
\{ ans_rule(20) \}


\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


############################
#  Answers 2

ANS( $answerc->cmp() );

;

ENDDOCUMENT();
