# DESCRIPTION
# Asks to find the difference of two rational expressions
# ENDDESCRIPTION

## DBsubject('Intermediate Algebra')
## DBchapter('Ch 06: Rational Expressions, Functions and Equations')
## DBsection('Rational Expressions and Functions: Adding and Subtracting')
## KEYWORDS('least common denominator')
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","z","a","b","c");
$ch = random(0,5,1);
$v = $vars[$ch];
Context()->variables->are($v=>"Real");

$case = random(0,3,1);

if ($case == 0) {
  $b = non_zero_random(-12,12);
  $k = random(2,5);
  $b0 = random(-1,1,2)*$b;
  $e = $k*$b0;
  $frac = "\left(\frac{$v}{$v+$b} - \frac{$v}{$v-$b}\right) \div \frac{$v}{$k $v + $e}";
  if ($b0 == $b) {
    $num = Formula("-2*$b*$k")->reduce;
    $den = Formula("$v - $b")->reduce;
  } else {
    $num = Formula("-2*$b*$k")->reduce;
    $den = Formula("$v + $b")->reduce;
  }
} elsif ($case == 1) {
  $b = non_zero_random(-12,12);
  $k = random(2,5);
  $b0 = random(-1,1,2)*$b;
  $e = $k*$b0;
  $frac = "\frac{$v}{$k $v + $e} \div \left(\frac{$v}{$v+$b} - \frac{$v}{$v-$b}\right)";
  if ($b0 == $b) {
    $den = Formula("-2*$b*$k")->reduce;
    $num = Formula("$v - $b")->reduce;
  } else {
    $den = Formula("-2*$b*$k")->reduce;
    $num = Formula("$v + $b")->reduce;
  }
} elsif ($case == 2) {
  $a = random(1,5);
  $b = non_zero_random(-7,7);
  $b2 = ($b)**2;
  do {$c = random(1,5)} until ($c != abs($b));
  $num = Formula("2*$a*$b*$c")->reduce;
  $den = Formula("$b2 $v - $v^3")->reduce;
  $frac = "\frac{$a}{$v} \cdot \left( \frac{$c}{$v + $b} - \frac{$c}{$v - $b}\right)";
} else {
  do {
    $a = non_zero_random(-7,7);
    $b = random(1,7);
    $c = random(1,7);
    do {$d = non_zero_random(-7,7)} until (abs($d) != abs($a));
  } until (gcd($b+$c,$b*$d) == 1);
  $e1 = $b+$c;
  $e2 = $b*$d;
  $num = Formula("$e1 $v + $e2")->reduce;
  $den = Formula("$v ($v+$a)($v+$d)")->reduce;
  $frac = "\frac{1}{$v+$a} \cdot \left( \frac{$b}{$v} + \frac{$c}{$v+$d}\right)";
}


$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Both answer blanks must be filled");
          $self->setMessage(2,"Both answer blanks must be filled");
          return [0,0];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Perform the indicated operation, and simplify if possible.  Assume any factors you cancel are not zero.
$BR
$BR
$BCENTER
$displayfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    \(\dfrac{$num}{$den} \)
    $PAR
    END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
