##DESCRIPTION
## Find limits using the squeeze theorem
##ENDDESCRIPTION

## Shotwell cleaned
## tcao , PAID on 11-24-2003

## DBsubject(Calculus - single variable)
## DBchapter(Limits and continuity)
## DBsection(Squeeze theorem)
## Date(6/3/2002)
## MLT(limits_squeeze_linear_and_quadratic)
## MLTleader(1)
## Level(3)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('6')
## Section1('2.3')
## Problem1('33')
## KEYWORDS('limits', 'sequences', 'squeeze theorem')
## MO(1)-TF-12/18/2014

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric");



$a = random(-4,4,1);
$b = random(-6,6,1);
$c = random(-6,6,1);
# coefficients of quadratic  x^2 + a1x + a2 passing through c at a
$a1 = $b-$a;
$a2 = -$a*$b+$c;
# coefficients of line b1x +b2  passing through c at a and tangent to quadratic above.
$b1 = $a + $b;
$b2 = -$a*($a +$b) + $c;

$pol = nicestring([1, $a1, $a2]);
$lhs = nicestring([$b1, $b2]);
$showPartialCorrectAnswers=1;

BEGIN_TEXT
Suppose \[ $lhs  \le f(x) \le $pol \] $BR 

Use this to compute the following limit. $BR

\[ \lim_{x\to $a} f(x) \]

Answer: \{ans_rule(25)\} $BR$BR

What theorem did you use to arrive at your answer? $BR
Answer: \{ ans_rule(50) \}

                  	                  
END_TEXT

Context('Numeric');
ANS(Real($c)->cmp);

Context("ArbitraryString");

#ANS(Real(1)->cmp);

ANS(Compute("The Squeeze Theorem")->cmp(checker=>sub {
  my ($correct,$student,$ans) = @_;
  $correct = $correct->value; # get perl string from String object
  $student = $student->value; # ditto
  $student =~ s/ //g;
  $student = lc($student);
  return 1 if($student =~ /squeeze/ or $student =~ /pinch/  or $student =~/sandwich/);
  return 0;
}));


ENDDOCUMENT();        # This should be the last executable line in the problem.
