# WeBWorK problem written by Carl Yao
# Portland Community College
#
# A speed/distance related application problem involving solving a rational equation.
# 
# Last edited: Yao 08/30/15
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','equation','application','rational')
## DBCCSS('6.EE.7','7.EE.4','8.EE.7')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

##############################################

Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');
Context()->flags->set(reduceConstants=>0, reduceConstantFunctions=>0, formatStudentAnswer=>parsed);

do {
   do {
      $a = random(200,1000,10);
      $b = random(200,1000,10);
      $c = random(5,30,1);
      $d = random(2,20,1);
      $a1 = $d;
      $b1 = $b-$a;
      $c1 = $b*$c-$d*$c**2+$a*$c;
      $d1 = ($b1)**2-4*$a1*$c1;
   } until ($d1>0);
   $x1 = (-$b1+sqrt($d1))/(2*$a1);
   $x2 = (-$b1-sqrt($d1))/(2*$a1);
} until ( ($x1==int($x1)) && ($x1>$c) && ($x2>$c) );

$commonD = Formula("(x+$c)*(x-$c)")->reduce;

@answers = ($x1,$x2);

$multians = MultiAnswer( @answers )->with(
  singleResult => 0, allowBlankAnswers=>1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @sAns = @{$student};
      if ( ($sAns[0]==$x1) && ($sAns[1]==$x2) ) {
         return [1,1];
      } elsif ( ($sAns[1]==$x1) && ($sAns[0]==$x2) ) {
         return [1,1];
      } elsif ( ($sAns[0]==$x1) || ($sAns[0]==$x2) ) {
         return [1,0];
      } elsif ( ($sAns[1]==$x1) || ($sAns[1]==$x2) ) {
         return [0,1];
      } else {
         return [0,0];
      }
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

A river flows at [`[$c]`] miles per hour. A boat traveled with the current from Town A to Town B, which are [`[$a]`] miles apart. Then, the boat turned around, and traveled against the current to reach Town C, which is [`[$b]`] miles away from Town B. The second leg of the trip (Town B to Town C) took [`[$d]`] hours longer than the first leg (Town A to Town B). During this whole trip, the boat was driving at a constant still-water speed. Find this speed. Note that there are two possible answers:

During this trip, the boat's speed on still water was [______]{$multians} miles per hour, or [______]{$multians} miles per hour.

END_PGML

##############################################

$s1 = -$b1;
$s2 = sqrt($d1);

##############################################


BEGIN_PGML_SOLUTION

Assume that the boat's still-water speed was [`x`] miles per hour.

Since the current's speed was [`[$c]`] miles per hour, the boat traveled from Town A  to Town B (with the current) at [`x+[$c]`] miles per hour, and from Town B to Town C (against the current) at [`x-[$c]`] miles per hour.

Using the formula [`time=\frac{distance}{speed}`], we know the boat traveled from Town A to Town B in [`\frac{[$a]}{x+[$c]}`] hours; and it traveled from Town B to Town C in [`\frac{[$b]}{x-[$c]}`] hours.

It's given that the second leg of the trip took [`[$c]`] hours longer, so we can write an equation and solve it:

    [``
\begin{aligned}
  &\frac{[$a]}{x+[$c]}=\frac{[$b]}{x-[$c]} + [$d] \\
  &\frac{[$a]}{x+[$c]} \cdot [$commonD] = \frac{[$b]}{x-[$c]} \cdot [$commonD] + [$d] \cdot [$commonD] \\
  &[$a](x-[$c]) = [$b](x+[$c]) +[$d](x+[$c])(x-[$c]) \\
  &[$a]x-[$a*$c] = [$b]x+[$b*$c] +[$d](x^{2}-[$c**2]) \\
  &[$a]x-[$a*$c] = [$b]x+[$b*$c] +[$d] x^{2}-[$d*$c**2] \\
  & 0= [$b]x+[$b*$c] +[$d] x^{2}-[$d*$c**2] -[$a]x +[$a*$c] \\
  & 0= [$d]x^{2}+[$b-$a]x +[$c1]
\end{aligned}
    ``]

When numbers are big, we use the quadratic formula to solve the quadratic equation:

    [``
\begin{aligned}
   x &= \frac{-b \pm \sqrt{b^{2}-4ac}}{2a} \\
   x &= \frac{-([$b1]) \pm \sqrt{([$b1])^{2}-4 \cdot ([$a1]) \cdot ([$c1])}}{2 \cdot [$a1]} \\
   x &= \frac{[$s1] \pm \sqrt{[$d1]}}{[$a1+$a1]} \\
   x &= \frac{[$s1] \pm [$s2]}{[$a1+$a1]} \\
   x &= \frac{[$s1] + [$s2]}{[$a1+$a1]} \:\: or \:\: x= \frac{[$s1] - [$s2]}{[$a1+$a1]} \\
   x &= [$x1] \:\: or \:\: x= [$x2]
\end{aligned}
    ``]

The boat's still-water speed was [`[$x1]`] miles per hour or [`[$x2]`] miles per hour.

END_PGML_SOLUTION


ENDDOCUMENT();
