# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Solve equations like a/(x+b)+cx/(x+d)=e/( (x+c)(x+d) )
# 
# There are two solutions.
#
# Last edited: Yao 08/30/15
# ENDDESCRIPTION

## DBCCSS('6.EE.7','7.EE.4','8.EE.7')
## DBsubject(Algebra)
## DBchapter(Basic Algebra)
## DBsection(Algebraic Expressions)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('solve','equation','rational')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserAssignment.pl",
  "PCCmacros.pl",
  "answerHints.pl",
  "contextFraction.pl",
  "SolveLinearEquationPCC.pl",
  "PGcourse.pl"
);

##############################################

Context("Fraction")->flags->set(
            reduceFractions => 0,
            showMixedNumbers=>0,
            showExtraParens=>0);
Context()->noreduce('(-x)-y','(-x)+y');
parser::Assignment->Allow;
Context()->operators->redefine(',',using=>',',from=>'Numeric');
Context()->operators->redefine('or',using=>',',from=>'Numeric');
Context()->operators->set(
  ','=>{string=>' or ',TeX=>'\hbox{ or }'},
  'or'=>{string=>' or ',TeX=>'\hbox{ or }'}
);
Context()->lists->set(List => {separator => " or "});
Context()->{error}{msg}{"Function 'sqrt' is not allowed in this context"} 
  = "Please simplify your answer further";

$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

# add solution strings to context- this means that if
# students enter these (and they are not correct), then 
# WW will not give a Context warning
Context()->strings->add("no real solutions"=>{},
  "no real solution"=>{alias=>'no real solutions'}, 
  "none"=>{alias=>'no real solutions'}, 
  "no solution"=>{alias=>'no real solutions'}, 
  "no solutions"=>{alias=>'no real solutions'}, 
  );


do {
   do {
      $a = non_zero_random(-9,9,1);
      $b = non_zero_random(-9,9,1);
      $c = non_zero_random(2,9,1);
      do {$d = non_zero_random(-9,9,1);} until ($b!=$d);
      $e = non_zero_random(-9,9,1);
      $a1 = $c;
      $b1 = $a+$b*$c;
      $c1 = $a*$d-$e;
      $d1 = ($b1)**2-4*$a1*$c1;
      $temp1 = gcf(abs($a1),abs($b1));
      $comD = gcf($temp1,abs($c1));
   } until ( ($c1!=0) && ($d1>0) && ($comD!=1) );
   $x2 = ( sqrt($d1)==int(sqrt($d1))) ? Fraction(-$b1+sqrt($d1),(2*$a1))->reduce : -100;
   $x1 = ( sqrt($d1)==int(sqrt($d1))) ? Fraction(-$b1-sqrt($d1),(2*$a1))->reduce : -100;
} until ( ($x1!=-$b) && ($x2!=-$b) && ($x1!=-$d) && ($x2!=-$d) && ($x1!=$x2) );

($numX1,$denX1) = $x1->value;
($numX2,$denX2) = $x2->value;

$commonD = Formula("($var+$b)($var+$d)")->reduce;
$left = Formula("$a/($var+$b)+$c*$var/($var+$d)")->reduce;
$bPlusD = $b+$d;
$bd = $b*$d;
$right = Formula("$e/($var**2+$bPlusD*$var+$bd)")->reduce;

$ans = Compute("$var = $x2, $var = $x1");


##############################################

TEXT(beginproblem());

BEGIN_PGML

Solve the following rational equation; the answer could be in the form *[`[$var]`] = ___ *, *no solution*, or *all real numbers*. 

    [`` [$left]=[$right]  ``]  

    [_______________________________]

END_PGML

##############################################

$showPartialCorrectAnswers = 1;
ANS($ans->cmp(
  entry_type => "a solution",
  checker => sub {
    my ($correct,$student,$ans,$nth,$value) = @_;
    if ($correct->type eq "Assignment") {
      my ($svar,$sfrac) = $student->value; # get the variable and fraction
      #return 0 unless Value::classMatch($sfrac,'Fraction') && $sfrac->isReduced;
      if(Value::classMatch($sfrac,'Fraction'))
      {
        return 0 unless $sfrac->isReduced;
      }
    }
    return $correct == $student;
  },
  extra => sub {
    my ($student,$ansHash,$nth,$value) = @_;
    if($student eq "no real solutions")
    {
         $student->context->setError("This equation does have some solutions- try using the square root method","",undef,undef,$Value::CMP_WARNING)
         unless $ans->{isPreview};
         return;
    }
    if ($student->type ne "Assignment" && $ansHash->{student_formula}->type ne "Assignment") {
      $student->context->setError("Your $nth solution should be written $var = $US$US$US","",undef,undef,$Value::CMP_WARNING)
         unless $ans->{isPreview};
      return;
    }
    my ($svar,$sfrac) = $student->value; # get the variable and fraction
    if (Value::classMatch($sfrac,'Fraction') && !$sfrac->isReduced) {
      $student->context->setError("Your $nth $value is not reduced","",undef,undef,$Value::CMP_WARNING)
         unless $ans->{isPreview};
      return;
    }
    return Value::Real->typeMatch($student);
  }
)->withPostFilter(AnswerHints(
  ["$var=$a","$var=-$a"] => "Are you sure you have all the solutions?",
  [$a,-$a] => ["Your solution is a correct one, but you should write $var = $US$US$US<br>Are you sure you have all the solutions?",replaceMessage=>1],
  ["$a,-$a","-$a,$a"] => ["Your solutions are correct, but you should write $var = $US$US$US",replaceMessage=>1],
)));

##############################################

$s1 = Formula("$e/( ($var+$b)($var+$d) )")->reduce;
$s2 = Formula("$a/($var+$b)*$commonD+$c*$var/($var+$d)*$commonD")->reduce;
$s3 = Formula("$e/( ($var+$b)($var+$d) )*$commonD")->reduce;
$ad = $a*$d;
$bc = $b*$c;
$s4 = Formula("$a*($var+$d)+$c*$var*($var+$b)")->reduce;
$s5 = $e;
$s6 = Formula("$a*$var+$ad+$c*$var**2+$bc*$var")->reduce;
$o = $a+$bc;
$s7 = Formula("$c*$var**2+$o*$var+$ad")->reduce;
$p = $ad-$e;
$s8 = Formula("$c*$var**2+$o*$var+$p")->reduce;
$r = $c/$comD;
$s = $o/$comD;
$t = $p/$comD;
$s9 = Formula("$comD*($r*$var**2+$s*$var+$t)")->reduce;
$s10 = Formula("$denX1*$var-$numX1")->reduce;
$s11 = Formula("$denX2*$var-$numX2")->reduce;

$nb = -$b;
$nd = -$d;


#############################################

BEGIN_PGML_SOLUTION

To clear fractions in an equation, multiply each side of the equation by a common denominator. For this problem, a common denominator is [`[$commonD]`].

    [``
\begin{aligned}
  &[$left] = [$right] \\
  &[$left] = [$s1] \\
  &[$s2] = [$s3] \\
  &[$s4] = [$s5] \\
  &[$s6] = [$s5] \\
  &[$s7] = [$s5] \\
  &[$s8] = 0 \\
  &[$s9] = 0 \\
  &[$com]([$s10])([$s11]) = 0 \\
   &[$s10]=0 \:\: or\:\: [$s11]=0 \\
   &[$var]=[$x1] \:\: or\:\: [$var]=[$x2]
\end{aligned}
    ``]

We must check solutions when we solve rational equations. In the original equation, both denominators cannot be [`0`]. We set them to [`0`], so we can find out what values to exclude from the solution set.

For this example, it's easy to see we must exclude [`[$nb]`] and [`[$nd]`] from the solution set. There is no conflict with those two solutions we got, so the final solution is [`[$var]=[$x1]`] and [`[$var]=[$x2]`].

END_PGML_SOLUTION


ENDDOCUMENT();
