## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Kernel and image)
## Date(July 2013)
## Institution(TCNJ and Hope College)
## Author(Paul Pearson)
## MLT(find_x_in_preimage2x2)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.1')
## Problem1('')
## KEYWORDS('vector space', 'linear transformation' 'matrix' 'image')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

install_problem_grader(~~&std_problem_grader);

Context('Matrix');
Context()->flags->set(
  tolerance => 0.001,
  tolType => "absolute",
);

do {

$a11 = random(2,6,2) * random(-1,1,2);
$a12 = random(1,5,2) * random(-1,1,2);
$e = non_zero_random(-4,4,1);
$f = non_zero_random(-4,4,1);
$a21 = $e * $a11;
$a22 = $e * $a12;

$u1 = $f* $a11;
$u2 = $f* $a21;

$A = Matrix([
[$a11, $a12], 
[$a21, $a22],
]);

} until ($A->det == 0);

$u = Matrix([[$u1], [$u2]]);

$b = $A * $u;

$popup = PopUp(['choose','unique','not unique'],'not unique');

Context()->texStrings;
BEGIN_TEXT
Let \[ A = $A \ \mbox{ and } \ \vec{b} = $b. \]
Define the linear transformation \( T: {\mathbb R}^2 \rightarrow {\mathbb R}^2\) by \(T(\vec{x}) = A\vec{x}\).  
Find a vector \(\vec{x}\) whose image under \(T\) is \(\vec{b}\).
$BR
$BR
\( \vec{x} = \) \{ $u->ans_array \}.
$BR
Is the vector \(\vec{x}\) unique? \{ $popup->menu \}
END_TEXT
Context()->normalStrings;

ANS($u->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
$s = Matrix($student);
return $A * $s == $b;
}));
ANS($popup->cmp);

COMMENT('Solution is never unique.');
ENDDOCUMENT();
