## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Kernel and image)
## Date(July 2013)
## Institution(TCNJ and Hope College)
## Author(Paul Pearson)
## MLT(find_x_in_preimage)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.1')
## Problem1('')
## KEYWORDS('vector space', 'linear transformation' 'matrix' 'image')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

install_problem_grader(~~&std_problem_grader);

Context('Matrix');
Context()->flags->set(
  tolerance => 0.001,
  tolType => "absolute",
);

do {

$a11=1;
$a12 = non_zero_random(-3,3,1);
$a13 = non_zero_random(-3,3,1);
$a21=0;
$a22=1;
$a23 = non_zero_random(-3,3,1);

$m1 = non_zero_random(-3,3,1);
$m2 = non_zero_random(-3,3,1);
$m3 = non_zero_random(-3,3,1);

$a31=$m1;
$a32=$m1*$a12 +$m2;
$a33=$m1*$a13 + $m2*$a23;

$A = Matrix([
[$a11, $a12, $a13], 
[$a21, $a22, $a23],
[$a31, $a32, $a33],
]);

} until ($A->det == 0);

$u = Matrix([[$m1], [$m2],[$m3]]);

$b = $A * $u;

$popup = PopUp(['choose','unique','not unique'],'not unique');

Context()->texStrings;
BEGIN_TEXT
Let \[ A = $A \ \mbox{ and } \ \vec{b} = $b. \]
Define the linear transformation \( T: {\mathbb R}^3 \rightarrow {\mathbb R}^3\) by \(T(\vec{x}) = A\vec{x}\).  
Find a vector \(\vec{x}\) whose image under \(T\) is \(\vec{b}\).
$BR
$BR
\( \vec{x} = \) \{ $u->ans_array \}.
$BR
Is the vector \(\vec{x}\) unique? \{ $popup->menu \}
END_TEXT
Context()->normalStrings;

ANS($u->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
$s = Matrix($student);
return $A * $s == $b;
}));
ANS($popup->cmp);

COMMENT('Solution is never unique.');
ENDDOCUMENT();
