# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 08/18/10

## DBsubject(Calculus - multivariable)
## DBchapter(Fundamental theorems)
## DBsection(Stokes' theorem)
## Date(08/18/10)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(Stokes_verify)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('20.4')
## Section2('20.2')
## Section3('20.2')
## Problem1('1')
## Problem2('3')
## Problem3('6')
## KEYWORDS('stokes theorem', 'circulation integral', 'curl', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Calculus of Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Vector");
Context()->variables->add( t=>'Real' );
Context()->flags->set( ijk => 1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"},
);
$showPartialCorrectAnswers = 1;

$a = random(1,8,1);
$b = random(-8,-1,1);
$c = random(1,8,1);
$ffunc = Compute( "$a*x + $b*y + $c*z" )->reduce();

$x0 = random(2,8,1);
$y0 = random(2,8,1);

$r1 = Vector( Compute( "$x0*t" ), Formula( "0" ), Formula( "0") );
$r2 = Vector( Formula( "$x0" ), Compute( "$y0*t" ), Formula( "0") );
$r3 = Vector( Compute( "$x0*(1-t)" ), Compute( "$y0*(1-t)" ), Formula( "0") );

$cint1 = Compute( "$a*$x0^2/2" );
$cint2 = Compute( "$a*$x0*$y0 + $b*$y0^2/2" );
$cint3 = Compute( "-1*($a*$x0/2 + $b*$y0/2)*($x0 + $y0)" );
$cint  = Compute( "($a - $b)*$x0*$y0/2" );

$curlF = Vector( Formula("-1*$c"), Formula("$c"), Formula("$a - $b") );
$xlim0 = Compute( "0" );
$xlim1 = Compute( "$x0" );
$ylim0 = Formula( "0" );
$ylim1 = Compute( "$y0*x/$x0" );
$integ = Formula( "$a - $b" );
$sint  = Compute( "($a - $b)*$x0*$y0/2" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Calculate the circulation, \( \int_C\vec F\cdot d\vec r \), in two ways,
directly and using Stokes' Theorem.  The vector field
\( \vec F = ($ffunc)(\vec i+\vec j) \) and  \( C \) is the triangle with
vertices \( (0,0,0) \), \( ($x0,0,0) \), \( ($x0,$y0,0) \), traversed in
that order.
$PAR
Calculating directly, we break \(C\) into three paths.  For each, give
a parameterization \(\vec r\,(t)\) that traverses the path from start
to end for \(0\le t\le1\).
$BR
On \(C_1\) from \((0,0,0)\) to \(($x0,0,0)\),
\(\vec r(t) = \) \{ ans_rule(25) \}
$BR
On \(C_2\) from \(($x0,0,0)\) to \(($x0,$y0,0)\),
\(\vec r(t) = \) \{ ans_rule(25) \}
$BR
On \(C_3\) from \(($x0,$y0,0)\) to \((0,0,0)\),
\(\vec r(t) = \) \{ ans_rule(25) \}
$PAR
So that, integrating, we have
\(\int_{C_1}\,\vec F\cdot d\vec r = \) \{ ans_rule(25) \}
$BR
\(\int_{C_2}\,\vec F\cdot d\vec r = \) \{ ans_rule(25) \}
$BR
\(\int_{C_3}\,\vec F\cdot d\vec r = \) \{ ans_rule(25) \}
$BR
and so 
\(\int_C\vec F\cdot d\vec r = \) \{ ans_rule(25) \}.

$PAR
Using Stokes' Theorem, we have
$BR
\( \mbox{curl }\vec F = \) \{ ans_rule(25) \}
$BR
So that the surface integral on \(S\), the triangular region on the
plane enclosed by the indicated triangle, is
$BR
\(\int_S\mbox{curl }\vec F\cdot d\vec A =
\int_a^b\int_c^d \) \{ ans_rule(25) \} \( dy\,dx \),
$BR
where \(a = \) \{ ans_rule(5) \}, \(b = \) \{ ans_rule(5) \},
\(c = \) \{ ans_rule(5) \}, and \(d = \) \{ ans_rule(5) \}.
$BR
Integrating, we get
\(\int_C\vec F\cdot d\vec r =
  \int_S\mbox{curl }\vec F\cdot d\vec A = \) \{ ans_rule(15) \}

END_TEXT
Context()->normalStrings;

ANS($r1->cmp( checker=>sub{
    my ( $c, $s, $ans ) = @_;
    if ( $c == $s ) {
	return 1;
    } else {
	$s0 = $s->eval( t=>0 );
        $s1 = $s->eval( t=>1 );
        return ( $s0 == Vector( Compute("0"), Compute("0"), Compute("0") ) &&
		 $s1 == Vector( Compute("$x0"), Compute("0"), Compute("0") ) ) ?
		   1 : 0; 
    }
} ) );
ANS($r2->cmp( checker=>sub{
    my ( $c, $s, $ans ) = @_;
    if ( $c == $s ) {
	return 1;
    } else {
	$s0 = $s->eval( t=>0 );
        $s1 = $s->eval( t=>1 );
        return ( $s0 == Vector( Compute("$x0"), Compute("0"), Compute("0") ) &&
		 $s1 == Vector( Compute("$x0"), Compute("$y0"), Compute("0"))) ?
		   1 : 0; 
    }
} ) );
ANS($r3->cmp( checker=>sub{
    my ( $c, $s, $ans ) = @_;
    if ( $c == $s ) {
	return 1;
    } else {
	$s0 = $s->eval( t=>0 );
        $s1 = $s->eval( t=>1 );
        return ( $s0 == Vector( Compute("$x0"), Compute("$y0"), Compute("0")) &&
		 $s1 == Vector( Compute("0"), Compute("0"), Compute("0") ) ) ?
		   1 : 0; 
    }
} ) );
ANS($cint1->cmp() );
ANS($cint2->cmp() );
ANS($cint3->cmp() );
ANS($cint->cmp() );
ANS($curlF->cmp() );
ANS($integ->cmp() );
ANS($xlim0->cmp() );
ANS($xlim1->cmp() );
ANS($ylim0->cmp() );
ANS($ylim1->cmp() );
ANS($sint->cmp() );

$ax0 = $a*$x0;
$ax0sq = $ax0*$x0;
$cint1d = ( $ax0sq/2 == int($ax0sq/2) ) ? $ax0sq/2 : "\frac{$ax0sq}{2}";
$by0 = -1*$b*$y0;
$by0sq = $by0*$y0;
$cint2d = ( $by0sq/2 == int($by0sq/2) ) ? $ax0*$y0-$by0sq/2 :
    "\frac{" . ($ax0*$y0*2 - $by0sq) . "}{2}";
$bxy0 = $by0*$x0;
$axy0 = $ax0*$y0;
$n = -1*$ax0sq - $axy0 + $bxy0 + $by0sq;
$cint3d = ( $n/2 == int($n/2) ) ? $n/2 : "\frac{$n}{2}";

$n = ($a-$b)*$x0*$y0;
$cintd = ( $n/2 == int($n/2) ) ? $n/2 : "\frac{$n}{2}";

$amb = $a - $b;

($n, $d) = reduce( $y0, $x0 );
$y1limd = ( $n == 1 && $d == 1 ) ? 'x' : ( $d == 1 ? "$n x" : "$n x/$d");

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

To calculate \(\int_C\vec F\cdot d\vec r\) directly, we compute the
integral along each of the sides of the triangle, as indicated.
On \(C_1\), the easiest indicated parameterization is
\[ \vec r_1(t) = $x0 t\vec i. \]
On \(C_2\) and \(C_3\), we similarly can take
\[ \vec r_2(t) = $x0\vec i + $y0 t\vec j \quad\mbox{and}\quad
   \vec r_3(t) = $x0(1 - t)\vec i + $y0(1 - t)\vec j. \]
Thus
\[ \int_{C_1}\vec F\cdot d\vec r =
   \int_0^1 $ax0 t(\vec i + \vec j)\cdot($x0\vec i)\,dt =
   \int_0^1 $ax0sq\,t\,dt = $cint1d. \]
Similarly,
\[ \int_{C_2}\vec F\cdot d\vec r =
   \int_0^1 ($ax0 - $by0 t)(\vec i + \vec j)\cdot( $y0\vec j)\,dt =
   \int_0^1 $ax0 - $by0sq\,t\,dt = $cint2d, \]
and
\[ \int_{C_3}\vec F\cdot d\vec r =
   \int_0^1 ($ax0(1 - t) - $by0(1 - t))(\vec i + \vec j)\cdot
	(-$x0\vec i - $y0\vec j)\, dt = $cint3d. \]
Then
\[ \int_C\vec F\cdot d\vec r = \int_{C_1}\vec F\cdot d\vec r +
   \int_{C_2}\vec F\cdot d\vec r + \int_{C_3}\vec F\cdot d\vec r = $cintd. \]

$PAR
To calculate \(\int_C\vec F\cdot d\vec r\) using Stokes' Theorem, we
find
\[
\mbox{curl }\vec F = \left|\begin{array}{ccc}
\vec i & \vec j & \vec k\\
\frac{\partial}{\partial x} & \frac{\partial}{\partial y} &
  \frac{\partial}{\partial z}\\
$ffunc & $ffunc & 0
\end{array}\right| = -$c\vec i+ $c\vec j+ $amb\vec k.
\]
For Stokes' Theorem, the triangular region \(S\) in
contained by the triangle is oriented upward, so \(d\vec A=\vec
k\,dx\,dy\).  Thus
\[
\int_C\vec F\cdot d\vec r = \int_S\mbox{curl }\vec F\cdot d\vec A =
\int_S(-$c\vec i+$c\vec j+$amb\vec k)\cdot \vec k\, dx\, dy
= \int_0^{$x0}\int_0^{$y1limd}\,$amb\,dx\,dy = $cintd.
\]

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
