## DESCRIPTION
## Compute a basis for the null space of a 3x3 matrix.
##   All matrix entries will be small real integers. The null space will be 2-dimensional, so 2 basis vectors are expected.
##   There will be a solution (i.e., a basis) with small integer entries.
##   There are infinitely many different solutions (i.e., bases). They are all accepted.
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Row, column, and null spaces)
## Date(March 22, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('2.4')
## Problem1('')
## KEYWORDS('linear algebra','matrix','null space','basis')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Matrix");

do {
    $c1 = random(-2,3);
    $c2 = random(-2,3);
    $c3 = random(-2,3);
    $c4 = random(-2,3);
    $c5 = random(-2,3);
    $c6 = random(-2,3);
    $c7 = random(-2,3);
    $c8 = random(-2,3);
    $c9 = random(-2,3);
    $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
    
    $r1 = random(-2,4);
    $r2 = random(-2,4);
    $R = Matrix([[1,$r1,$r2],[0,0,0],[0,0,0]]);
    
    $A = $C * $R;
} until (abs($C->det) > 0.1);

Context("Vector");
Context()->lists->set(Vector => {open => "(", close => ")"});
Context()->parens->set("(" => {type => "Vector", close => ")"});
$v1 = Vector(-$r1,1,0);
$v2 = Vector(-$r2,0,1);

$multians = MultiAnswer($v1,$v2)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if((Vector($A * $s0) == Vector(0,0,0)) && (Vector($A * $s1) == Vector(0,0,0)) && $s0->norm > 0.001 && $s1->norm > 0.001) {
            return 1 - ($s0->isParallel($s1));
        } else {
            return 0;
        }
    }
);

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Compute a basis of the null space of the matrix \( \displaystyle $A \).
$BR
$BR
\( \Bigg\lbrace \)\{$multians->ans_array(3)\}, \{$multians->ans_array(3)\}\( \Bigg\rbrace \).
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
