# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Given the slope and a point, find a line's equation in slope-intercept form.
# Application problem. Slope is negative.
# 
# Last updated: Jordan 8/31/13; Carl Yao, 7/9/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Lines')
## KEYWORDS('linear','equation','slope','slope-intercept','application')
## DBCCSS('F-LE.2','8.F.3','A-CED','S-ID.7')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserAssignment.pl",
  "PCCmacros.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
Context()->variables->add(y=>'Real');
Context()->noreduce('(-x)-y','(-x)+y');
Context()->flags->set(showExtraParens=>0);

$x4=random(40,50,1);
$y4=0;

$m= -random(11,99,1)/10;
$absM = abs($m);
$b= -$m*$x4;

$x1=random(5,10,1);
$y1=$m*$x1+$b;
$x2=random(11,19,1);
$y2=$m*$x2+$b;
$x3=random(30,40,1);
$y3=$m*$x3+$b;

parser::Assignment->Allow;
$ansSI=Formula("y=$m*x+$b")->reduce;

Context("Numeric");
parser::Assignment->Allow;
Context()->variables->add(y=>'Real');
Context()->flags->set(tolType=>'absolute');
Context()->variables->set(x=>{limits=>[$x1,$x2]});
Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
);

$ansPS = Formula("y=$m*(x-$x1)+$y1");
$ansPS_evaluator = $ansPS -> cmp(
      checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is a line equation for the right line, but it is not in the expected point-slope form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
      return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Scientists are conducting an experiment with a gas in a sealed container. The mass of the gas is measured, and the scientists realize that the gas is leaking over time in a linear way. Each minute, they lose [`[$absM]`] grams. [@numberWord($x1,capital=>1)@] minutes since the experiment started, the remaining gas had a mass of [`[$y1]`] grams.

Let [`x`] be the number of minutes that have passed since the experiment started, and let [`y`] be the mass of the gas in grams at that moment. Use a linear equation to model the weight of the gas over time.
a) A point-slope equation to model this is [______________]{$ansPS_evaluator}.

b)  [`[$x3]`] minutes after the experiment started, there would be [___________]{$y3} grams of gas left.

c)  If a linear model continues to be accurate, [___________]{$x4} minutes since the experiment started, all gas in the container will be gone.

END_PGML



##############################################

$s1=$y2-$y1;
$s2=$x2-$x1;
$s3=$m*$x1;
$s4=-$s3;
$s5=$m*$x3;
$s6=$y4-$b;

$outputy1 = $y1<0 ? "($y1)" : $y1;
$outputx1 = $x1<0 ? "($x1)" : $x1;

BEGIN_PGML_SOLUTION

a)    A line's equation in point-slope form looks like [` y=m(x-x_0)+y_0 `], where [`m`] is the slope, and [`(x_0,y_0)`] is some point we know to be on the line. We first need to find the line's slope. Since the gas is leaking at a rate of [`[$absM]`] grams per minute, we know that the slope of the linear model is [`[$m]`].

    So now we have that [` y=[$m](x-x_0)+y_0 `]. We can use [`(x_0,y_0)=([$x1],[$y1])`] to get [`[$ansPS]`].

b)    After [`[$x3]`] minutes, we can find the mass of the remaining gas if we substitute [`[$x3]`] in for [`x`] in the equation [` y=[$m](x-[$x1])+[$y1] `], and we have:

    [``
\begin{aligned}
  y &= [$m]([$x3]-[$x1]) +[$y1] \\
  y &= [$y3]
\end{aligned}
``]

    So after [`[$x3]`] minutes, there are [`[$y3]`] grams of gas remaining.

c)    To find when the gas will be all gone, we need to find when there are  [`[$y4]`] grams of gas left.  We can substitute [`[$y4]`] in for [`y`] in the equation [` y=[$m](x-[$x1])+[$y1] `], and we have:

    [``
\begin{aligned}
  y &= [$m](x-[$x1]) +[$y1] \\
  [$y4] &= [$m](x-[$x1]) +[$y1] \\
  [@$y4-$y1@] &= [$m](x-[$x1])\\
  \frac{[@$y4-$y1@]}{[$m]} &= x - [$x1]\\
  [@($y4-$y1)/$m@] &= x - [$x1]\\
  [@($y4-$y1)/$m+$x1@] &= x\\
  [$x4] &= x
\end{aligned}
``]

    So after [`[$x4]`] minutes, the gas will be all gone.

END_PGML_SOLUTION

ENDDOCUMENT();



