# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/20/17
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','application','subtract','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextCurrency.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(b=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$original = random(10,100,10);
$pNum = random(10,20,1);
$p = $pNum/100;
$new = $original*(1+$p);

$ansString = "b+$p*b=$new";
$ans = Compute("$ansString");

Context("Currency");
$newC = Currency($new);


##############################################

TEXT(beginproblem());

BEGIN_PGML

The price of a restaurant bill, including an [`[$pNum]%`] gratuity charge, was [`[$newC]`]. What was the price of the bill before gratuity was added?

Assume the bill without gratuity is [`b`] dollars. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

Assume the bill without gratuity is [`b`] dollars. A [`[$pNum]%`] gratuity is [`[$p]b`] dollars. Since the original bill plus gratuity is [`[$newC]`], we can write the equation:

    [``b+[$p]b=[$new]``]

END_PGML_SOLUTION

ENDDOCUMENT();
