# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/20/17
# ENDDESCRIPTION
 
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('solve','linear','equation','application','subtract','divide')
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextTypeset.pl",
  "contextCurrency.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(m=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$base = random(20,40,0.05);
$rate = random(0.5,0.9,0.01);
$mile = random(50,200,1);
$total = $base+$rate*$mile;

$ansString = "$rate*m+$base=$total";
$ans = Compute("$ansString");

Context("Currency");
$baseC = Currency($base);
$rateC = Currency($rate);
$totalC = Currency($total);

##############################################

TEXT(beginproblem());

BEGIN_PGML

Uhaul charges an initial fee of [`[$baseC]`] and then [`[$rateC]`] per mile to rent a [`15`]-foot truck for a day. If the total bill is [`[$totalC]`], how many miles were driven?

Assume [`m`] miles were driven. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

A rental bill at this company consists of the base fee, [`[$baseC]`], and the cost of driving a certain number of miles. Assume [`m`] miles were driven. Since the company charges [`[$rateC]`] per mile, the cost of driving [`m`] miles is [`[$rate]m`] dollars. Now we can write the equation:

    [``[$rate]m+[$base]=[$total]``]

END_PGML_SOLUTION

ENDDOCUMENT();
