# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Template:
# Simplify the rational expression like (2x+5)(2x-5)/(2x+5)(x-2) / 2(2x-5)
#
# Last updated: Carl Yao 08/14/2015
# ENDDESCRIPTION

## DBCC('A-APR.1','A-SSE.2')
## DBsubject(Algebra)
## DBchapter(Basic Algebra)
## DBsection(Exponents', 'Factoring)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('polynomial','divide','factoring','simplify','rational')

##############################################
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextLimitedFactor.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);

##############################################

Context("LimitedFactor");
Context()->noreduce('(-x)-y','(-x)+y');
$var1 = RandomVariableName(type=>'variable');
do {$var2 = RandomVariableName(type=>'variable');} until ($var1 ne $var2);
Context()->variables->are($var1=>'Real',$var2=>'Real');

$a = list_random(2,3,5);
do {$b=random(2,10,1);} until (gcd($a,$b)==1);
$c = non_zero_random(-6,6,1);

$num1 = Formula("$a**2*$var1**2-$b**2")->reduce;
$den1 = Formula("$a*$var1**2+($a*$c+$b)*$var1+$b*$c")->reduce;
$num1Factored = Formula("($a*$var1+$b)*($a*$var1-$b)")->reduce;
$den1Factored = Formula("($a*$var1+$b)*($var1+$c)")->reduce;

$num2 = Formula("$b-$a*$var1")->reduce;
$den2 = Formula("1")->reduce;
$num2Factored = Formula("-($a*$var1-$b)")->reduce;
$den2Factored = Formula("1")->reduce;

$frac1= Formula("$num1/$den1");
$frac2= Formula("$num2/$den2")->reduce;

$frac3= Formula("$den2/$num2");
$num3Factored = $den2Factored;
$den3Factored = $num2Factored;

$ans = Formula("-1/($var1+$c)")->reduce;

Context("LimitedFraction")->flags->set(
            reduceFractions => 0,
            showMixedNumbers=>0,
            showExtraParens=>0);
Context()->operators->redefine(',',using=>',',from=>'Numeric');
Context()->operators->redefine('or',using=>',',from=>'Numeric');
Context()->operators->set(
  ','=>{string=>' and ',TeX=>'\hbox{ and }'},
  'or'=>{string=>' and ',TeX=>'\hbox{ and }'}
);
Context()->lists->set(List => {separator => "\; and \;"});
$lostFrac1 = Fraction($b,$a)->reduce;
$lostFrac2 = -Fraction($b,$a)->reduce;
$lostDomain = Compute("$lostFrac1,$lostFrac2");

##############################################
TEXT(beginproblem());

BEGIN_PGML

Simplify this expression. Don't forget to write the domain conditions lost during the simplification, if any. Please write your answer in *factored* form if possible.

    [``[$frac1] \div ([$frac2]) =``] [___________________]{$ans}

The domain condition(s) lost during the simplification is (are) [`[$var] \ne`][_____]. If you need to enter more than one number, use comma to separate them.

END_PGML

###################################################

$showPartialCorrectAnswers = 1;
    ANS($lostDomain->cmp(
      entry_type => "a fraction",
      checker => sub {
        my ($correct, $student, $ans, $nth, $value) = @_;
        $correct->context->setError("Your $nth $value is not reduced","",undef,undef,$Value::CMP_WARNING)
          if $student->{notReduced};
        return 0 unless $correct == $student;
        return 1 unless Value::classMatch($student,'Fraction') && !$student->isReduced;
        $student->{notReduced} = 1 unless $ans->{isPreview};
        return 0;
      }
    ));

###################################################

BEGIN_PGML_SOLUTION

    [``
\begin{aligned}
   [$frac1] \div [$frac2] 
      &= [$frac1] \div \frac{[$num2]}{1} \\
      &= [$frac1] \cdot [$frac3] \\
      &=  \frac{[$num1Factored]}{[$den1Factored]} \cdot \frac{[$num3Factored]}{[$den3Factored]} \\
      &= [$ans]
\end{aligned}
    ``]

Since the terms [`[$a][$var1]+[$b]`] and [`[$a][$var1]-[$b]`] were crossed out, the domain condition lost during the simplification are [`[$var1] \ne [$lostDomain]`].

END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
