## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Singular value decomposition)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');

$choice = random(0,2,1);

if ($choice == 0) { 
        $s = random(-1,1,2);
	$a = non_zero_random(-2,2,1);
	$b = 2/$a * random(-1,1,2);
	$d = $a * $s;
	$bc = (5 - $a*$a)*$s; 
	$c = $bc/$b;
} 

if ($choice == 1) {
        $s = random(-1,1,2);  
        $a = random(-3,3,2);
        $b = 3/$a * random(-1,1,2);
        $d = $a * $s;
        $bc = (10 - $a*$a)*$s; 
        $c = $bc/$b;
} 

if ($choice == 2) { 
	$s = random(-1,1,2);
	$a = random(-3,3,2);
	$d = 3/$a * $s;
	$b = random(-3,3,2);
	$c = - 3/$b * $s;
}

## A^T * A is  
##  a^2+c^2  ab+cd
##  ab+cd    b^2+d^2

$tr = $a*$a + $b*$b + $c*$c + $d*$d;
$det = ($a*$d - $b*$c)**2;
$discr = $tr*$tr - 4*$det;

## eigenvalues are (tr +- sqrt(discr))/2 
## singular values are sqrt of these

$eig1 = ($tr + sqrt($discr))/2;
$eig2 = ($tr - sqrt($discr))/2;

$sing1 = Compute("sqrt($eig1)");
$sing2 = Compute("sqrt($eig2)");

# eigenvectors v and u 

$v1 = $a*$b + $c*$d;
$v2 = $eig1 - $a*$a - $c*$c;

$u1 = $a*$b + $c*$d;
$u2 = $eig2 - $a*$a - $c*$c;

# make sure the eigenvectors are not 0

if (($v1 == 0) and ($v2 == 0)) { 
	$v1 = $eig1 - $b*$b - $d*$d; 
	$v2 = $a*$b + $c*$d;
} 

if (($u1 == 0) and ($u2 == 0)) { 
        $u1 = $eig2 - $b*$b - $d*$d;  
        $u2 = $a*$b + $c*$d; 
}

$nor1 = $v1*$v1+$v2*$v2;
$nor1=sqrt($nor1);
$nor2 = $u1*$u1+$u2*$u2;
$nor2=sqrt($nor2);

$basis1 = Matrix([$v1, $v2])->transpose/$nor1;
$basis2 = Matrix([$u1, $u2])->transpose/$nor2;

$A = Matrix([[$a, $b], [$c, $d]]);

Context()->texStrings;
BEGIN_TEXT
Find the singular values \( \sigma_1 \ge \sigma_2 \) of 
\[ A = $A.\]
$BR
\( \sigma_1 = \) \{ans_rule(10)\},
$BR
\( \sigma_2 = \) \{ans_rule(10)\}.
$BR
$BR
Find unit vectors \(\vec{v}_1\) and \(\vec{v}_2\) such that \( \| A\vec{v}_1 \| = \sigma_1 \) and \( \| A\vec{v}_2 \| = \sigma_2 \).
$BR
$BR
\(\vec{v}_1=\) \{ $basis1->ans_array(15) \}
$BR
$BR
\(\vec{v}_2=\) \{ $basis2->ans_array(15) \}
END_TEXT
Context()->normalStrings;

ANS($sing1->cmp);

ANS($sing2->cmp);

ANS($basis1->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return (($student == $correct) or ($student == -1*$correct));
}));

ANS($basis2->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return (($student == $correct) or ($student == -1*$correct));
}));

ENDDOCUMENT();       # This should be the last executable line in the problem.

