## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Integrating factor)
## Institution(The College of Idaho)
## Author(LA Danielson)
## Level(3)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Rogawski')
## EditionText1('2')
## Section1('9.5')
## Problem1('11')
## KEYWORDS('calculus', 'differential equations', 'linear equation', 'first order')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->set(x => { limits => [ .1, 2 ] });

##xy' = $a y - $b x

$a = Real(random(2,9,1));
$b = Real(random(1,9,1));

$ap1 = $a+1;
$am1 = $a-1;

$y1 = Real(random(-9,9,1));

$alpha = Formula("x^(-$a)")->cmp(
    checker => sub {
        my ($correct, $student, $ans) = @_;
        return 0 if $student == Formula(0); 
        if ($self->{_filter_name} ne 'produce_equivalence_message') {
            my $context = Context()->copy;
            $context->flags->set(no_parameters => 0);
            $context->variables->add(C0 => 'Parameter');
            $student = Formula($context, $student);
            $correct = Formula($context, "C0 * x**(-$a)");
        }
        return $correct == $student;
    }
);

Context("Fraction");
Context()->variables->add(C=>'Real');

$bdivam1 = Fraction("$b/$am1");

$ans1 = Formula("$bdivam1*x + C*x^$a")->reduce;

$C = Fraction("$y1-$b/$am1");

$ans2 = Formula("$bdivam1*x + $C*x^$a")->reduce;

BEGIN_PGML
In this problem you will solve the initial value problem    
    [` x y' = [$a] y - [$b] x\ `] with [`\ y(1) = [$y1] `].

    (a) Identify an integrating factor, [`\alpha(x)`].    
        [`\alpha(x)`] = [_]{$alpha} 

    (b) Find a general solution of the differential equation [` x y' = [$a] y - [$b] x`].    
    Your "general solution" should be valid at every point in an open rectangle that includes the initial condition [`(x,y)=(1,[$y1])`].     
    Use [`C`] for the arbitrary constant.      
        [`y(x)`] = [_]{$ans1}{15} 

    (c) Solve the initial value problem.    
        [`y(x)`] = [_]{$ans2}{15}
END_PGML

BEGIN_PGML_SOLUTION
(a) Divide by [`x`] to rewrite the differential equation in the standard form for linear ODEs    
    [`` y' - \frac{[$a]}{x} y = -[$b] \ .``]    
An integrating factor is    
    [`` \begin{aligned} \alpha(x) &= \exp\left(\int -\frac{[$a]}{x} \;dx\right) \\
       & = \exp(-[$a]\ln |x|+K) = e^K |x|^{-[$a]} \end{aligned}``]    
where [`K`] is some constant.  Any choice of [`K`] will work; [`K=0`] gives the 
simplest formula, so a reasonable candidate for [`\alpha(x)`] is  
    [``|x|^{-[$a]} = \begin{cases} x^{-[$a]} \text{ if } x>0 \\ (-x)^{-[$a]} \text{ if } x<0 \ . \end{cases} ``]    
But that is not continuous where [`x=0`], and [``\alpha(x)``] needs to be continuous so we can integrate it.  Also it needs to be defined near [`x=1`] because of the initial conditions, so it's clear that we should restrict its domain to [`0<x<\infty`], where     
    [``\alpha(x)=x^{-[$a]}``].    
(Any positive multiple [``e^K x^{-[$a]}``] of this would also work).
     
(b) When multiplied by the integrating factor [`\alpha(x)=x^{-[$a]}`], the standard form of the linear ODE becomes:
    [`` \begin{gathered} x^{-[$a]} y' - [$a] x^{-[$ap1]} y = -[$b] x^{-[$a]} \\    
      \Big( x^{-[$a]} y\Big)' = -[$b] x^{-[$a]} \end{gathered} ``]    
That was the whole point of the integrating factor -- multiplying the standard form of the ODE by the integrating factor turns the left-hand side into a derivative.  Now one can integrate both sides 
    [``x^{-[$a]} y = \int -[$b] x^{-[$a]} \; dx = [$bdivam1] x^{-[$am1]} + C``]    
then multiply by [`x^{[$a]}`] to solve for [`y`] and obtain the general solution 
    [``y = [$ans1]``]    
where [`C`] is an arbitrary constant.

(c) The initial condition says [`y=[$y1]`] when [`x=1`] so, plugging these values into the general solution, one obtains    
    [``\begin{gathered} [$y1] = [$bdivam1] + C \\
     [$y1]-[$bdivam1] = [$C] = C \end{gathered}``]    
Thus the solution to the initial value problem is    
    [`` y = [$ans2] ``]
END_PGML_SOLUTION

ENDDOCUMENT();
