## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Rank)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('4.3')
## Problem1('')
## KEYWORDS('linear algebra','matrix','rank')


######################################
#  Initialization

DOCUMENT() ;        
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserVectorUtils.pl",
"PGcourse.pl"
) ;
TEXT(beginproblem());

#######################################
#  Setup

Context('Matrix')->variables->are(k=>'Real');

foreach $i (0..2) {
	foreach $j (0..1) {
		$a[$i][$j] = non_zero_random(-9,9,1);
	}
}

while($a[0][0]*$a[1][1]==$a[0][1]*$a[1][0]) {
        $a[1][1]=non_zero_random(-9,9,1);
}

$l = non_zero_random(-2,2,1);
$m = non_zero_random(-2,2,1);

foreach $i (0..2) {
	$a[$i][2] = $l * $a[$i][0] + $m * $a[$i][1];
}



$a = non_zero_random(-9,9,1);
$b = non_zero_random(-9,9,1);

$x = Compute("1 - $b");
$y = Compute("$b * (1 - $b)/$a");

$A = Matrix([
[$a[0][0], $a[0][1], $a[0][2]], 
[$a[1][0], $a[1][1], $a[1][2]],
[$a[2][0], $a[2][1], 'k']
]);


$answer = $A;
$vecans=ColumnVector($l,$m,-1);

#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
a) Find the value of \(k\) for which the matrix 
\[ A = $A \]
has rank \(2\). 
$BR
\(k=\) \{ans_rule(10)\}
$BR
b) For this value of \(k\), find a basis of \(\ker(A)\).
$BR
Basis = \{$vecans->ans_array()\}.
END_TEXT
Context()->normalStrings;

#######################################
#  Answer evaluation

$showPartialCorrectAnswers = 1 ;

ANS( Compute("$a[2][2]")->cmp() );
ANS( $vecans->cmp(
         checker=>sub {
	my ( $cor, $stu, $ans ) = @_;
        if ($stu->isParallel($vecans)) {
        return 1;
        }
        return 0;
        }
) );
;

ENDDOCUMENT();
