## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Systems of linear equations)
## DBsection(Testing linear combinations)
## Date(August 2015)
## Institution(University of Michigan)
## Author(David Sher)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Otto Bretscher')
## EditionText1('5')
## Section1('1.3')

## New problem for section 1.3, number 2

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Vector");

$a=non_zero_random(-5,5,1);
$b=non_zero_random(-5,5,1);
$c=non_zero_random(-3,3,1);
$d=non_zero_random(-3,3,1);
$e=non_zero_random(-3,3,1);
$f=non_zero_random(-3,3,1);
$g=non_zero_random(2,5,1);
$h=non_zero_random(2,5,1);

$v1=ColumnVector($a,$b,$c);
$v2=ColumnVector($d,$e,$f);
$vcomb=$g*$v1+$h*$v2;

$c1=$a*$g+$d*$h;
$c2=$b*$g+$e*$h;

if ( $e*$a - $d*$b == 0 ) {
    $answ = Compute( "0" );
    $anyAns = 1;
} else{
    $answ=Compute($c*$g+$f*$h);
    $anyAns = 0;
}

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find a value of \(k\) for which \(\vec w=\begin{bmatrix}$c1\\ $c2\\ k\end{bmatrix}\) is a linear combination of \(\vec v_1=$v1\) and \(\vec v_2=$v2\).
$BR
\(k=\) \{ $answ->ans_rule()\}

END_TEXT
Context()->normalStrings;

if ( $anyAns ) {
    ANS( $answ->cmp( checker=>sub { return 1; } ) );
} else {
    ANS( $answ->cmp() );
}
# ANS( $answ->cmp() );

ENDDOCUMENT();
