# ************************************
#   Comments + Keyword Clarification
#
#    Anton, Calculus, 0.3.5
# ************************************
# DESCRIPTION
#   Find an equation for the family of lines tangent to the circle
#   with center at origin and radius = R
#
# ENDDESCRIPTION
## DBsubject('Calculus')
## DBchapter('Before Calculus')
## DBsection('Families of Functions')
## KEYWORDS('Tangent', 'Circle', 'Line', 'Function')
## TitleText1('Calculus')
## EditionText1('9')
## AuthorText1('Anton')
## Section1('0.3')
## Problem1('5')
## Author('')
## Institution('John Wiley & Sons')
#***************************************************************
#   Initialization - Load standard PG and MathObjects macros   
#***************************************************************
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "littleneck.pl",
  "PGcourse.pl"
);
#
rand_button();
#
# Set Up - Set up variable context and formulation of problem
#
Context("Numeric");
#************************
#   Circle of Radius R
#************************
$R = random(2,10,1);
$Rsqr = Compute("$R^2");
#**************************************************
#   Answer, separated into slope and y intercept
#**************************************************
Context()->variables->add(a=>'Real',b=>'Real');
$Slope = Formula("a/sqrt($Rsqr - a^2)");
$Yint = Formula("$Rsqr/sqrt($Rsqr - a^2)");
$Eqn1 = Formula("$Slope*x + $Yint")->reduce;
$Eqn2 = Formula("$Slope*x - $Yint")->reduce;
$Eqn3 = Formula("-$Slope*x + $Yint")->reduce;
$Eqn4 = Formula("-$Slope*x - $Yint")->reduce;
# *********************
#   Problem Statement
# *********************
TEXT(beginproblem());
#
# Strings in Tex form (formula)
#
Context()->texStrings;
BEGIN_TEXT
Below is the graph of a circle with center at the origin and a radius of $R.
Line \( \small{l_2} \) is tangent to the circle at \( \small{(a,b)} \), and
line \( \small{l_1} \) passes through the origin and intersects \( \small{l_2} \)
at \( \small{(a,b)} \).
$PAR
$BCENTER
\{ image("Q5-Circle.png", height=>225, width=>225) \}
$ECENTER
$PAR
Using the point \( \small{(a,b)} \) and the relationship between \( \small{l_1} \)
and \( \small{l_2} \), find an equation for the family of lines tangent to this circle.
Express your final answer in terms of \( \small{x \;\mbox{and} \; a} \).
$PAR
$BCENTER
\( \small{y = } \) \{ ans_rule(30) \}
$ECENTER
END_TEXT
#
# Expect strings in standard form
#
Context()->normalStrings;
# *******************************
#   Answer and Solution section
# *******************************
#
# Check student solution
#
#ANS($Slope->cmp);
#ANS($Yint->cmp);
#***********************************
#   Accept any of the 4 equations
#***********************************
ANS( $Eqn1->cmp(checker=>sub { my ($correct,$student) = @_;
                               return ( ($Eqn1 == $student) || ($Eqn2 == $student) ||
                                        ($Eqn3 == $student) || ($Eqn4 == $student) )
                             } ) );
#***************************
#  Explanation of solution  
#***************************
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
Since line \( \small{l_1} \) passes through the center of the circle and \( \small{l_2} \) is
tangent to the circle, these lines are perpendicular. Since line \( \small{l_1} \) passes through
the origin, its slope = \( \large{\frac{b}{a}} \). Therefore, the slope of \( \small{l_2} \) =
\( \large{-\frac{a}{b}} \).
To find the equation of line \( \small{l_2} \), use the point-slope formula:
$BCENTER
\( \small{y - b = } \large{-\frac{a}{b}} \small{(x - a)} \Rightarrow
\small{y = } \large{-\frac{a}{b}} \small{x \;+\;} \large{\frac{a^2}{b}} \small{+\; b} \Rightarrow
\small{y = } \large{-\frac{a}{b}} \small{x \;+\; } \large{\frac{a^2 + b^2}{b}} \).
$ECENTER
$BR
$BR
To express this solely in \( \small{a} \) note that since \( \small{(a,b)} \) lies on the
circle, \( \small{a^2 + b^2 = $Rsqr} \) and so \( \small{b = \pm\sqrt{$Rsqr - a^2}} \).
Replacing these into the above equation yields,
$BR
$BCENTER
\( \small{y = \pm} \large{\frac{a}{\sqrt{$Rsqr - a^2}}} \small{x \;\pm\; } \large{\frac{$Rsqr}{\sqrt{$Rsqr - a^2}}} \).
$PAR 
END_SOLUTION
Context()->normalStrings;
ENDDOCUMENT();
