## DESCRIPTION
## Linear transformations
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Properties)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(4)
## TitleText1('Multivariable Mathematics')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('linear transformations')

DOCUMENT(); 
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "unionLists.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;

Context('Numeric')->variables->are('x'=>'Real');
Context()->noreduce('(-x)-y');
Context()->constants->add(
'a2'=>7.281,
'a1'=>2.272,
'a0'=>3.273,
'b2'=>8.261,
'b1'=>4.299,
'b0'=>9.826,
'c'=>1.2379,
);

# this means that student answers are not reduced to the numerical value specified in the Context
Context()->flags->set(
    formatStudentAnswer=>'parsed'
);

$p = Formula("a2 x^2 + a1 x + a0")->reduce;
$q = Formula("b2 x^2 + b1 x + b0")->reduce;

$popup = PopUp(['choose','f is a linear transformation','f is not a linear transformation'],'f is a linear transformation');
$popup_add = PopUp(['choose','Yes, they are equal','No, they are not equal'],'Yes, they are equal');
$popup_mult = PopUp(['choose','Yes, they are equal','No, they are not equal'],'Yes, they are equal');

$ans[1] = Formula("(a2+b2) x^3 + (a1+b1) x^2 + (a0+b0) x")->reduce;
$ans[2] = Formula("a2 x^3 + a1 x^2 + a0 x")->reduce;
$ans[3] = Formula("b2 x^3 + b1 x^2 + b0 x")->reduce;

$ans[4] = Formula("a2 c x^3 + a1 c x^2 + a0 c x")->reduce;
$ans[5] = Formula("c")->reduce;
$ans[6] = Formula("a2 x^3 + a1 x^2 + a0 x")->reduce;

$multians[1] = MultiAnswer($ans[2], $ans[3])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s + %s",
  tex_format => "%s + %s",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);

$multians[2] = MultiAnswer($ans[5], $ans[6])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s ( %s )",
  tex_format => "%s \left( %s \right)",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);



Context()->texStrings;
BEGIN_TEXT
Let \( \mathcal{P}_n \) denote the vector space of polynomials in the variable \( x \) of degree \( n \) or less with real coefficients.
Let \( f : \mathcal{P}_2 \to \mathcal{P}_3 \) be defined by \( f(p(x)) = x p(x) \) for all polynomials \( p(x) \in \mathcal{P}_2 \).  
Is \( f \) a linear transformation?
$BR
$BR
Let \( p(x) = a_2 x^2 + a_1 x + a_0 \) and \( q(x) = b_2 x^2 + b_1 x + b_0 \) be any two polynomials in \( \mathcal{P}_2 \) and \( c \in \mathbb{R} \).
$BR
$BR
\{ BeginList('OL',type=>'a') \}
$ITEM
\( f(p(x)+q(x)) = \) \{ $ans[1]->ans_rule(50) \}.  (Enter \( a_2 \) as \(\verb+a2+\), etc.)
$BR
\( f(p(x)) + f(q(x)) = \) \{ $multians[1]->ans_rule(25) \} \( + \) \{ $multians[1]->ans_rule(25) \}.
$BR
Does \( f(p(x)+q(x)) = f(p(x)) + f(q(x)) \) for all \( p(x), q(x) \in \mathcal{P}_2 \)? \{ $popup_add->menu \}
$ITEMSEP
$ITEM
\( f(c p(x)) = \) \{ $ans[4]->ans_rule(40) \}.
$BR
\( c(f(p(x))) = \) \{ $multians[2]->ans_rule(5) \} \( \Big( \) \{ $multians[2]->ans_rule(20) \} \( \Big) \).
$BR
Does \( f(c p(x)) = c(f(p(x)))\) for all \( c \in \mathbb{R} \) and all \( p(x) \in \mathcal{P}_2 \)? \{ $popup_mult->menu \}
$ITEMSEP
$ITEM Is \( f \) a linear transformation? \{ $popup->menu \}
\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;

ANS($ans[1]->cmp);
ANS($multians[1]->cmp);
ANS($popup_add->cmp);

ANS($ans[4]->cmp);
ANS($multians[2]->cmp);
ANS($popup_mult->cmp);

ANS($popup->cmp);

COMMENT('Always a linear transformation.');
ENDDOCUMENT();
