## DESCRIPTION
## Linear transformations
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Properties)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(4)
## TitleText1('Multivariable Mathematics')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('linear transformations')

DOCUMENT(); 
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "unionLists.pl",
  "parserFunction.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;

Context('Numeric')->variables->are('x'=>'Real');
Context()->noreduce('(-x)-y');
Context()->constants->add(
'a3'=>3.729,
'a2'=>7.281,
'a1'=>2.272,
'a0'=>3.273,
'b3'=>2.371,
'b2'=>8.261,
'b1'=>4.299,
'b0'=>9.826,
'c'=>1.2379,
);

# this means that student answers are not reduced to the numerical value specified in the Context
Context()->flags->set(
    formatStudentAnswer=>'parsed'
);

$popup = PopUp(['choose','f is a linear transformation','f is not a linear transformation'],'f is a linear transformation');
$popup_add = PopUp(['choose','Yes, they are equal','No, they are not equal'],'Yes, they are equal');
$popup_mult = PopUp(['choose','Yes, they are equal','No, they are not equal'],'Yes, they are equal');

$ans[1] = Formula("6(a3+b3) x + 2(a2+b2)")->reduce;
$ans[2] = Formula("6 a3 x + 2 a2")->reduce;
$ans[3] = Formula("6 b3 x + 2 b2")->reduce;

$ans[4] = Formula("c(6 a3 x + 2 a2)")->reduce;
$ans[5] = Formula("c")->reduce;
$ans[6] = Formula("6 a3 x + 2 a2")->reduce;

$multians[1] = MultiAnswer($ans[2], $ans[3])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s + %s",
  tex_format => "%s + %s",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);

$multians[2] = MultiAnswer($ans[5], $ans[6])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s ( %s )",
  tex_format => "%s \left( %s \right)",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);



Context()->texStrings;
BEGIN_TEXT
Let \( \mathcal{P}_n \) denote the vector space of polynomials in the variable \( x \) of degree \( n \) or less with real coefficients.
Let \( D^2 : \mathcal{P}_3 \to \mathcal{P}_1 \) be the function that sends a polynomial to its second derivative.  That is, \( D^2(p(x)) = p''(x) \) for all polynomials \( p(x) \in \mathcal{P}_3 \).  
Is \( D^2 \) a linear transformation?
$BR
$BR
Let \( p(x) = a_3 x^3 + a_2 x^2 + a_1 x + a_0 \) and \( q(x) = b_3 x^3 + b_2 x^2 + b_1 x + b_0 \) be any two polynomials in \( \mathcal{P}_3 \) and \( c \in \mathbb{R} \).
$BR
$BR
\{ BeginList('OL',type=>'a') \}
$ITEM
\( D^2(p(x)+q(x)) = \) \{ $ans[1]->ans_rule(50) \}.  (Enter \( a_3 \) as \(\verb+a3+\), etc.)
$BR
\( D^2(p(x)) + D^2(q(x)) = \) \{ $multians[1]->ans_rule(25) \} \( + \) \{ $multians[1]->ans_rule(25) \}.
$BR
Does \( D^2(p(x)+q(x)) = D^2(p(x)) + D^2(q(x)) \) for all \( p(x), q(x) \in \mathcal{P}_3 \)? \{ $popup_add->menu \}
$ITEMSEP
$ITEM
\( D^2(c p(x)) = \) \{ $ans[4]->ans_rule(40) \}.
$BR
\( c(D^2(p(x))) = \) \{ $multians[2]->ans_rule(5) \} \( \Big( \) \{ $multians[2]->ans_rule(20) \} \( \Big) \).
$BR
Does \( D^2(c p(x)) = c(D^2(p(x)))\) for all \( c \in \mathbb{R} \) and all \( p(x) \in \mathcal{P}_3 \)? \{ $popup_mult->menu \}
$ITEMSEP
$ITEM Is \( D^2 \) a linear transformation? \{ $popup->menu \}
\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;

ANS($ans[1]->cmp);
ANS($multians[1]->cmp);
ANS($popup_add->cmp);

ANS($ans[4]->cmp);
ANS($multians[2]->cmp);
ANS($popup_mult->cmp);

ANS($popup->cmp);

COMMENT('Always a linear transformation.');
ENDDOCUMENT();
