## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('8.3')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');
Context()->flags->set(
  tolerance => 0.001,
  tolType => "absolute",
);

# orthogonal matrix: 

$a[1][1] = 1/sqrt(3);  $a[1][2] = -2/sqrt(6);  $a[1][3] = 0;
$a[2][1] = 1/sqrt(3);  $a[2][2] = 1/sqrt(6);   $a[2][3] = 1/sqrt(2);
$a[3][1] = 1/sqrt(3);  $a[3][2] = 1/sqrt(6);   $a[3][3] = -1/sqrt(2);

# slice columns and rows: 

@slice1 = NchooseK(3,3);
@slice2 = NchooseK(3,3);

foreach $i (1..3) { 
	foreach $j (1..3) { 
		$u[$i][$j] = $a[$slice1[$i-1]+1][$slice2[$j-1]+1];
#		$U -> assign($i,$j,$u[$i][$j]);
	}
}

$U = Matrix([
[ $u[1][1], $u[1][2], $u[1][3] ],
[ $u[2][1], $u[2][2], $u[2][3] ],
[ $u[3][1], $u[3][2], $u[3][3] ],
]);

if ($slice2[0]==0) { $eig1 = random(-9,-3,3); } 
if ($slice2[0]==1) { $eig1 = -6; }
if ($slice2[0]==2) { $eig1 = random(-8,-2,2); } 

$eig2 = 0;

if ($slice2[2]==0) { $eig3 = random(3,9,3); }
if ($slice2[2]==1) { $eig3 = 6; } 
if ($slice2[2]==2) { $eig3 = random(2,8,2); } 

$E = Matrix([
[$eig1,0,0],
[0,$eig2,0],
[0,0,$eig3],
]);

$I = $U->inverse;

$A = $U * $E * $I;


$lambda[1] = Compute("$eig1");
$lambda[2] = Compute("$eig2");
$lambda[3] = Compute("$eig3");

$basis[1] = Matrix([ $u[1][1], $u[2][1], $u[3][1] ])->transpose;
$basis[2] = Matrix([ $u[1][2], $u[2][2], $u[3][2] ])->transpose;
$basis[3] = Matrix([ $u[1][3], $u[2][3], $u[3][3] ])->transpose;

Context()->texStrings;
BEGIN_TEXT
Find the eigenvalues \( \lambda_1 < \lambda_2 < \lambda_3 \) and associated unit eigenvectors \( \vec{u}_1, \vec{u}_2, \vec{u}_3 \) of the symmetric matrix 
\[ A = $A. \]
$BR
The eigenvalue \( \lambda_1 = \) \{ $lambda[1]->ans_rule(10) \} 
has associated unit eigenvector \( \vec{u}_1 = \) \{ $basis[1]->ans_array(15) \}.
$BR
$BR
The eigenvalue \( \lambda_2 = \) \{ $lambda[2]->ans_rule(10) \} 
has associated unit eigenvector \( \vec{u}_2 = \) \{ $basis[2]->ans_array(15) \}.
$BR
$BR
The eigenvalue \( \lambda_3 = \) \{ $lambda[3]->ans_rule(10) \} 
has associated unit eigenvector \( \vec{u}_3 = \) \{ $basis[3]->ans_array(15) \}.
$BR
Note: The eigenvectors above form an orthonormal eigenbasis for \(A\). 
END_TEXT
Context()->normalStrings;

foreach my $i (1..3) {

ANS($lambda[$i]->cmp);
ANS($basis[$i]->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return (($student == $correct) or ($student == -1*$correct));
}));

}

COMMENT('One of the eigenvalues is zero and the others are nonzero.');
ENDDOCUMENT();
