## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(4)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('8.3')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');
Context()->flags->set(
  tolerance => 0.001,
  tolType => "absolute",
);

do {

$a = random(1,9,1);
$b = random(-9,-1,1);

} until (-1*$a < $b);

$A = Matrix([[0, 0, $a], [0, $b, 0], [$a, 0, 0]]);

$lambda[1] = Compute("-1*$a");
$lambda[2] = Compute("$b");
$lambda[3] = Compute("$a");

$basis[1] = Matrix(["-1/sqrt(2)", 0, "1/sqrt(2)"])->transpose;
$basis[2] = Matrix([0, 1, 0])->transpose;
$basis[3] = Matrix(["1/sqrt(2)",0, "1/sqrt(2)"])->transpose;

Context()->texStrings;
BEGIN_TEXT
Find the eigenvalues \( \lambda_1 < \lambda_2 < \lambda_3 \) and associated unit eigenvectors \( \vec{u}_1, \vec{u}_2, \vec{u}_3 \) of the symmetric matrix 
\[ A = $A. \]
$BR
The eigenvalue \( \lambda_1 = \) \{ $lambda[1]->ans_rule(10) \} 
has associated unit eigenvector \( \vec{u}_1 = \) \{ $basis[1]->ans_array(15) \}.
$BR
$BR
The eigenvalue \( \lambda_2 = \) \{ $lambda[2]->ans_rule(10) \} 
has associated unit eigenvector \( \vec{u}_2 = \) \{ $basis[2]->ans_array(15) \}.
$BR
$BR
The eigenvalue \( \lambda_3 = \) \{ $lambda[3]->ans_rule(10) \} 
has associated unit eigenvector \( \vec{u}_3 = \) \{ $basis[3]->ans_array(15) \}.
$BR
Note: The eigenvectors above form an orthonormal eigenbasis for \(A\). 
END_TEXT
Context()->normalStrings;

foreach my $i (1..3) {

ANS($lambda[$i]->cmp);
ANS($basis[$i]->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return (($student == $correct) or ($student == -1*$correct));
}));

}

COMMENT('One of the eigenvalues is zero and the others are nonzero.');
ENDDOCUMENT();
