# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/02/10


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(pc_lines_2d)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.1')
## Section2('17.1')
## Section3('17.1')
## Problem1('2')
## Problem2('5')
## Problem3('3')
## KEYWORDS('parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "parserParametricLine.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector2D");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"}
);
Context()->variables->are( t=>'Real' );
$showPartialCorrectAnswers = 1;

$x0 = list_random(1,2,3);
$y0 = list_random(1,2,3);
$dx = list_random(1,2,3);
$dy = list_random(1,2,3);
$x1 = $x0 + $dx;
$y1 = $y0 + $dy;

## then the graph is
$gr = init_graph( -1, -1, $x1+1, $y1+1, size=>[250,250], axes=>[0,0],
		  grid=>[$x1+2, $y1+2] );
$gr->stamps( closed_circle( $x0, $y0, 'black' ) );
$gr->stamps( closed_circle( $x1, $y1, 'black' ) );
$gr->moveTo( $x0,$y0 );
$gr->lineTo( $x1, $y1, 'blue', 2 );
$gr->moveTo( $x0,$y0 );
$gr->arrowTo( $x0+$dx/2, $y0+$dy/2, 'blue', 2 );

$xt = Compute( "$x0 + $dx*t" );
$yt = Compute( "$y0 + $dy*t" );
$r = Vector( "<$xt,$yt>" );
$t0 = Compute( "0" );
$t1 = Compute( "1" );

$pchk = MultiAnswer( $r, $t0, $t1 )->with(
    singleAnswer => 0,
    checker => sub {
	my ( $c, $s, $ans ) = @_;
	my $cr = $c->[0];
	my $cline = ParametricLine( $cr );
	my $startPt = $cline->eval( t=>$t0 );
	my $stopPt  = $cline->eval( t=>$t1 );
	my $sr = $s->[0];
	my $sline = ParametricLine( $sr );
	my $stuStart = $sline->eval( t=>$s->[1] );
	my $stuStop  = $sline->eval( t=>$s->[2] );

	if ( $cline == $sline && $startPt == $stuStart &&
	     $stopPt == $stuStop ) {
	    return [ 1, 1, 1 ];
	} elsif ( $cline == $sline && $startPt == $stuStart ) {
	    return [ 1, 1, 0 ];
	} elsif ( $cline == $sline && $stopPt == $stuStop ) {
	    return [ 1, 0, 1 ];
	} elsif ( $cline == $sline ) {
	    return [ 1, 0, 0 ];
	} else {
	    return [ 0, 0, 0 ];
	}
} );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a parameterization for the curve shown in the figure below.
$BR
$BCENTER
\{ image( insertGraph($gr), tex_size=>500, height=>250, width=>250,
	  extra_html_tags=>'alt="graph of the line segment from (' .
	  $x0 . ', ' . $y0 . ') to (' . $x1 . ', ' . $y1 . ')."' ) \}
$ECENTER
$BR
\( \vec r(t) = \) \{ $pchk->ans_rule(35) \},
$BR
where
\{ $pchk->ans_rule(15) \} \( \le t \le \) \{ $pchk->ans_rule(15) \}
$BR
(Note: you need to enter answers in all three answer blanks to receive feedback.)
END_TEXT
Context()->normalStrings;

ANS( $pchk->cmp() );

$x0d = ( $x0 == 1 ) ? '' : $x0;
$y0d = ( $y0 == 1 ) ? '' : $y0;
$dxd = ( $dx == 1 ) ? '' : $dx;
$dyd = ( $dy == 1 ) ? '' : $dy;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

We want the straight line segment from \(($x0,$y0)\) to \(($x1,$y1)\).  The
position vector of \((1,1)\) is \($x0d\vec i + $y0d\vec j\) and the
displacement vector from \(($x0,$y0)\) to \(($x1,$y1)\) is
\($dxd\vec i + $dyd\vec j\), so the line has
equation
\[
\vec r = $x0d \vec i + $y0d\vec j + t($dxd\vec i + $dyd\vec j),
\]
or
\[
x = $x0 + $dxd t, \quad y = $y0 + $dyd t.
\]
This passes \(($x0,$y0)\) when \(t=0\) and \(($x1,$y1)\) when
\(t=1\), so a possible
parameterization is
\[
\vec r = x(t)\vec i + y(t)\vec j =
($x0 + $dxd t)\vec i + ($y0 + $dyd t)\vec j,
\quad 0 \le t \le 1.
\]

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
