# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/02/10


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(pc_circles_center_radius)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.1')
## Section2('17.1')
## Section3('17.1')
## Problem1('19')
## Problem2('21')
## Problem3('21')
## KEYWORDS('parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are( t=>'Real' );
$showPartialCorrectAnswers = 1;

$r0 = random(2,8,1);
$pchk = MultiAnswer( Compute( "$r0*sin(t)" ), Compute( "$r0*cos(t)" ),
		     Compute( "0" ) )->with(
    singleAnswer => 0,
    checker => sub {
	my ( $c, $s, $ans ) = @_;
	my ($xs, $ys, $zs) = @{$s};

	## get some formulas to be sure that we can evaluate and
	##    differentiate them
	my ( $xsf, $ysf ) = ( Formula( $xs ), Formula( $ys ) );
	## plug in a value to be able to test the derivatives
	my ( $x0, $y0 ) = ( $xsf->eval(t=>0), $ysf->eval(t=>0) );
	my ( $xd0, $yd0 ) = ( $xsf->D->eval(t=>0), $ysf->D->eval(t=>0) );

	my @ret = (0,0,0);
	$ret[2] = ( $c->[2] == $zs ) ? 1 : 0;

	if ( $xs*$xs + $ys*$ys == $r0*$r0 &&
	     ( ( $x0 < 0 && $yd0 > 0 ) || ( $x0 == 0 && $yd0 == 0 ) ||
	       ( $x0 > 0 && $yd0 < 0 ) ) &&
	     ( ( $y0 < 0 && $xd0 < 0 ) || ( $y0 == 0 && $xd0 == 0 ) ||
	       ( $y0 > 0 && $xd0 > 0 ) ) ) {
	    $ret[0] = 1;
	    $ret[1] = 1;
	} elsif ( $xs*$xs + $ys*$ys == $r0*$r0 && $xsf->D != 0 ) {
	    $ans->setMessage( 1, "Check that your expressions parameterize " .
			      "the circle in a clockwise direction." );
	}
	return [ @ret ];
} );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a parameterization for the
circle of radius $r0
in the \( xy \)-plane, centered at the origin,
clockwise.

$PAR
\( x(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( y(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( z(t) = \) \{ $pchk->ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($pchk->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The \(xy\)-plane is where \(z=0\), and to make the particle go in
the clockwise direction we start at \(($r0,0,0)\) and head in the
negative \(y\)-direction. Thus one possible answer is
\[ x = $r0 \cos t, \quad y = -$r0 \sin t, \quad z = 0. \]



END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
