# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/02/10


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(pc_cartesian_to_parametric)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.1')
## Section2('17.1')
## Section3('17.1')
## Problem1('27')
## Problem2('31')
## Problem3('31')
## KEYWORDS('parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are( t=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(1,9,1);
$b = non_zero_random(-5,5,1);
$r = list_random(3,4);
$z0 = non_zero_random(-5,5,1);
$t0 = random(-3,3,1);
$ycurve = ( $b < 0 ) ? "$a $b x^{$r}" : "$a + $b x^{$r}";

$pchk = MultiAnswer( Compute( "t-$t0" ), Compute( "$a + $b*(t-($t0))^{$r}" ),
		     Compute( "$z0" ) )->with(
    singleAnswer => 0,
    checker => sub {
	my ( $c, $s, $ans ) = @_;
	my ($xs, $ys, $zs) = @{$s};

	my @ret = (0,0,0);
	$ret[2] = ( $c->[2] == $zs ) ? 1 : 0;

	my $ycheck = Compute( "$a + $b*$xs^{$r}" );

	if ( $ycheck == $ys && $xs->eval(t=>$t0) == 0 &&
	     $ys->eval(t=>$t0) == $a ) {
	    $ret[0] = 1;
	    $ret[1] = 1;
	} elsif ( $ycheck == $ys ) {
	    $ans->setMessage(1, 'Check that your parameterization ' .
			     'passes through the indicated point.');
	}
	return [ @ret ];
} );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a parameterization for the
curve
\( y = $ycurve \) that passes through the point
\( (0, $a , $z0) \) when \(t = $t0\) and is parallel to
the \( xy \)-plane.

$PAR
\( x(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( y(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( z(t) = \) \{ $pchk->ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($pchk->cmp() );

$targ = ( $t0 == 0 ) ? "t" : Compute("t - $t0")->reduce();
$s = ( $b < 0 ) ? '-' : '+';
$bd = abs($b);

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Since the curve is parallel to the \(xy\)-plane, \(z\) is constant, and
since it passes through \((0,$a,$z0)\), we have \(z=$z0\). One possible answer
is
\[
x = $targ, \quad y = $a $s $bd ($targ)^{$r}, \quad z=$z0,
\]
which passes through the indicated point when \(t = $t0\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
