# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/02/10


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(pc_cartesian_to_parametric)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.1')
## Section2('17.1')
## Section3('17.1')
## Problem1('27')
## Problem2('31')
## Problem3('31')
## KEYWORDS('parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(1,9,1);
$b = non_zero_random(-5,5,1);
$r = list_random(3,4);
$z0 = non_zero_random(-5,5,1);
$t0 = random(-3,3,1);
#$ycurve = ( $b < 0 ) ? "$a $b x^{$r}" : "$a + $b x^{$r}";
$ycurve = Formula("$a + $b x^{$r}")->reduce;

$pchk = MultiAnswer( Compute( "t-$t0" ), Compute( "$a + $b*(t-($t0))^{$r}" ),
		     Compute( "$z0" ) )->with(
    singleAnswer => 0,
    checker => sub {
	my ( $c, $s, $ans ) = @_;
	my ($xs, $ys, $zs) = @{$s};

        if ( Formula($xs)->eval(t=>$t0) != Formula(0) or 
	     Formula($ys)->eval(t=>$t0) != Formula($a) or
             Formula($zs)->eval(t=>$t0) != Formula($zs) )
        { $ans->setMessage(1,"Check that your parametrization passes ".
             "through \( (0, $a , $z0) \) when \(t = $t0\).");
          return (0,0,0);
        }
        if ( $xs->D(t)->eval(t=>$t0) == Formula(0) ) 
        {
           $ans->setMessage(1,"Check that \(x'($t0)\neq 0\)");
           return (0,0,0);
        }
        if ( $c->[2] != $zs )
        {
           $ans->setMessage(3,"Is your curve parallel to the \(x,y\)-plane?");
           return (0,0,0);  
        }
        if (Compute( "$a + $b*$xs^{$r}" )!=$ys) {
           $ans->setMessage(2,"Your curve doesn't follow the right path.");
           return (1,0,1);
        }
        else { return (1,1,1);}
} );



Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a parameterization \( (x(t),y(t),z(t)) \) for the
curve
\( y = $ycurve \) that is parallel to
the \( xy \)-plane.  Your parametrized curve should pass through the point
\( (0, $a , $z0) \) when \(t = $t0\), and have \(x'($t0)\neq 0\).  

$PAR
\( x(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( y(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( z(t) = \) \{ $pchk->ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($pchk->cmp() );

$xtarg = Compute("t-$t0")->reduce;
$ytarg = Compute("$a+$b($xtarg)^{$r}")->reduce;
#$s = ( $b < 0 ) ? '-' : '+';
#$bd = abs($b);

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
\(z\) is constant since the curve is parallel to the \(xy\)-plane, so \(z=$z0\) 
since it passes through \((0,$a,$z0)\). One possible answer
is
\[
x = $xtarg, \quad y = $ytarg, \quad z=$z0,
\]
which passes through the indicated point, with \(x'($t0)=1\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
