# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/02/10


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.1')
## Section2('17.1')
## Section3('17.1')
## Problem1('37')
## Problem2('41')
## Problem3('41')
## KEYWORDS('parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are( t=>'Real' );
$showPartialCorrectAnswers = 1;

$r0 = random(2,9,1);
$twor0 = 2*$r0;

$pchk = MultiAnswer( Compute( "$r0 - $r0*cos(t)" ),
		     Compute( "$r0*sin(t)" ), Compute("0"),
		     Compute("pi") )->with(
    singleAnswer => 0,
    checker => sub {
	my ( $c, $s, $ans ) = @_;

	## make sure that the student answers are formulas
	my ( $xsf, $ysf ) = ( Formula( $s->[0] ), Formula( $s->[1] ) );

	my @ret = ( 0, 0, 0, 0 );
	if ( ($xsf - $r0)*($xsf - $r0) + $ysf*$ysf == $r0*$r0 ) {
	    $ret[0] = 1;
	    $ret[1] = 1;
	}
	if ( Compute($s->[2]) < Compute($s->[3]) ) {
	    ## check that the endpoints evaluate correctly
	    if ( $xsf->eval(t=>$s->[2]) == 0 &&
		 $ysf->eval(t=>$s->[2]) == 0 ) {
		$ret[2] = 1;
	    }
	    if ( $xsf->eval(t=>$s->[3]) == $twor0 &&
		 $ysf->eval(t=>$s->[3]) == 0 ) {
		$ret[3] = 1;
	    }
	    if ( $ret[2] && $ret[3] ) {
  	        ## and that we don't get any points that have x and y
	        ##    coordinates that they shouldn't
	        ## check 10 points
		$delt = ($s->[3] - $s->[2])/10;
	        ## are we going up or down?
		$which = ( $ysf->eval(t=>($s->[2] + $delt)) > 0 ) ? 1 : 0;
		for ( my $i=1; $i<10; $i++ ) {
		    $tcheck = $s->[2] + $i*$delt;
		    $stux = $xsf->eval(t=>$tcheck);
		    $stuy = $ysf->eval(t=>$tcheck);
		    if ( $which ) {
			if ( $stux < 0 || $stux > $twor0 || $stuy < 0 ) {
			    $ret[3] = 0;  # we know the endpoints evaluated
			    break;        #    to the correct values, so give
			}                 #    credit for the left one
		    } else {
			if ( $stux < 0 || $stux > $twor0 || $stuy > 0 ) {
			    $ret[3] = 0;  # we know the endpoints evaluated
			    break;        #    to the correct values, so give
			}                 #    credit for the left one
		    }
		}

	    }
	} else {
	    $ret[2] = 0;
	    $ret[3] = 0;
	}
	return [@ret];
    }
);


Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find parametric equations for the
arc of a circle of radius $r0 from \( P=(0,0) \) to \( Q=($twor0,0) \).

$PAR
\( x(t) = \) \{ $pchk->ans_rule(35) \}, $BR
\( y(t) = \) \{ $pchk->ans_rule(35) \}, $BR
with
\{ $pchk->ans_rule(15) \} \( \le t\le \) \{ $pchk->ans_rule(15) \}

END_TEXT
Context()->normalStrings;

ANS($pchk->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The line segment \(PQ\) has length $twor0, so it must be a diameter of
the circle.  The center of the circle is therefore the midpoint of
\(PQ\), which is the point \(($r0,0)\). The  upper arc of the circle between
\(P\) and \(Q\) can be parameterized as follows:
\[
\vec r (t) = $r0\vec i +  $r0(-\cos t \vec i + \sin t \vec j),
\quad 0 \le t \le \pi.
\]
The  lower arc can be parameterized as follows:
\[
\vec r (t) = $r0\vec i + $r0(\cos t \vec i + \sin t \vec j),
\quad \pi \le t \le 2\pi.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
