# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/02/10


## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.1')
## Section2('17.1')
## Section3('17.1')
## Problem1('63')
## Problem2('73')
## Problem3('87')
## KEYWORDS('parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserParametricLine.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"}
);
Context()->variables->are( t=>'Real' );
Context()->constants->add( a=>1.404, b=>2.71828, c=>1.1771 );
Context()->flags->set(
    formatStudentAnswer=>'parsed',
    formatCorrectAnswer=>'parsed',
);

$showPartialCorrectAnswers = 1;

## pick a point
$x0 = random(1,6,1);
$y0 = random(1,6,1);
$z0 = random(1,6,1);
$pt0 = Point( $x0, $y0, $z0 );

## the direction vector is
$v0 = Vector( Compute("a"), Compute("b"), Compute("c") );

$line = ParametricLine( "$pt0 + t*$v0" );

$abc = MultiAnswer( Compute("$x0"), Compute("$y0"), Compute("$z0"),
		    Compute("-$x0"), Compute("-$y0"), Compute("-$z0") )->with(
    singleAnswer=>0,
    checker=>sub {
	my ( $c, $s, $ans ) = @_;
	## the three values must be in the ratios x0:y0:z0 and the two
	##    pairs of values must be different
	my $validRatio1 = 0;
	my $validRatio2 = 0;
	my $different = 0;

	my ( $stuA1, $stuB1, $stuC1, $stuA2, $stuB2, $stuC2 ) = @{$s};
	my $ratio1 = $stuA1/$x0;
	if ( $stuB1 == $ratio1*$y0 && $stuC1 == $ratio1*$z0 ) {
	    $validRatio1 = 1;
	}
	my $ratio2 = $stuA2/$x0;
	if ( $stuB2 == $ratio2*$y0 && $stuC2 == $ratio2*$z0 ) {
	    $validRatio2 = 1;
	}
	$different = 1 if ( $stuA1 != $stuA2 );

	if ( $validRatio1 && $validRatio2 && $different ) {
	    return [ 1,1,1, 1,1,1 ];
	} elsif ( $validRatio1 && $validRatio2 ) {
	    $ans->setMessage( 4, "The two pairs of (a,b,c) values you " .
			      "give must be different." );
	    return [ 1,1,1, 0,0,0 ];
	} elsif ( $validRatio1 ) {
	    return [ 1,1,1, 0,0,0 ];
	} elsif ( $validRatio2 ) {
	    return [ 0,0,0, 1,1,1 ];
	} else {
	    return [ 0,0,0, 0,0,0 ];
	}
    }
);

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

${BBOLD}(a)$EBOLD
Find a parametric equation for the line through the point \( $pt0 \)
and in the direction of \( a\vec i +b\vec j+c\vec k \).
$BR
\( \vec r(t) = \) \{ ans_rule(55) \}

$PAR
${BBOLD}(b)$EBOLD
Find conditions on \( a,b,c \) so that the line you found in part (a) goes
through the origin.  (Be sure you can give a reason for your answer.)  Then
use your work to give two distinct triples \(a,b,c\) that result in the
line passing through the origin.
$BR
Choice 1: \(a = \) \{ $abc->ans_rule(10) \}, \(b = \) \{ $abc->ans_rule(10) \},
\(c = \) \{ $abc->ans_rule(10) \};
$BR
Choice 2: \(a = \) \{ $abc->ans_rule(10) \}, \(b = \) \{ $abc->ans_rule(10) \},
\(c = \) \{ $abc->ans_rule(10) \}.

END_TEXT
Context()->normalStrings;

ANS($line->cmp() );
ANS($abc->cmp() );

$x0d = ( $x0 == 1 ) ? '' : $x0;
$y0d = ( $y0 == 1 ) ? '' : $y0;
$z0d = ( $z0 == 1 ) ? '' : $z0;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
Parametric equations are
\[ x = $x0 + a t,\quad y = $y0 + b t,\quad z = $z0 + c t, \]
or
\[ \vec r(t) = ($x0 + a t)\,\vec i + ($y0 + b t)\,\vec j +
    ($z0 + c t)\,\vec k. \]
$PAR

${BBOLD}(b)$EBOLD
The line goes through the origin if the position vector
\( $x0d \vec i + $y0d \vec j + $z0d \vec k\) is parallel to the vector
\(a\vec i +b\vec j +c\vec k\). This occurs if \(a,b,c\) are
in the ratio \($x0:$y0:$z0\); that is if
\[ \frac a{$x0} =\frac b {$y0} =\frac c{$z0}. \]
Any values satisfying this are appropriate.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
