## DBsubject(Number theory)
## DBchapter(Divisibility)
## DBsection(GCDs and LCMs)
## Institution(University of Massachusetts Amherst)
## Author(Siman Wong, Patrick O'Brien, and Dan Nichols)
## MLT(GCD_as_LinearCombo)
## Level(2)
## MO(1)
## KEYWORDS('GCD, Euclidean algorithm')



DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);




TEXT(beginproblem());
$showPartialCorrectAnswers = 0;


$z = random(0,99,1);


@a = (64, 66, 54, 52, 73, 48, 11, 54, 13, 42, 19, 35, 79, 73, 52, 36, 35, 56, 27, 44, 14, 68, 19, 47, 17, 28, 61, 70, 31, 47, 32, 24, 54, 40, 36, 23, 61, 73, 19, 5, 69, 17, 9, 77, 18, 22, 32, 79, 47, 49, 65, 22, 50, 27, 46, 21, 68, 58, 74, 72, 57, 56, 36, 28, 46, 76, 46, 66, 21, 80, 54, 57, 66, 60, 45, 51, 42, 72, 10, 74, 70, 46, 24, 40, 39, 16, 51, 50, 48, 70, 55, 54, 74, 27, 39, 34, 57, 44, 36, 34);

@b = (23, 37, 7, 23, 50, 19, 37, 7, 16, 19, 7, 26, 29, 50, 49, 41, 33, 25, 37, 17, 11, 7, 16, 43, 26, 25, 41, 11, 20, 13, 47, 41, 13, 19, 25, 48, 27, 13, 27, 43, 43, 41, 41, 8, 7, 31, 39, 44, 5, 38, 40, 28, 76, 33, 36, 33, 54, 64, 28, 33, 36, 22, 14, 72, 54, 64, 68, 46, 15, 62, 69, 75, 48, 76, 35, 69, 50, 38, 76, 60, 74, 30, 62, 65, 45, 38, 42, 54, 57, 48, 70, 22, 14, 39, 48, 74, 33, 10, 15, 54);

@x = (9, -14, 3, 4, -13, 2, -10, 3, 5, 5, 3, 3, -11, -13, -16, 8, -16, -4, 11, -5, 4, 3, -5, 11, -3, -8, -2, 3, -9, 5, -22, 12, -6, -9, -9, 23, 4, 5, 10, -17, 5, -12, -9, -3, 2, -7, 11, -5, -2, 7, -3, -5, -3, 5, -7, -3, 4, -11, -3, -5, -5, 2, 2, -5, -7, -5, 3, 7, -2, 7, 9, 4, 3, -5, -3, -4, 6, 9, -15, 13, 18, 2, 13, 5, 7, -7, 5, 13, 6, 11, -5, -2, -3, 3, 5, -13, -4, -2, -2, 8);

@y = (-25, 25, -23, -9, 19, -5, 3, -23, -4, -11, -8, -4, 30, 19, 17, -7, 17, 9, -8, 13, -5, -29, 6, -12, 2, 9, 3, -19, 14, -18, 15, -7, 25, 19, 13, -11, -9, -28, -7, 2, -8, 5, 2, 29, -5, 5, -9, 9, 19, -9, 5, 4, 2, -4, 9, 2, -5, 10, 8, 11, 8, -5, -5, 2, 6, 6, -2, -10, 3, -9, -7, -3, -4, 4, 4, 3, -5, -17, 2, -16, -17, -3, -5, -3, -6, 3, -6, -12, -5, -16, 4, 5, 16, -2, -4, 6, 7, 9, 5, -5);

@g = (1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 5, 2, 2, 3, 2, 3, 2, 2, 2, 3, 3, 2, 2, 4, 2, 4, 2, 2, 3, 2, 3, 3, 6, 4, 5, 3, 2, 2, 2, 2, 2, 2, 2, 5, 3, 2, 3, 2, 3, 2, 5, 2, 2, 3, 3, 2, 3, 2, 3, 2);

# multiAnswer to verify two distinct solutions (x,y) with |x| and |y| < 1000
$xyAnswer = MultiAnswer( $x[ $z ], $y[ $z ], $x[ $z ] + $b[ $z ], $y[ $z ] - $a[ $z ] )->with(
	singleResult => 0,
	checker => sub
	{
		my( $correct, $student, $self ) = @_;
		my( $x1, $y1, $x2, $y2 ) = @$student;
		
		my @result = ( 0, 0, 0, 0 );
		
		if( grep { $_ >= 1000 || $_ <= -1000 } @$student )
		{
			$self->setMessage( 1, "one of the numbers you entered is too large" );
			return ~~@result;
		}
		
		@result[ 0, 1 ] = ( 1, 1 ) if $x1 * $a[ $z ] + $y1 * $b[ $z ] == $g[ $z ];
		
		if( $x1 == $x2 || $x1 == $y2 || $y1 == $x2 || $y1 == $y2 )
		{
			$self->setMessage( 3, "Your two solutions must be different" );
			return ~~@result;
		}
		
		@result[ 2, 3 ] = ( 1, 1 ) if $x2 * $a[ $z ] + $y2 * $b[ $z ] == $g[ $z ];
		
		return ~~@result;
	}
);




TEXT(EV3(<<'EOT'));

Find two integer pairs of the form \((x, y)\) with \(|x| < 1000\) such that \($a[$z] x + $b[$z] y = \gcd($a[$z],$b[$z])\) $BR $BR

\( (x_1,y_1) = ( \) \{ $xyAnswer->ans_rule(5) \} \( , \) \{ $xyAnswer->ans_rule( 5 ) \} \( ) \) $BR
\( (x_2,y_2) = ( \) \{ $xyAnswer->ans_rule(5) \} \( , \) \{ $xyAnswer->ans_rule( 5 ) \} \( ) \) $BR



EOT

ANS( $xyAnswer->cmp() );


ENDDOCUMENT();       
