# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# 
# Use the associative property of addition.
#
# Last edited: Jordan 8/23/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions','Equivalent Expressions','Reordering and Regrouping')
## KEYWORDS('associative')
## DBCCSS('')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
);

   

    

##############################################

$var1 = RandomVariableName();
do {$var2 = RandomVariableName();} until ($var2 ne $var1);

Context("Numeric");
Context()->variables->are($var1=>'Real',$var2=>'Real');
Context()->noreduce('(-x)-y','(-x)+y');
Context()->lists->set(AbsoluteValue => {open => "(", close => ")"});
Context()->parens->set("(" => {removable=>0, type => "AbsoluteValue", close => ")"});

$a = random(1,100,1);
Context()->variables->set($var1 => {limits => [-$a-2,-$a+2]});
Context()->variables->set($var2 => {limits => [-$a-2,-$a+2]});

$giv = Compute("$var1+($a+$var2)");
$amb = Compute("$var1+$a+$var2");
$ans = Compute("($var1+$a)+$var2");



##############################################
TEXT(beginproblem());

BEGIN_PGML

Use the associative property of addition to write an equivalent expression.  

    [`[$giv]`]

    [____________]


END_PGML
##############################################

ANS($ans->cmp(checker => sub {
      my ($correct,$student,$ans) = @_;

      # Because of the absolute value hack for parens, make sure to test at both pos and neg input:
     $giv->{test_at} = [[-$a-1,-$a-1],[-$a-1,-$a+1],[-$a+1,-$a-1],[-$a+1,-$a+1]];
     $student ->{test_at} = [[-$a-1,-$a-1],[-$a-1,-$a+1],[-$a+1,-$a-1],[-$a+1,-$a+1]];
     $amb->{test_at} = [[-$a-1,-$a-1],[-$a-1,-$a+1],[-$a+1,-$a-1],[-$a+1,-$a+1]];
     $correct->{test_at} = [[-$a-1,-$a-1],[-$a-1,-$a+1],[-$a+1,-$a-1],[-$a+1,-$a+1]];

      Value->Error("This expression is algebraically equivalent to the given expression, but it is not what an application of the associative property of addition would give you") if (($student == $giv) or ($student == $amb) or (($student == $correct) and ($student ne $correct)));
      return ($student eq $correct);
    }));

##############################################

BEGIN_PGML_SOLUTION

To use the associative property of addition, we leave terms written in their original order, but change the grouping symbols so that a different addition has higher priority:

    [``\begin{aligned}
    [$giv] & =[$ans]
    \end{aligned}``]

END_PGML_SOLUTION

ENDDOCUMENT();


