# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# 
# Use the commutative property of addition.
#
# Last edited: Jordan 8/23/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions','Equivalent Expressions','Reordering and Regrouping')
## KEYWORDS('commutative')
## DBCCSS('')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
);

   

    

##############################################

do {$var = RandomVariableName();} until ($var ne uc($var));

Context("Numeric");
Context()->variables->are($var=>'Real');
Context()->reduction->set('(-x)-y'=>0);
Context()->reduction->set('(-x)+y'=>0);

$a = random(1,100,1);

$giv = Formula("$var+$a");
$ans = Formula("$a+$var");


##############################################
TEXT(beginproblem());

BEGIN_PGML

Use the commutative property of addition to write an equivalent expression to [`[$giv]`].

    [____________]


END_PGML
##############################################

ANS($ans->cmp(checker => sub {
      my ($correct,$student,$ans) = @_;
      Value->Error("This is the original expression") if ($student eq $giv);
      Value->Error("This expression is algebraically equivalent to the given expression, but it is not what an application of the commutative property of addition would give you") if (($student == $correct) and ($student ne $correct));
      return ($student eq $correct);
    }));

##############################################

BEGIN_PGML_SOLUTION

To use the commutative property of addition, we change the order in which two terms are added:

    [``\begin{aligned}
    [$giv] & =[$ans]
    \end{aligned}``]

END_PGML_SOLUTION

ENDDOCUMENT();


