# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Simplify an algebraic expression with context
#
# Last edited: Jordan 8/23/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('')
## KEYWORDS('')
## DBCCSS('')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "contextLimitedPolynomial.pl",
  "parserAssignment.pl",
  "answerHints.pl",
,
  "PGcourse.pl",
);

##############################################

$a = random(0.3,0.6,0.01);
$b = random(3,6,0.1);

$c = random(0.3,0.6,0.01);
$d = random(3,6,0.1);

$e = random(0.01,0.04,0.01);
$f = random(0.1,0.2,0.1);

$apc = $a+$c;
$bpd = $b+$d;

$apcpe = $apc+$e;
$bpdpf = $bpd+$f;

Context("LimitedPolynomial-Strict");
Context()->variables->are(T=>'Real', x=>'Real');
parser::Assignment->Allow;
Context()->{error}{msg}{"The left side of an assignment must be a variable or function"} 
  = "Your answer should be in the form T = $US$US$US";

$ans = Formula("$apcpe x + $bpdpf");
$ansEq=Formula("T=$apcpe x + $bpdpf");


##############################################
TEXT(beginproblem());

BEGIN_PGML

The number of students enrolled in math courses at Portland Community College has grown over the years. The formulas [```\begin{aligned}M&=[$a]x+[$b] & W&=[$c]x+[$d] & N&=[$e]x+[$f]\end{aligned}```] describe the numbers (of thousands) of men, women, and gender-non-binary students enrolled in math courses at PCC [`x`] years after 2005. (Note this is an exercise using randomized numbers, not actual data.) Give a simplified formula for the total number [`T`] of thousands of students at PCC taking math classes [`x`] years after 2005. Be sure to give the entire formula, starting with [|T=|]*.

     [_________________________]

END_PGML
##############################################

ANS($ansEq->cmp(cmp_class=>"a formula of the form T = $US$US$US", 
                checker => sub {
        my ($correct,$student,$ans) = @_;
        if ($correct->type eq "Assignment") {
          my ($svar,$sform) = $student->value; 
          my $t = Formula("$sform");
        }
        return $correct == $student;
      },
) ->
             withPostFilter(AnswerHints( 
             "$ans" => ["You have the right expression, but the whole formula should begin with T = $US$US$US" , replaceMessage => 1] 
             ))
             );

##############################################


BEGIN_PGML_SOLUTION

The total number of students [`T`] can be found by adding the men, women, and gender-non-binary students together. In the steps that follow, precisely one algebraic property at a time is applied. In practice, you might execute several of these steps together at once.

    [`\begin{aligned}
T & = M+W+N\\
 & = ([$a]x+[$b]) + ([$c]x+[$d]) + ([$e]x+[$f])\\
 & = ([$a]x+[$b]) + ([$d]+[$c]x) + ([$e]x+[$f])\\
 & = [$a]x+([$b] + ([$d]+[$c]x)) + ([$e]x+[$f])\\
 & = [$a]x+(([$b] + [$d])+[$c]x) + ([$e]x+[$f])\\
 & = [$a]x+([$bpd]+[$c]x) + ([$e]x+[$f])\\
 & = [$a]x+([$c]x+[$bpd]) + ([$e]x+[$f])\\
 & = ([$a]x+[$c]x)+[$bpd] + ([$e]x+[$f])\\
 & = ([$a]+[$c])x+[$bpd] + ([$e]x+[$f])\\
 & = ([$apc]x+[$bpd]) + ([$e]x+[$f])\\
 & = ([$apc]x+[$bpd]) + ([$f]+[$e]x)\\
 & = [$apc]x+([$bpd] + ([$f]+[$e]x))\\
 & = [$apc]x+(([$bpd] + [$f])+[$e]x)\\
 & = [$apc]x+([$bpdpf]+[$e]x)\\
 & = [$apc]x+([$e]x+[$bpdpf])\\
 & = ([$apc]x+[$e]x)+[$bpdpf]\\
 & = ([$apc]+[$e])x+[$bpdpf]\\
 & = [$ans]
\end{aligned}`]

So the formula for the total number of students taking math at PCC [`x`] years after 2005 is [`[$ansEq]`].

END_PGML_SOLUTION

ENDDOCUMENT();


