## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Lines)
## Author(Justin Freeman)
## MLT(lin_find_vec_eqn_two_points)
## Level(3)
## MO(1)

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##########################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[0,10]});

$a1 = random(-5,5,1);
$a2 = random(-5,5,1);
$a3 = random(-5,5,1);

$c1 = non_zero_random(-5,5,1);
$c2 = non_zero_random(-5,5,1);
$c3 = non_zero_random(-5,5,1);


$b1 = $a1+$c1;
$b2 = $a2+$c2;
$b3 = $a3+$c3;

$x = Formula("$a1 + ($c1)*t");
$y = Formula("$a2 + ($c2)*t");
$z = Formula("$a3 + ($c3)*t");

$multians = MultiAnswer($x, $y, $z)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $xstu, $ystu, $zstu ) = @{$student};
      return 0 unless $xstu->isFormula;
      if (($c2*$xstu-$c1*$ystu==$c2*$a1-$c1*$a2) && ($c3*$xstu-$c1*$zstu==$c3*$a1-$c1*$a3)) {
            return 1;
      } else {
            return 0;
      }
  }
);


##########################
#  Main Text 

Context()->texStrings;
BEGIN_TEXT
Give a parametric equation of the line which passes through \(A($a1,$a2,$a3)\) and \(B($b1,$b2,$b3)\).  Use \( t \) as the 
parameter for all of your answers.
$BR
$BR
\( x(t) = \) \{$multians->ans_rule(20)\}
\{ AnswerFormatHelp("formulas") \}
$BR
\( y(t) = \) \{$multians->ans_rule(20)\}
\{ AnswerFormatHelp("formulas") \}
$BR
\( z(t) = \) \{$multians->ans_rule(20)\}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


########################
#  Answer Evaluation

$showPartialCorrectAnswers = 0;

ANS($multians->cmp());



Context()->normalStrings;

COMMENT('MathObject version.');

ENDDOCUMENT();
