## DESCRIPTION
## Algebra: Complete the square (quadratic equation)
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Completing the square)
## Level(2)
## KEYWORDS('algebra', 'rules of exponents')


###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGcourse.pl",
"PGML.pl",
"scaffold.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());


###########################
#  Setup

$showPartialCorrectAnswers = 1;

Context("Numeric");

$sign = list_random(-1,1);
$b1abs = random(4,20,2);
$b1 = $sign*$b1abs;

$a1 = Compute("$b1/2");
$a1abs = Compute("$b1abs/2");
$a1neg = Compute("-($a1)");
$d1 = Compute("($a1)^2");

$r1 = random(2,7,1);

$r1s = ($r1)**2;

$c1 = $d1 - $r1s;

$lhs = Formula("x^2 + ($b1)*x + ($c1)")->reduce;;

$lhsstart = Formula("x^2 + ($b1)*x")->reduce;

$squareof = Formula("x + ($a1)")->reduce;

if ($sign == -1) {
    $perfsquare = "x^2 - 2ax + a^2 = (x - a)^2";
} else {
    $perfsquare = "x^2 + 2ax + a^2 = (x + a)^2";
}

$sol1 = Compute("-($a1)-$r1");
$sol2 = Compute("-($a1)+$r1");

$mp = MultiAnswer($sol1, $sol2)->with(
   singleResult => 0,
   allowBlankAnswers => 1,
   checker => sub {
     my $correct= shift; my $student = shift;
     my ($ca,$cb) = @$correct;
     my ($sa,$sb) = @$student;
     return (1,1) if ($ca == $sa and $cb == $sb) or
                     ($ca == $sb and $cb == $sa);
     return (1,0) if $ca == $sa or $cb == $sa;
     return (0,1) if $ca == $sb or $cb == $sb; 
     return (0,0);
   },
 );

###########################
#  Main text

BEGIN_PGML

Fill in the blanks to solve the equation

[` [$lhs] = 0 `]

by completing the square:

END_PGML

Scaffold::Begin(
    can_open => "when_previous_correct",
    is_open => "correct_or_first_incorrect",
);
Section::Begin("Step 1: Complete the square.");

BEGIN_PGML
Add a number to both sides so the left side will become a perfect square:

[` [$lhsstart] + `][_____]{$d1} [` = `][____]{$r1s}
END_PGML

#Hint for step 1:
$showHint = 2;

BEGIN_PGML_HINT
You want to make the left side look like [` x^2 + 2ax + a^2 `] or [` x^2 - 2ax + a^2 `].
Comparing that to the original left hand side: [` [$lhs] `], what should [`a`] be? 
END_PGML_HINT

BEGIN_PGML_SOLUTION
Comparing the left hand side of the equation: [`[$lhs]`] with the perfect square 
equation [` [$perfsquare] `], we can see that [` 2a = [$b1abs] `].  Then 

[`` a = \frac{[$b1abs]}{2} = [$a1abs] ``]

and

[`` d = a^2 = [$a1abs]^2 = [$d1] ``]

To complete the square, we need to add [` [$d1] - ([$c1]) = [$r1s] `] to both sides of the equations. 
After completing the square, the equation will look like

[`` [$lhsstart] + [$d1] = [$r1s] ``]
END_PGML_SOLUTION

Section::End();

######################## Section 2:

Section::Begin("Step 2: Rewrite the left side as perfect square.");

BEGIN_PGML
Next fill in the blank so that the left side of the equation is 
written as an actual square:

[` ( `][______]{$squareof}[` )^2 = [$r1s] `]

END_PGML

#hint:
$showHint = 2;

BEGIN_PGML_HINT
You already wrote the left side as [` x^2 + 2ax + a^2 `] or
[` x^2 - 2ax + a^2 `].  In this atep you have to write it as 
[` (x + a)^2 `] or [` (x - a)^2 `].
END_PGML_HINT

BEGIN_PGML_SOLUTION

We already figured out in Step 1 that [` a = [$a1abs] `], so the perfect square on the left side is [` ([$squareof])^2 `].  The whole equations is now

    [` ([$squareof])^2 = [$r1s] `]

END_PGML_SOLUTION

Section::End();

######################## Section 3:

Section::Begin("Step 3: Use the square root property");

BEGIN_PGML

Now use the square root property and solve for [`x`]:

[` x = `][________]{-$a1}[` \pm `][________]{$r1}

The two solutions are [________]{$mp} and [________]{$mp}.

END_PGML

#hint:
$showHint = 2;

BEGIN_PGML_HINT
The square root property says that the equation 

    [` x^2 = d `]

is equivalent to the pair of eqautions

    [` x = \pm\sqrt{d} `]
END_PGML_HINT

BEGIN_PGML_SOLUTION
According to the square root property, the equation 

    [` ([$squareof])^2 = [$r1s] `]

is equivalent to 

    [` [$squareof] = \pm\sqrt{[$r1s]} = \pm [$r1] `]

Solving that for [`x`] will give us

    [` x = [$a1neg] \pm [$r1] `]

The two solutions will be [$sol1] and [$sol2].
END_PGML_SOLUTION

Section::End();
Scaffold::End();

ENDDOCUMENT();
