## DESCRIPTION
## Algebra: Complete the square (quadratic equation with leading coefficient)
# written by Jan Hlavacek (jhlavace@svsu.edu)
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Completing the square)
## Level(2)
## KEYWORDS('algebra', 'quadratic equation', 'completing the square')


###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGcourse.pl",
"PGML.pl",
"contextFraction.pl",
"scaffold.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());


###########################
#  Setup

$showPartialCorrectAnswers = 1;

Context("LimitedFraction")->flags->set(showMixedNumbers => 0);

$b = random(2,15,1);
do { $a = random(2,9,1); } until ( gcd($a, $b) == 1 );

$c = non_zero_random(-20,$b**2/(4*$a)-1,1);

$r1 = 4*$a*$c;

$rtop = $b**2 - $r1;

$a3bot = 2*$a;
$d3top = $b**2;
$d3bot = $a3bot**2;

$b3 = Compute("$b/$a");
$c3 = Compute("$c/$a");
$a3 = Compute("$b/$a3bot");
$d3 = Compute("$d3top/$d3bot");
$r3 = Compute("$rtop/$d3bot");

Context("Numeric");

$squareof = Formula("x - ($a3)")->reduce;

$sol1 = Real("$a3-sqrt($r3)");
$sol2 = Real("$a3+sqrt($r3)");

$mp = MultiAnswer($sol1, $sol2)->with(
   singleResult => 0,
   allowBlankAnswers => 1,
   checker => sub {
     my $correct= shift; my $student = shift;
     my ($ca,$cb) = @$correct;
     my ($sa,$sb) = @$student;
     return (1,1) if ($ca == $sa and $cb == $sb) or
                     ($ca == $sb and $cb == $sa);
     return (1,0) if $ca == $sa or $cb == $sa;
     return (0,1) if $ca == $sb or $cb == $sb; 
     return (0,0);
   },
 );

###########################
#  Main text

BEGIN_PGML

Fill in the blanks to solve the equation

[` [$a]x^2 - [$b]x + [$c] = 0 `]

by completing the square:

END_PGML

Scaffold::Begin(
    can_open => "when_previous_correct",
    is_open => "correct_or_first_incorrect",
);
Section::Begin("Step 1: Complete the square.");

BEGIN_PGML
Divide by [`a`] and complete the square (fill in the blanks as integers or fractions in the simplest form):

[` x^2 - `][______]{$b3}[` x + `][_____]{$d3} [` = `][____]{$r3}
END_PGML

#Hint for step 1:
$showHint = 2;

BEGIN_PGML_HINT
First you want  to divide both side of the equation by [$a].
END_PGML_HINT

BEGIN_PGML_SOLUTION
First divide both sides of the equations by [$a] to get

[` x^2 - [$b3]x + [$c3] = 0 `]

Subtracting [` [$c3] `] from both sides and adding [` [$d3] `] to both sides gives you

[` x^2 - [$b3]x + [$d3] = [$r3] `]

The left side is a complete square.
END_PGML_SOLUTION

Section::End();

######################## Section 2:

Section::Begin("Step 2: Rewrite the left side as perfect square.");

BEGIN_PGML
Next fill in the blank so that the left side of the equation is 
written as an actual square:

[` \big( `][______]{$squareof}[` \big)^2 = [$r3] `]

END_PGML

#hint:
$showHint = 2;

BEGIN_PGML_HINT
You already wrote the left side as [` x^2 - 2ax + a^2 `].  In this step you
have to write it as [` (x - a)^2 `].
END_PGML_HINT

BEGIN_PGML_SOLUTION

We already figured out in Step 1 that [` a = [$a3] `], so the perfect square on the left side is 

[`` \left(x - [$a3]\right)^2 ``].  

The whole equations is now

[`` \left(x - [$a3]\right)^2 = [$r3] ``].  
END_PGML_SOLUTION

Section::End();

######################## Section 3:

Section::Begin("Step 3: Use the square root property");

BEGIN_PGML

Now use the square root property and solve for [`x`]:

The two solutions are [________________]{$mp} and [________________]{$mp}. (You can use sqrt() for square roots, if necessary).

END_PGML

#hint:
$showHint = 2;

BEGIN_PGML_HINT
The square root property says that the equation 

    [` x^2 = d `]

is equivalent to the pair of equations

    [` x = \pm\sqrt{d} `]
END_PGML_HINT

BEGIN_PGML_SOLUTION
According to the square root property, the equation 

    [` (x - [$a3])^2 = [$r3] `]

is equivalent to 

    [`  x - [$a3] = \pm\sqrt{[$r3]} `]

Solving that for [`x`] will give us

    [` x = [$a3] \pm \sqrt{[$r3]} `]

The two solutions will be [$sol1] and [$sol2].
END_PGML_SOLUTION

Section::End();
Scaffold::End();

ENDDOCUMENT();
