## DESCRIPTION
## Answer is a pair of matrices that require a custom answer checker
## ENDDESCRIPTION


## DBsubject(WeBWorK)
## DBchapter(WeBWorK tutorial)
## DBsection(Fort Lewis tutorial 2011)
## Date(06/30/2014)
## Institution(Hope College)
## Author(Paul Pearson)
## MO(1)
## KEYWORDS('matrix','custom answer checker','multianswer')


#############################
#  Initialization

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 


#############################
#  Setup

Context('Matrix');

$A = Matrix([[1,1],[0,1]]);
$B = Matrix([[1,0],[1,1]]);

$multians = MultiAnswer($A, $B)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @s = @{$student};
      $s0 = Matrix($s[0]); 
      $s1 = Matrix($s[1]);
      return $s0 * $s1 != $s1 * $s0;
  }
);


############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Give an example of two \( 2 \times 2 \) matrices \( A \) and \( B \)
such that \( AB \ne BA \).
$BR
$BR
\( A = \)
\{ $multians->ans_array(5) \}
$BR
$BR
\( B = \)
\{ $multians->ans_array(5) \}
END_TEXT
Context()->normalStrings;


############################
#  Answer evaluation

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

;

ENDDOCUMENT();
