## DESCRIPTION
## Answer is a pair of matrices that require a custom answer checker
## ENDDESCRIPTION

## DBsubject(WeBWorK)
## DBchapter(WeBWorK tutorial)
## DBsection(PGML tutorial 2015)
## Date(06/01/2015)
## Institution(Hope College)
## Author(Paul Pearson)
## MO(1)
## KEYWORDS('matrix','custom answer checker','multianswer')


#############################
#  Initialization

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGML.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

install_problem_grader(~~&std_problem_grader);


#############################
#  Setup

Context('Matrix');

$A = Matrix([[1,1],[0,1]]);
$B = Matrix([[1,0],[1,1]]);

$multians = MultiAnswer($A, $B)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @s = @{$student};
      $s0 = Matrix($s[0]); 
      $s1 = Matrix($s[1]);
      return $s0 * $s1 != $s1 * $s0;
  }
);


############################
#  Main text

BEGIN_PGML
Give an example of two [` 2 \times 2 `] matrices [` A `] and [` B `]
such that [` AB \ne BA `].
 
[` A = `] [_____]*{$multians}

[` B = `] [_____]*{$multians}
END_PGML

COMMENT('Uses PGML.');

ENDDOCUMENT();
