##DESCRIPTION
##  Number theory: primitive root
##ENDDESCRIPTION

##KEYWORDS('number theory', 'primitive element')

## DBsubject('Number Theory')
## DBchapter('Congruences')
## DBsection('Primitive elements')
## Date('7/25/2018')
## Author('John Perry')
## Institution('University of Southern Mississippi')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

@primes = ( 17, 19, 23 );
$p = @primes[ random(0,2) ];

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

Find a primitive element of \( \mathbb Z_{$p} \).
$PAR
\{ ans_rule(5) \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS( Compute(0) -> cmp(
  checker => sub {
    my ( $junk, $a, $hash ) = @_;
    $result = 1;
    for ( $i = 1; $result && $i < $p - 1; ++$i ) {
      if ( ( $a ** $i ) % $p == 1 ) {
        $result = 0;
        $b = $p - 1;
        Value->Error( "\($a^{$i}=1\), need \(a^k\neq1\) when \(k<$b\)" );
      }
    }
    return $result;
  }
) );


ENDDOCUMENT();        
