## DESCRIPTION
## Rational Expressions: Multiplication and Division
## ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: multiply, divide)
## Date(Spring 2012)
## Institution(University of Missouri)
## MLT(RatExp2)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.6')
## Problem1('')
## KEYWORDS('Algebraic fractions')
## Adapted from a problem written by Paul Pearson (Fort Lewis)


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


###################################
# Setup

Context("Numeric");

Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real");

$var = list_random('x','y','z','w','p','q','r','s','t','u','v');

do {
$a = random(-1,1,2)*random(2,9);
do {$b = random(-1,1,2)*random(2,9)} until (abs($a+$b) > 1);
$c = random(-1,1,2)*random(2,9);
do {$d = random(-1,1,2)*random(2,9)} until (abs($c+$d) > 1 && abs($b + $d) > 1);
do {$e = random(-1,1,2)*random(2,9)} until (gcd ($c, $e) == 1 && abs($a + $e) > 1);

$n1 = $a + $b;
$n2 = $a * $b;
$n3 = $c + $d;
$n4 = $c * $d;
$n5 = $b + $d;
$n6 = $b * $d;
$n7 = $a + $e;
$n8 = $a * $e;
} until ($n1 != $n3 && $n1 != $n5 && $n1 != $n7 && $n3 != $n5 && $n3 != $n7 && $n5 != $n7);

$frac = "\frac{$var^2 + $n1 $var + $n2}{$var^2 + $n5 $var + $n6 } \cdot \frac{$var^2 + $n3 $var + $n4}{$var^2 + $n7 $var + $n8}";

$num = Formula("$var + $c")->reduce;
$den = Formula("$var + $e")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}



###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Perform the indicated operation, and simplify if possible.  Assume any factors you cancel are not zero.
$BR
$BR
$BCENTER
$displayfrac
$ECENTER

END_TEXT

Context()->normalStrings;



###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

COMMENT('MathObject version');
ENDDOCUMENT();
