## DESCRIPTION
## Higher order ODEs: constant coefficient second order nonhomogeneous linear ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Undetermined coefficients)
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('2.5')
## Problem1('3')
## KEYWORDS('differential equations','second order linear ODE','constant coefficients','nonhomogeneous')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserAssignment.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup 1

Context("Numeric");
Context()->variables->add(
#'c1'=>"Real",'c2'=>"Real",
y=>"Real",
'a'=>'Parameter','b'=>'Parameter'
);

$k = random(2,5,1);
$k2 = 2 * $k;
$kk = $k * $k;

$diffeq = "y^{\,\prime\prime} - $k2 y^{\,\prime} + $kk y = e^{$k x}"; # tex

$yparticular = Compute("1/2 x^2 e^($k x) + (a x + b) e^($k x)");


######################
#  Main text 1

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

$ITEM Find a particular solution to the 
nonhomogeneous differential equation \( $diffeq \).
$BR
$BR
\( y_p = \) 
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;

####################################
#  Answer evaluation 1

$showPartialCorrectAnswers = 1;

ANS( $yparticular->cmp() );



####################################
#  Setup 2

Context("Numeric");
Context()->variables->add(
'c1'=>"Real",'c2'=>"Real",y=>"Real",
);
Context()->variables->set(
'c1'=>{limits=>[2,4]},
'c2'=>{limits=>[2,4]}
);


$yhomogeneous = Compute("(c1 x + c2) e^($k x)");
$ygeneral = Compute("1/2 x^2 e^($k x) + (c1 x + c2) e^($k x)");


Context()->texStrings;
BEGIN_TEXT
$ITEMSEP
$ITEM Find the most general solution to the 
associated homogeneous differential equation.  
Use \( c_1 \) and \( c_2 \) in your answer to denote 
arbitrary constants and enter them as ${BTT}c1${ETT} 
and ${BTT}c2${ETT}.
$BR
$BR
\( y_h = \)
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}


$ITEMSEP
$ITEM Find the most general solution to the 
original nonhomogeneous differential equation.  
Use \( c_1 \) and \( c_2 \) in your answer to denote 
arbitrary constants.
$BR
$BR
\( y = \)
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation


ANS( $yhomogeneous->cmp(
  checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $stu   = Formula($student);

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('c1'))==Formula(0) || 
      Formula($stu->D('c2'))==Formula(0) 
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(1.24);

    my $a11 = $stu->eval('c1'=>1,'c2'=>0,x=>$x,y=>0);
    my $a12 = $stu->eval('c1'=>0,'c2'=>1,x=>$x,y=>0);
 
    my $a21 = $stu->D('x')->eval('c1'=>1,'c2'=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval('c1'=>0,'c2'=>1,x=>$x,y=>0);

    # my $wronskian = $a11*$a22 - $a21*$a12;
    # Value->Error("Your functions are not linearly independent or your answer is not complete") 
    # if ($wronskian==Real(0));

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($a11 * $a22 == $a21 * $a12);


    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu1 = Formula($stu->D('x'));
    my $stu2 = Formula($stu->D('x','x'));
    return $stu2 + $kk * $stu == $k2 * $stu1;

#    # old code that didn't work well
#    # apparently checking whether something is zero is not numerically stable
#    my $stuDE = Formula($stuxx - $k2*$stux + $kk*$stu)
#    ->with(test_at=>[[1,-1,1,-1],[1,2,3,1],[2,1,1,3],[1,2,4,-1],[2,4,1,-3]]);
#    return ($stuDE==Formula("0"));

    # new code that seems to work very reliably is obtained by simply rewriting 
    # the equation $stuDE = 0, so that both sides of the equation vary.
    # return $stuxx + $kk * $stu == $k2 * $stux;

}));





ANS( $ygeneral->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $stu   = Formula($student);

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('c1'))==Formula(0) || 
      Formula($stu->D('c2'))==Formula(0) 
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(1.24);

    my $a11 = $stu->eval('c1'=>1,'c2'=>0,x=>$x,y=>0);
    my $a12 = $stu->eval('c1'=>0,'c2'=>1,x=>$x,y=>0);
 
    my $a21 = $stu->D('x')->eval('c1'=>1,'c2'=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval('c1'=>0,'c2'=>1,x=>$x,y=>0);

    # my $wronskian = $a11*$a22 - $a21*$a12;
    # Value->Error("Your functions are not linearly independent or your answer is not complete") 
    # if ($wronskian==Real(0));

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($a11 * $a22 == $a21 * $a12);


    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu1 = Formula($stu->D('x'));
    my $stu2 = Formula($stu->D('x','x'));
    return ($stu2 - $k2 * $stu1 + $kk * $stu) == Formula("e^($k x)");

    # my $stuDE = Formula($stuxx - $k2*$stux + $kk*$stu)
    # ->with(test_points=>[[1,1,1,0],[1,1,0,0],[2,1,1,0],[1,2,0,0]]);
    # return ($stuDE==Formula("e^($k x)"));

}));


COMMENT("Characteristic polynomial (r - a)^2.");

ENDDOCUMENT();

