## DESCRIPTION
## Higher order ODEs: constant coefficient second order nonhomogeneous linear ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Undetermined coefficients)
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(undet_05)
## Level(3)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('2.5')
## Problem1('4')
## KEYWORDS('differential equations','second order linear ODE','constant coefficients','nonhomogeneous')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserAssignment.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->add(
y=>"Real",
'a'=>'Parameter','b'=>'Parameter'
);

$k = random(2,5,1);
$k2 = 2 * $k;
$kk = $k * $k;

$y0 = random(2,9,1);
do { $yprime0 = random(1,5,1); } until ($yprime0 != $y0);

$num = 2*$k*$yprime0+1;
$den = 2*$k**2;

$diffeq = "y^{\,\prime\prime} + $kk y = \cos($k x) + \sin($k x)"; # tex

$yparticular = Compute("-1/$k2 x cos($k x) + 1/$k2 x sin($k x) + a cos($k x) + b sin($k x)");


Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

$ITEM Find a particular solution to the 
nonhomogeneous differential equation \( $diffeq \).
$BR
$BR
\( y_p = \) 
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;

$showPartialCorrectAnswers = 1;

ANS( $yparticular->cmp() );


##################################
#  reset the context and get rid of adaptive parameters

Context("Numeric");
Context()->variables->add(
'y' =>'Real',
'c1'=>'Real',
'c2'=>'Real',
);
Context()->variables->set(
'c1'=>{limits=>[2,4]},
'c2'=>{limits=>[2,4]}
);


$yhomogeneous = Compute("c1 cos($k x) + c2 sin($k x)");
$yinitvalprob = Compute("-1/$k2 x cos($k x) + 1/$k2 x sin($k x) + $y0 cos($k x) + $num/$den sin($k x)");


Context()->texStrings;
BEGIN_TEXT

$ITEMSEP
$ITEM Find the most general solution to the 
associated homogeneous differential equation.  
Use \( c_1 \) and \( c_2 \) in your answer to denote 
arbitrary constants.  Enter 
\( c_1 \) as ${BTT}c1${ETT} and
\( c_2 \) as ${BTT}c2${ETT}.
$BR
$BR
\( y_h = \)
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}


$ITEMSEP
$ITEM Find the solution to the original 
nonhomogeneous differential equation satisfying
the initial conditions \( y(0) = $y0 \) and 
\( y^{\,\prime}(0) = $yprime0 \).
$BR
$BR
\( y = \)
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation



ANS( $yhomogeneous->cmp( 
  checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    my $stu   = Formula($student);

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('c1'))==Formula(0) || 
      Formula($stu->D('c2'))==Formula(0) 
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(1.24);

    my $a11 = $stu->eval('c1'=>1,'c2'=>0,x=>$x,y=>0);
    my $a12 = $stu->eval('c1'=>0,'c2'=>1,x=>$x,y=>0);
 
    my $a21 = $stu->D('x')->eval('c1'=>1,'c2'=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval('c1'=>0,'c2'=>1,x=>$x,y=>0);

    # my $wronskian = $a11*$a22 - $a21*$a12;
    # Value->Error("Your functions are not linearly independent or your answer is not complete") 
    # if ($wronskian==Real(0));

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($a11 * $a22 == $a21 * $a12);

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stu1 = Formula($stu->D('x'));
    my $stu2 = Formula($stu->D('x','x'));
    return ($stu2) == (-1 * $kk * $stu);


    # my $stuDE = Formula($stuxx + $kk*$stu)
    # ->with(test_points=>[[1,1,1,0],[1,1,0,0],[2,1,1,0],[1,2,0,0]]);
    # return ($stuDE==Formula(0));
    # return ($stuxx) == (-1 * $kk * $stu);

}));



ANS( $yinitvalprob->cmp() );


COMMENT("Characteristic polynomial r^2 + a^2.");

ENDDOCUMENT();
