## DESCRIPTION
## College Algebra
## ENDDESCRIPTION
 
 
## DBsubject()
## DBchapter()
## DBsection()
## Date(6/21/2017)
## Institution(Community College of Denver)
## Author(Michael Sulock)
## MO(1)
## KEYWORDS('college algebra')
 
 
###########################
#  Initialization
 
DOCUMENT();
 
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"contextLimitedPoint.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
);
 
TEXT(beginproblem());

 
###########################
#  Setup
Context("Numeric");
 
# Two random integer to grab letters
$n1=random(0,10,1);
$n2=random(0,10,1);

# Array of 11 letters to be grabbed
@letter=qw/a b c d e f g h i j k/;

# Grabbing 2 letters
$var1 = $letter[$n1] ;
$var2=$letter[$n2] ;

# Adding the two letters as variables 
Context()->variables->add($var1 => "Real");
Context()->variables->add($var2 => "Real");

# The two sides of the correct answer
$ans1=Formula("$var1");
$ans2=Formula("log($var2)");
 

# Allows for the correct equation to be given in either order
$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2) = @{$student};
      my ( $f1, $f2) = @{$correct};
      if (   ( $f1 == $stu1 && $f2 == $stu2)     ||
             ( $f1 == $stu2 && $f2 == $stu1)      )
{return [1,1];} 
 
else { return [0,0];  }
 
} );
 
 
###########################
#  Main text
 
BEGIN_PGML


Rewrite the following equation in logarithmic form.

[`10^{[$var1]}=[$var2] `] 
 
 [____]{$multians} [`=`] [____]{$multians} [@ AnswerFormatHelp("logarithms") @]*
  
END_PGML
 
 

############################
#  Solution
 
#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION
 
COMMENT('MathObject version. Uses PGML.');
 
ENDDOCUMENT();