# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Polynomial and Rational Functions)
## DBsection(Polynomial Functions and their Graphs)
## Institution(PCC)
## Author(Alex Jordan)
## MO(1)



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PCCgraphMacros.pl",
  "PGcourse.pl",

);
TEXT(beginproblem());



##############################################


Context("Numeric");

$r[0] = random(-4,-1,1);
for my $i (1..2)
{
  $r[$i] = $r[$i-1] + random(1,5,1);
}

$p = Compute("(x-$r[0])*(x-$r[1])*(x-$r[2])")->reduce;

for my $i ($r[0]..$r[2]-1)
{$vals[$i-$r[0]] = abs($p->eval(x=>$i+0.5));}

$ymax = 1.3*max(@vals);

$xmax = max($r[2]+1,5);
$gr = init_graph(-5,-$ymax,$xmax,$ymax,axes=>[0,0],grid=>[$xmax+5,10],size=>[400,400]);
$gr -> lb('reset');


foreach my $i (1..$xmax) {
  $gr -> lb(new Label ( $i,-0.05*$ymax, $i,'black','center','middle'));
  $gr -> lb(new Label (-$i,-0.05*$ymax,-$i,'black','center','middle'));
}

add_functions($gr, "$p for x in <-5,$xmax> using color:blue and weight:2");

##############################################


BEGIN_PGML

Find a formula for a degree three polynomial function [`f`] that might have the graph below. There are many possible correct answers.

[@EnlargeImageStatementPGML()@]**

>>[@image(insertGraph($gr), width=>400, height=>400, tex_size=>700) @]* <<

    [`f(x)={}`][________________________________]


END_PGML

##############################################

$showPartialCorrectAnswers = 1;

ANS( $p->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"abs($c0) * $p");
      return $correct == $student;
    }
) );

##############################################

BEGIN_PGML_SOLUTION

We can see from the graph that [`[@$r[0]@]`], [`[@$r[1]@]`], and [`[@$r[2]@]`] are zeros. Therefore [`(x-[@$r[0]@])`], [`(x-[@$r[1]@])`], and [`(x-[@$r[2]@])`], are factors of [`f(x)`]. Since [`f`] is cubic, there can be no more factors involving [`x`]. However there may be an additional constant factor [`c`], and we can say that for some [`c\neq0`], [`f(x)=c[$p]`]. 

The graph indicates that as [`x\to\infty`], then [`f(x)\to\infty`]. In other words, the right arm of the graph points upward. This indicates that the leading coefficient of [`f(x)`] must be positive. So [`c`] must be positive. You can use any positive value for [`c`] to submit an answer. For instance, taking [`c=1`], we have that [`f(x)=[$p]`] is one possible answer to this question.



END_PGML_SOLUTION

ENDDOCUMENT();
