# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Polynomial and Rational Functions)
## DBsection(Polynomial Functions and their Graphs)
## Institution(PCC)
## Author(Alex Jordan)
## MO(1)



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
);
TEXT(beginproblem());



##############################################


Context("Numeric");

$r[0] = random(1,6,1);
$r[1] = $r[0]+random(1,6,1);


$p = Compute("(x-$r[0])*(x-$r[1])**2")->reduce;

##############################################

BEGIN_PGML

Find a formula for a degree three polynomial function [`f`] whose zeros are [`[@$r[0]@]`] with multiplicity [`1`], and [`[@$r[1]@]`] with multiplicity [`2`]. There are many possible correct answers to this question.

    [`f(x)={}`][____________________________________]

END_PGML


##############################################

$showPartialCorrectAnswers = 1;

ANS( $p->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $p");
      return $correct == $student;
    }
) );

##############################################

BEGIN_PGML_SOLUTION

Since [`[@$r[0]@]`] is a zero with multiplicity [`1`], [`(x-[@$r[0]@])`] is a factor. And since [`[@$r[1]@]`] is a zero with multiplicity [`2`], [`(x-[@$r[1]@])^2`] is a factor. The polynomial is degree three, so there can be no more factors involving [`x`]. There might be an additional constant factor, so we can say that for some [`c\neq 0`], [`f(x)=c(x-[@$r[0]@])(x-[@$r[1]@])^2`]. You could take [`c=1`] for instance, and [`f(x)=(x-[@$r[0]@])(x-[@$r[1]@])^2`] is a formula that meets the necessary conditions.




END_PGML_SOLUTION

ENDDOCUMENT();
