## MO(1)
# WeBWorK problem written by Alex Jordan an Debbie Neft
# Portland Community College
# ENDDESCRIPTION




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "niceTables.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl",
);



##############################################

TEXT(beginproblem());
Context("Numeric");

$a = non_zero_random(-8,8,1);
$b = random(1,10,1);
$c = random(1,8,1)+$b;

$f = Formula("$a (x-$b)(x-$c)^2");

$zeros = List("$b,$c,$c");

$crosses = "crosses";
$touches = "touches";

Context()->strings->add("crosses"=>{},"touches"=>{});

$multians = MultiAnswer(Compute("$b"), Compute("1"),Compute("$crosses"),Compute("$c"), Compute("2"),Compute("$touches"))->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $r1stud, $m1stud, $b1stud, $r2stud, $m2stud, $b2stud  ) = @{$student};
      my ( $r1cor, $m1cor, $b1cor, $r2cor, $m2cor, $b2cor  ) = @{$correct};
      if ( ($r1stud == $r1cor && $m1stud == $m1cor && $b1stud == $b1cor && $r2stud == $r2cor && $m2stud == $m2cor && $b2stud == $b2cor) ||
           ($r1stud == $r2cor && $m1stud == $m2cor && $b1stud == $b2cor && $r2stud == $r1cor && $m2stud == $m1cor && $b2stud == $b1cor) ) {
          return 1;}
       else {return 0;};
      
  }
);

if ($a > 0)
{$right = Compute("INF");
$left = Compute("-INF");
$rdir = 'an upward';
$ldir = 'a downward';}
else {$right = Compute("-INF");
$left = Compute("INF");
$ldir = 'an upward';
$rdir = 'a downward';
};




##############################################

BEGIN_PGML

Let [`f(x)=[$f]`].


Fill out the table below.  In the third column, type the word *crosses* for a zero where the curve crosses through to the other side of the [`x`]-axis, or *touches* for a zero where the curve approaches the [`x`]-axis, touches it, and reverses direction without crossing over it.

END_PGML
BEGIN_TEXT
$PAR

\{ 
DataTable([
[["zero",headerrow=>1],"multiplicity","crosses or touches the \(x\)-axis"], [$multians->ans_rule(5),$multians->ans_rule(5),$multians->ans_rule(15)], [$multians->ans_rule(5),$multians->ans_rule(5),$multians->ans_rule(15)],
],
allcellcss=> 'border: solid 1px;',
)
\}

$PAR
END_TEXT

BEGIN_PGML

What is the number of turning points of the graph of [`f`]? (A turning point is a point where the curve changes from increasing to decreasing, or from decreasing to increasing.)

     [________________________]

As [`x\to\infty,\quad f(x)\to{}`][______________]  

As [`x\to-\infty,\quad f(x)\to{}`][______________]  


END_PGML

##############################################

$showPartialCorrectAnswers = 1;


ANS( $multians->cmp() ); 
ANS( Compute("2")->cmp() );
ANS( $right->cmp() );
ANS( $left->cmp() );

##############################################

$lt = Compute("$a x^3")->reduce;

BEGIN_PGML_SOLUTION

The zero at [`[$b]`] has degree [`1`], and since this is odd, the curve crosses the [`x`]-axis there.

The zero at [`[$c]`] has degree [`2`], and since this is even, the curve touches the [`x`]-axis there and turns around.

We can now sketch a rough graph of [`f`], and see that there would have to be [`2`] turning points (one in between [`[$b]`] and [`[$c]`], and one _at_ [`[$c]`]).

Since the leading term is [`[$lt]`], the long-term behavior of the graph on the right is [$rdir] pointing arm. So as [`x\to\infty,\quad f(x)\to[$right]`].

And the long-term behavior of the graph on the left is [$ldir] pointing arm. So as [`x\to-\infty,\quad f(x)\to[$left]`].


END_PGML_SOLUTION


ENDDOCUMENT();