##DESCRIPTION
##  Factoring Trinomials
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Institution(The College of Idaho)
## Author(RA Cruz)
## MLT(perfectSq)
## Level(2)
## MO(1)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('5.6')
## Problem1('')
## KEYWORDS('factoring')
## Date: 2007/11 --Updated 2013/09 -rac

DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "CofIdaho_macros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

######################################
#  Setup

$a= random(2,6,1);
$b = random(-1,1,2);

$polynomial = Formula("$a^2 x^2 - (2 $a) x + $b")->reduce->TeX;

######################################
#  Main text

BEGIN_TEXT
Factor completely: 
$PAR
\( $polynomial = \)  \{ ans_rule(30) \} 
END_TEXT

######################################
#  Answer

$showPartialCorrectAnswers = 1;

$answer = "($a x - 1)^2";
ANS(FactoringEvaluator($answer,["x"])); 

#Context("PolynomialFactors-Strict");
#Context()->flags->set(singleFactors=>1);
#Context()->strings->add('Does not factor'=>{});
#
#if ($b>0)
#   {
#   LimitedPowers::OnlyIntegers(
#      minPower => 0, maxPower => 1
#   );
   $factoredform = Compute("($a x - 1)^2");
#   ANS($factoredform->cmp); 
#   }
#else
#   {
#   LimitedPowers::OnlyIntegers(
#      minPower => 0, maxPower => 2
#   );
#   $factoredform  = Compute("$a^2 x^2 - 2*$a x + $b");
#   ANS( $factoredform->cmp( checker=>sub {
#      my ( $correct, $student, $ansHash ) = @_;
#      if ($b<=0) {$student==$correct || $student == "Does not factor";}
#      return $correct == $student;
#   } ) );
#   ANS(FactoringEvaluator("Does not factor",$answer,"x")); 
#   }

######################################
#  Solution

$tex = $factoredform->TeX;
$solution = "$polynomial = $tex";
if ($b<=0) {$solution = "\mbox{The polynomial does not factor.}";}

SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

\($solution\)

END_SOLUTION

;
ENDDOCUMENT();
