##DESCRIPTION
##TYPE('equation answer')
## finding an equation of the line with a given slope and a
## point on the line;
##ENDDESCRIPTION

## Tagged by sawblade

## DBsubject(Algebra)
## DBchapter(Cartesian coordinate system)
## DBsection(Circles)
## Date(6/3/2002)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('2.2')
## Problem1('52')
## KEYWORDS('algebra', 'lines')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add(y=>"Real");

$A = random(-10,10,1);
$B = random(-10,10,1);
$C = random(1,20,1);

$f = Formula("(x-$A)**2 + (y-$B)**2 - $C**2")->reduce->cmp(
    checker => sub {
        my ($correct, $student, $ansHash) = @_;
        my $mrsq_cor = Formula($correct->eval(x=>$A, y=>$B));
        my $mrsq_stu = Formula($student->eval(x=>$A, y=>$B));
        if (Formula(0)==$mrsq_stu){return 0};
        if (($mrsq_stu*$correct) == ($mrsq_cor*$student)){return 1;}
        else {return 0;}
    }
);

BEGIN_PGML
Write an equation for the circle that is centered at [`([$A] , [$B] )`] and has radius [`[$C]`]. 

    [_]{$f}{20} = 0.
END_PGML

BEGIN_PGML_SOLUTION
If [`(x,y)`] is any point on the circle then the radius [`[$C]`] is the distance from [`(x,y)`] to the center     
    [``\sqrt{(x-[$A])^2+(y-[$B])^2}=[$C]``]    
Square both sides to obtain the textbook equation for the circle
    [``(x-[$A])^2+(y-[$B])^2 = [@ $C**2 @]``]    
Then subtract [`[@ $C**2 @]`] from both sides to obtain an equation with zero on the right hand side. 
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.


