##DESCRIPTION
##TYPE('equation answer')
## find an equation of a circle given its center and
## a point on the circle
##ENDDESCRIPTION

## Tagged by sawblade

## DBsubject(Algebra)
## DBchapter(Cartesian coordinate system)
## DBsection(Circles)
## Date(6/3/2002)
## MLT(Circle_from_center_and_point)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('2.2')
## Problem1('54')
## KEYWORDS('algebra', 'lines')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add(y=>"Real");

$A = random(-10,10,1);
$B = random(-10,10,1);
$C = random(-20,20,1);
$D = random(-20,20,1);

while ($A==$C && $B==$D) {$D = random(-20,20,1);}

# r^2
$rsq = ($A-$C)**2 + ($B-$D)**2;

# any nonzero multiple of this will be accepted

$ans = Compute("(x-$A)^2 + (y-$B)^2 - $rsq")->cmp(
    checker => sub {
        my ($correct, $student, $ansHash) = @_;
        my $mrsq_cor = Formula($correct->eval(x=>$A, y=>$B));
        my $mrsq_stu = Formula($student->eval(x=>$A, y=>$B));
        if (Formula(0)==$mrsq_stu){return 0};
        if (($mrsq_stu*$correct) == ($mrsq_cor*$student)){return 1;}
        else {return 0;}
    }
);

BEGIN_PGML
Write an equation for a circle, centered at [`([$A] , [$B] )`], that passes through [`([$C],[$D])`].

    [_]{$ans}{15} = 0.
END_PGML

BEGIN_PGML_SOLUTION
    
    [``(x-a)^2+(y-b)^2 = r^2 ``]    
is an equation for the circle with center [`(a,b)`] and radius [`r`].  Subtract [`r^2`] from both sides to obtain an equivalent equation with zero on the right hand side    
    [``(x-a)^2+(y-b)^2-r^2 = 0 ``]    
In our case the circle passes through through [`([$C],[$D])`] so the radius is the distance from  [`([$C],[$D])`] to the center [`([$A] , [$B] )`]    
    [``r=\sqrt{([$A]-[$C])^2+([$B]-[$D])^2}=\sqrt{[$rsq]}``].    
Thus    
    [``(x-[$A])^2 + (y-[$B])^2 - [$rsq] = 0``]    
is an equation for this circle.
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.


