## DESCRIPTION
## Precalculus: answers up to multiplication
## ENDDESCRIPTION

## DBsubject(WeBWorK)
## DBchapter(WeBWorK tutorial)
## DBsection(PGML tutorial 2015)
## Date(06/01/2015)
## Institution(Hope College)
## Author(Paul Pearson)
## MO(1)
## KEYWORDS('precalculus', 'answer up to multiplication')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGML.pl",
  "PGcourse.pl"
);

TEXT(beginproblem()); 


###########################
#  Setup

Context("Numeric");

$aSolution = Compute("(x-2)(x+1)");


###########################
#  Main text

BEGIN_PGML
Find a quadratic equation in terms of the variable 
[` x `] with roots [` -1 `] and [` 2 `].

[` y = `] [______________] [@ AnswerFormatHelp("formulas") @]*
END_PGML

############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $aSolution->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $aSolution");
      return $correct == $student;
    }
) );


############################
#  Solution

BEGIN_PGML_SOLUTION
Solution explanation goes here.
END_PGML_SOLUTION

COMMENT('MathObject version. Uses PGML.');

ENDDOCUMENT();
