# DESCRIPTION
# Determine linear dependence of sets and the dimension of a span.
# ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear independence)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(3)
## KEYWORDS('subspace','span')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
do {
  $a[1][1] = random(-19,19);
  $a[2][1] = random(-19,19);
  $a[3][1] = random(-19,19);
  $a[1][2] = random(-19,19);
  $a[2][2] = random(-19,19);
  $a[3][2] = random(-19,19);
  $a[1][3] = random(-19,19);
  $a[2][3] = random(-19,19);
  $a[3][3] = random(-19,19);
  $A = Matrix([$a[1][1], $a[1][2], $a[1][3]], [$a[2][1], $a[2][2], $a[2][3]], [$a[3][1], $a[3][2], $a[3][3]]);
  $det = $A->det;
} until ($det != 0);

$c1 = random(0,2,1);
do {$c2 = random(0,2,1)} until ($c2 != $c1);
do {$c3 = random(0,2,1)} until ($c3 != $c1 && $c3 != $c2);

@v = (0,0,0);
@vstr = ("v_1", "v_2", "v_3");
for ($i = 0; $i < 3; $i++) {
  if ($i == $c1) {
    $v[$i] = $A->column($c1+1);
  }
  if ($i == $c2) { 
    $v[$i] = $A->column($c2+1);
  }
}

do {
  $a = non_zero_random(-8,8,1);
  $b = non_zero_random(-8,8,1);
  $v[$c3] = $a*$v[$c1] + $b*$v[$c2];
  $tmp3 = max(abs($v[$c3]->element(1,1)),abs($v[$c3]->element(2,1)),abs($v[$c3]->element(3,1)));
  $tmp1 = max(abs($v[$c1]->element(1,1)),abs($v[$c1]->element(2,1)),abs($v[$c1]->element(3,1)));
  $tmp2 = max(abs($v[$c2]->element(1,1)),abs($v[$c2]->element(2,1)),abs($v[$c2]->element(3,1)));
} until ($tmp3 < 2*max($tmp1,$tmp2));
$ans[$c1] = $a;
$ans[$c2] = $b;
$ans[$c3] = -1;
$d1 = random(0,2,1);
do {$d2 = random(0,2,1)} until ($d2 != $d1);
$c1s = min($d1,$d2);
$c2s = max($d1,$d2);

$mans = $multians = MultiAnswer($ans[0],$ans[1],$ans[2])->with(
  singleResult => 0,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $s1, $s2, $s3 ) = @{$student};
    my ( $a1, $a2, $a3 ) = @{$correct};
    $tmp = $s1/$a1;
    if ($s2/$a2 == $tmp && $s3/$a3 == $tmp && $s1 != 0 && $s2 != 0 && $s3 != 0) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Let \(v_1\), \(v_2\), \(v_3\) be the vectors in \(\mathbb{R}^3\) defined by
$PAR
\[
v_1 = $v[0] \quad v_2 = $v[1] \quad v_3 = $v[2]
\]
$PAR
(a) Is \($LBRACE v_1, v_2, v_3 $RBRACE\) linearly independent? Write all zeros if it is or if it is linearly dependent write zero as a non-trivial (not all zero coefficients) linear combination of \(v_1\), \(v_2\), and \(v_3\)
$PAR
$BCENTER
\(0 =\) \{$mans->ans_rule(5)\} \(v_1+\) \{$mans->ans_rule(5)\} \(v_2+\) \{$mans->ans_rule(5)\} \(v_3\)
$ECENTER
$PAR
(b) Is \($LBRACE $vstr[$c1s], $vstr[$c2s] $RBRACE\) linearly independent? Write all zeros if it is or if it is linearly dependent write zero as a non-trivial (not all zero coefficients) linear combination of \($vstr[$c1s]\) and \($vstr[$c2s]\).
$PAR
$BCENTER
\(0 = \) \{ans_rule(5)\} \($vstr[$c1s]+\) \{ans_rule(5)\} \($vstr[$c2s]\)
$ECENTER
$PAR
(c) Type the dimension of \(\mbox{span}$LBRACE v_1, v_2, v_3$RBRACE\): \{ans_rule(5)\}
$PAR
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($mans->cmp());
ANS(Compute("0")->cmp());
ANS(Compute("0")->cmp());
ANS(Compute("2")->cmp());


###########################################################################
# solution
###########################################################################
Context()->texStrings;
BEGIN_SOLUTION
$PAR
${BBOLD}Solution:${EBOLD} $PAR
(a) Notice that $SPACE $SPACE\( 0 = $ans[0] v_1 + $ans[1] v_2 + $ans[2] v_3\)$SPACE $SPACE so that the set is linearly dependent.
$PAR
(b) Notice that \($vstr[$c2s]\) is not a constant multiple of \($vstr[c1s]\), so they are linearly independent.
$PAR
(c) The given span in this part must have dimension 2 since the set of all three vectors is linearly dependent, but \($LBRACE $vstr[$c1s], $vstr[$c2s] $RBRACE\) is linearly independent.
$PAR
END_SOLUTION
Context()->normalStrings;

# Edited and updated in 2012/2013 by Rick Lynch @ Mizzou.
# Desc: Determine linear dependence of sets and the dimension of a span.
ENDDOCUMENT();
