# DESCRIPTION
# Determine when a vector is in a span and the dimension of a subspace.
# ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear independence)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(3)
## KEYWORDS('subspace','span')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
do {
  $a[1][1] = random(-19,19);
  $a[2][1] = random(-19,19);
  $a[3][1] = random(-19,19);
  $a[4][1] = random(-19,19);
  $a[1][2] = random(-19,19);
  $a[2][2] = random(-19,19);
  $a[3][2] = random(-19,19);
  $a[4][2] = random(-19,19);
  $a[1][3] = random(-19,19);
  $a[2][3] = random(-19,19);
  $a[3][3] = random(-19,19);
  $a[4][3] = random(-19,19);
  $a[1][4] = random(-19,19);
  $a[2][4] = random(-19,19);
  $a[3][4] = random(-19,19);
  $a[4][4] = random(-19,19);
  $A = Matrix([$a[1][1], $a[1][2], $a[1][3], $a[1][4]], [$a[2][1], $a[2][2], $a[2][3], $a[2][4]], [$a[3][1], $a[3][2], $a[3][3], $a[3][4]],[$a[4][1], $a[4][2], $a[4][3], $a[4][4]]);
  $det = $A->det;
} until ($det != 0);

$c1 = random(0,3,1);
do {$c2 = random(0,3,1)} until ($c2 != $c1);
do {$c3 = random(0,3,1)} until ($c3 != $c1 && $c3 != $c2);
do {$c4 = random(0,3,1)} until ($c4 != $c1 && $c4 != $c2 && $c4 != $c3);

@vstr = ("v_1", "v_2", "v_3");
$v1 = $A->column($c1+1);
$v2 = $A->column($c2+1);
$v3 = $A->column($c3+1);
$u = $A->column($c4+1);

do {
  $a = non_zero_random(-8,8,1);
  $b = non_zero_random(-8,8,1);
  $c = non_zero_random(-8,8,1);
  $w = $a*$v1 + $b*$v2 + $c*$v3;
  $tmp4 = max(abs($w->element(1,1)), abs($w->element(2,1)), abs($w->element(3,1)), abs($w->element(4,1)));
  $tmp1 = max(abs($v1->element(1,1)), abs($v1->element(2,1)), abs($v1->element(3,1)), abs($v1->element(4,1)));
  $tmp2 = max(abs($v2->element(1,1)), abs($v2->element(2,1)), abs($v2->element(3,1)), abs($v2->element(4,1)));
  $tmp3 = max(abs($v3->element(1,1)), abs($v3->element(2,1)), abs($v3->element(3,1)), abs($v3->element(4,1)));
} until ($tmp4 < 2*max($tmp1,$tmp2,$tmp3));

$mans = $multians = MultiAnswer(-1,$a,$b,$c)->with(
  singleResult => 0,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $s1, $s2, $s3, $s4 ) = @{$student};
    my ( $a1, $a2, $a3, $a4 ) = @{$correct};
    $tmp = $s1/$a1;
    if ($s2/$a2 == $tmp && $s3/$a3 == $tmp && $s4/$a4 == $tmp && $s1 != 0 && $s2 != 0 && $s3 != 0 && $s4 != 0) {
      return [1,1,1,1];
    } else {
      return [0,0,0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Let \(u\), \(w\), \(v_1\), \(v_2\), and \(v_3\) be the vectors in \(\mathbb{R}^4\) defined by
$PAR
\[
u = $u \quad w = $w \quad v_1 = $v1 \quad v_2 = $v2 \quad v_3 = $v3
\]
$PAR
(a) Is \(u \in \mbox{span}$LBRACE v_1, v_2, v_3 $RBRACE\)? Write all zeros if it is not in the span or write zero as a non-trivial (not all zero coefficients) linear combination of \(u\), \(v_1\), \(v_2\), and \(v_3\) if \(u\) is in the span.
$PAR
$BCENTER
\(0 =\) \{ans_rule(5)\} \(u +\) \{ans_rule(5)\} \(v_1+\) \{ans_rule(5)\} \(v_2+\) \{ans_rule(5)\} \(v_3\)
$ECENTER
$PAR
(b)  Is \(w \in \mbox{span}$LBRACE v_1, v_2, v_3 $RBRACE\)? Write all zeros if it is not or if it is in the span write zero as a non-trivial (not all zero coefficients) linear combination of \(w\), \(v_1\), \(v_2\), and \(v_3\) if \(w\) is in the span.
$PAR
$BCENTER
\(0 =\) \{$mans->ans_rule(5)\} \(w + \) \{$mans->ans_rule(5)\} \(v_1+\) \{$mans->ans_rule(5)\} \(v_2+\) \{$mans->ans_rule(5)\} \(v_3\)
$ECENTER
$PAR
(c) Type the dimension of \(\mbox{span}$LBRACE v_1, v_2, v_3, u$RBRACE\): \{ans_rule(5)\}
$PAR
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS(Compute("0")->cmp());
ANS(Compute("0")->cmp());
ANS(Compute("0")->cmp());
ANS(Compute("0")->cmp());
ANS($mans->cmp());
ANS(Compute("4")->cmp());


###########################################################################
# solution
###########################################################################
Context()->texStrings;
BEGIN_SOLUTION
$PAR
${BBOLD}Solution:${EBOLD} $PAR
(a) One checks that \(Ax = u\) has no solution, where \(A\) is the matrix formed by taking  \(v_1,v_2, v_3\) as it columns. This means that \(u\) is not in the span.
$PAR
(b) Notice that $SPACE $SPACE \( 0 = -1w + $a v_1 + $b v_2 + $c v_3\) $SPACE $SPACE so that \(w\) is in the span. One way to find such a linear combination (where the coefficient of \(w\) is -1) is to simply find a solution of the linear system whose augmented matrix is (\(v_1,v_2,v_3|w\)). Finding such a linear combination simply means that \(w\) is in the span of \(v_1,v_2,v_3\).
$PAR
(c) One can check that \(\det(v_1,v_2,v_3,u) \neq 0\) so that the set of all four vectors is linearly independent and so has dimension 4. That is, \($LBRACE v_1, v_2, v_3, u$RBRACE\) is a basis for \(\mathbb{R}^4\).
$PAR
END_SOLUTION
Context()->normalStrings;

# Edited and updated in 2012/2013 by Rick Lynch @ Mizzou.
# Desc: Determine when a vector is in a span and the dimension of a subspace.
ENDDOCUMENT();
