##DESCRIPTION
##  College Algebra Review: Substitution
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Date(Spring 2012)
## Institution(University of Missouri)
## Level(3)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.4')
## Problem1('')
## KEYWORDS('algebra', 'factoring', 'quadratic')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPolynomial.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#######################################
#  Set-up

Context("LimitedPolynomial-Strict");
Context()->flags->set(singlePowers=>1);
Context()->variables->are(x=>"Real",t=>"Real");

$m = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6,7,8,9,10);
do {$n = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6,7,8,9,10)} until ($n != $m && $n != -$m);
do {$d = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,-1,1,2,3,4,5,6,7,8,9,10)} until ($d + $m != 0 && $d + $n != 0);
$a = 1; #do {$a = random(2,16)} until (gcd($a, $d + $m) == 1 && gcd($a,$d + $n) == 1);
$e2 = $m + $n;
$e3 = $m*$n;
$u = $d + $m;
$v = $d + $n;


$fac1 = Formula(" $a x + $u ")->reduce;
$fac2 = Formula(" $a x+ $v ")->reduce;
$texanswer = "\( \left($a x + $u \right)\left($a x + $v\right) \)";

if ($e2 == 1) {
  $expression = "\(  ( x + $d)^2 + (x + $d)  + $e3 \)";
} elsif ($e2 == -1) {
  $expression = "\(  ( x + $d)^2 - (x + $d)  + $e3 \)";
} else {
  $expression = "\(  ( x + $d)^2 + $e2 (x + $d)  + $e3 \)";
}

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your 
answer as much as possible:
$PAR
${BCENTER}
$expression \(= \) 
\( \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \)
${ECENTER}
END_TEXT

################################################
#  Answer evaluation

$texanswer = "\(\left($fac1\right)\left($fac2\right)\)";
Context()->normalStrings;

################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

ENDDOCUMENT();
