## DESCRIPTION
## Discrete Mathematics
## ENDDESCRIPTION

## Tagged by cmd6a 8/6/06

## DBsubject(Set theory and logic)
## DBchapter(Propositional logic)
## DBsection(Rules of inference)
## Institution(SUNYSB)
## Level(2)
## MO(1)
## KEYWORDS('discrete mathematics','logic','proof by contradiction')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "niceTables.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->strings->add(a=>{}, b=>{}, c=>{}, d=>{}, e=>{}, f=>{}, 
  g=>{}, h=>{}, i=>{}, j=>{}, k=>{}, l=>{}, m=>{}, n=>{}, o=>{}, 
  p=>{}, q=>{}, r=>{}, s=>{} );

TEXT(beginproblem()); 
$showPartialCorrectAnswers = 0;

BEGIN_TEXT
    For the following $BBOLD proof by contradiction $EBOLD provide the justifications at each step,
using the following equivalences and inference rules.  Use the following keys:

\{ 
  DataTable( 
  [
         [ 'a', 'Idempotent Law'],
         [ 'b', 'Double Negation'],
         [ 'c', "De Morgan's Law"], 
	 [ 'd', 'Commutative Properties '],
         [ 'e', 'Associative Properties '],
         [ 'f', 'Distributive Properties '],
         [ 'g', 'Equivalence of Contrapositive '],
         [ 'h', 'Definition of Implication '],
         [ 'i', 'Definition of Equivalence '],
         [ 'j', 'Identity Laws \((p \vee F = p \wedge T = p) \) '],
         [ 'k', 'Tautology \((p \vee \neg p = T) \) '], 
         [ 'l', 'Contradiction \((p \wedge \neg p = F) \) '], 
         [ 'm', 'Negation of the goal to prove '],
         [ 'n', 'Modus Ponens'],
         [ 'o', 'Modus Tollens'],
         [ 'p', 'Transitivity of Implication'],
         [ 'q', 'Conjunctive Simplification'],
         [ 'r', 'Conjunctive Addition'],
         [ 's', 'Disjunctive Addition'],
  ],
  align => 'c|c',
  midrules => 1,
  center =>0,
); 
\}
END_TEXT

# three versions of the problem, using different letters, otherwise identical:
$version = random(1,3,1);  

if ($version == 1)
{
Context()->texStrings;
BEGIN_TEXT
    $PAR
    We want to prove \(s\) by a proof by contradiction from the following propositions.
$PAR
\{ 
  DataTable(
    [
      ['\( p \rightarrow q \)' ],
      ['\( r \rightarrow q \)' ],
      ['\( \neg q \)' ], 
      ['\( \neg(s \wedge T) \rightarrow p \)' ]
    ],
    align => '|c|',
    midrules => 1,
    center => 0,
  )
\}

$PAR
    \( \neg s \) by \{ ans_rule(1) \}

    $PAR
    \( \neg p \) by  \{ ans_rule(1) \} between \( p \rightarrow q \) and \( \neg q \)

    $PAR
    \( s \wedge T \) by  \{ ans_rule(1) \} between \( \neg(s \wedge T) \rightarrow p \) and \( \neg p \) previously deduced.

    $PAR
    \( s \) by  \{ ans_rule(1) \} from \( s \wedge T \)

    $PAR
    We have \( s \) and \( \neg s \) true, therefore we have a contradiction. 
END_TEXT
Context()->normalStrings;
}
## end of version 1

elsif ($version == 2)
{
Context()->texStrings;
BEGIN_TEXT
    $PAR
    We want to prove \(d\) by a proof by contradiction from the following propositions.

\{ 
  DataTable(
    [
      ['\( a \rightarrow b \)' ],
      ['\( r \rightarrow b \)' ],
      ['\( \neg b \)' ], 
      ['\( \neg(d \wedge T) \rightarrow a \)' ]
    ],
    align => '|c|',
    midrules => 1,
    center => 0,
  )
\}

    $PAR
    \( \neg d \) by \{ ans_rule(1) \}

    $PAR
    \( \neg a \) by  \{ ans_rule(1) \} between \( a \rightarrow b \) and \( \neg b \)

    $PAR
    \( d \wedge T \) by  \{ ans_rule(1) \} between \( \neg(d \wedge T) \rightarrow a \) and \( \neg a \) previously deduced.

    $PAR
    \( d \) by  \{ ans_rule(1) \} from \( d \wedge T \)

    $PAR
    We have \( d \) and \( \neg d \) true, therefore we have a contradiction. 
END_TEXT
Context()->normalStrings;
}

# end of version 2

# start version 3
else
{
Context()->texStrings;
BEGIN_TEXT
    $PAR
    We want to prove \(s\) by a proof by contradiction from the following propositions.

$PAR
\{ 
  DataTable(
    [
      ['\( p \rightarrow b \)' ],
      ['\( r \rightarrow b \)' ],
      ['\( \neg b \)' ], 
      ['\( \neg(s \wedge T) \rightarrow p \)' ]
    ],
    align => '|c|',
    midrules => 1,
    center => 0,
  )
\}

    $PAR
    \( \neg s \) by \{ ans_rule(1) \}

    $PAR
    \( \neg p \) by  \{ ans_rule(1) \} between \( p \rightarrow b \) and \( \neg b \)

    $PAR
    \( s \wedge T \) by  \{ ans_rule(1) \} between \( \neg(s \wedge T) \rightarrow p \) and \( \neg p \) previously deduced.

    $PAR
    \( s \) by  \{ ans_rule(1) \} from \( s \wedge T \)

    $PAR
    We have \( s \) and \( \neg s \) true, therefore we have a contradiction. 
END_TEXT
Context()->normalStrings;
}
## end of version 3

# answers are the same for all three versions 

ANS( String("m")->cmp());
ANS( String("o")->cmp());
ANS( String("o")->cmp());
ANS( String("q")->cmp(
  checker => sub{ 
    my ($correct, $student, $anshash)=@_;
    if ($correct == $student or String("j")==$student){
      return 1;
    } else {
      return 0;
    }
  }
));

ENDDOCUMENT();

