# DESCRIPTION
# Simplify exponents in an expression.
# ENDDESCRIPTION

## DBsubject('Algebra') 
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponents')
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "MUHelp.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real", a=>"Real", b=>"Real", A=>"Real", B=>"Real");
LimitedPowers::OnlyPositiveIntegers();
Context()->operators->undefine("/");

do {
  $c1 = random(2,15);
  $c2 = random(2,15);
} until (gcd($c1,$c2) == 1 && (($c1 < 10 && $c2 => 10) || ($c1 > 10 && $c1 <= 10)));

$p1 = random(4,7);
$q1 = random(8,13);
do {$p2 = random(2,6)} until ($p2 < $p1);
do {$q2 = random(2,7)} until ($q2 != $p2 && $q2 != $p1);

$nc1 = ($c1)**2;
$nc2 = ($c2)**2;
$e1 = 2*($p1-$p2);
$e2 = 2*($q1-$q2);

$num = Formula("$nc1 a^{$e1} b^{$e2}")->reduce;
$den = Formula("$nc2")->reduce;
$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1x^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01x^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);

Context()->texStrings;
$num = $num;  $den = $den;
$texanswer = "\(\displaystyle \frac{$num}{$den}\)";

$frac = "\left( \frac{$c1 a^{$p1} b^{$q1}}{$c2 a^{$p2} b^{$q2}} \right)^2";
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write the following using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
Context()->texStrings;
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} $texanswer
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $BR
  ${BBOLD}Help:${EBOLD} \{ MUHelp("algfrac") \}
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization. Edited and updated in 2012/2013.
$BR
Desc: Simplify exponents in an expression.');
ENDDOCUMENT();
