# DESCRIPTION
# Simplify exponents in an expression.
# ENDDESCRIPTION

## DBsubject('Algebra') 
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponents')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "MUHelp.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real", y=>"Real", "X"=>"Real", Y=>"Real");
LimitedPowers::OnlyPositiveIntegers();
Context()->operators->undefine("+","-","/");

$n = random(2,4,1);
$a = random(1,5,1);
$b = ($a == 1) ? random(2,5,1) : 1;
$c = random(2,3,1);
$d = ($c == 2) ? 3 : 2;
$do {$e = random(2,4,1)} until (2 < $e*$n);

$tmp1 = ($a**2)*($d**$e);
$tmp2 = ($b**2)*($c**$e);
$tmpn = $tmp1/gcd($tmp1,$tmp2);
$tmpd = $tmp2/gcd($tmp1,$tmp2);
$e1 = $n*$e - 2;
$num = Formula("$tmpn")->reduce;
$den = Formula(" $tmpd*(x^($e1))*(y^2)")->reduce;
$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1x^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01x^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);
$texanswer = ($tmpd == 1) ? "\(\displaystyle \frac{$tmpn}{x^{$e1} y^2}\)" : "\(\displaystyle \frac{$tmpn}{$tmpd x^{$e1} y^2}\)";

$frac = ($a == 1) ? "\left(\dfrac{x}{$b y}" : "\left(\dfrac{$a x}{y}";
$frac = $frac."\right)^{2} \left(\frac{$c x^{$n}}{$d}\right)^{-$e}";
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write the following using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
Context()->texStrings;
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} $texanswer
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $BR
  ${BBOLD}Help:${EBOLD} \{ MUHelp("algfrac") \}
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization. Edited and updated in 2012/2013.
$BR
Desc: Simplify exponents in an expression.');
ENDDOCUMENT();
