##DESCRIPTION
## Simplify Exponents
## 
##ENDDESCRIPTION
## DBsubject('Intermediate Algebra')
## DBchapter('Ch 01: Algebra, Mathematical Models and Problem Solving')
## DBsection('Properties of Integral Exponents')
## KEYWORDS('real numbers','algebraic expressions')
## TitleText1('Essentials of Intermediate Algebra')
## EditionText1('1')
## AuthorText1('Blitzer')
## Section1('1.6')
## Problem1('')
## Author('RA Cruz')
## Institution('The College of Idaho')
## Date: 2007/09
# Pulled from NPL: Library/CollegeOfIdaho/setAlgebra_01_06_Exponents/16IntAlg_11_Exponents.pg
# Edited for Periodic Re-randomization by Rick Lynch @ Mizzou 
# Significant changes made also by Rick Lynch

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPowers.pl",
  "answerHints.pl",
  "PGunion.pl",
  "parserMultiAnswer.pl",
  "contextLimitedNumeric.pl",
  "problemRandomize.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
######################################
#  Setup
Context("Numeric");
@alphabet = ("a","b","x","y");
$choose = random(0,2,2);
$var1 = $alphabet[ $choose ];
$var2 = $alphabet[ $choose+1 ];
Context()->variables->are($var1=>'Real',
                          $var2=>'Real');
$n1 = random(7,20,1);
$m1 = random(7,20,1);
$m2 = $m1 - random(2,5,1);
$num2 = random(2,7,1);
$num1 = $num2 * random(2,5,1); #To keep whole numbers in the answre
$frac = "\frac{$num1 $var1^{$n1}$var2^{$m1}}{$num2 $var1 $var2^{$m2}}";

$num = Formula("($num1/$num2)*$var1**{$n1-1}*$var2**{$m1-$m2}")->reduce;
$den = Formula("1");
$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);
Context()->operators->undefine("+","-","/");

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


######################################
#  Main text
Context()->operators->undefine("+","-");
Context()->texStrings;
BEGIN_TEXT

Simplify and write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER


END_TEXT

######################################
#  Answers

Context()->operators->set(@LimitedPowers::OnlyPositiveIntegers);

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );

Context()->texStrings;
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \($num\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
$PAR
  ${BBOLD}Note: $EBOLD   If your answer is not a fraction, write "1" in the denominator.
$BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
