##DESCRIPTION
##   Evaluation of a simple algebraic expression
##ENDDESCRIPTION

##KEYWORDS('algebra','exponent')
## hcao tagged and PAID on 12-12-2003

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## Date('6/3/2002')
## Author('')
## Institution('')
## TitleText1('College Algebra')
## EditionText1('3')
## AuthorText1('Stewart, Redlin, Watson')
## Section1('1.3')
## Problem1('27')

# Pulled from NPL: Library/Rochester/setAlgebra02ExponentsRadicals/sw1_3_27.pg
# Edited for Periodic Re-randomization by Rick Lynch @ Mizzou 
# Significant changes made also by Rick Lynch

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "problemRandomize.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
Context("Numeric")->variables->add(y=>"Real");

do {
$n1 = random(3,6,1);
$n2 = random(2,6,1);
$n3 = random(3,6,1);
$n4 = random(2,6,1);
$n5 = random(-1,1,2)*random(2,3,1);
$n6 = ($n5 < 0) ? random(2,3,1) : random(-3,-2,1);
} until ($n1*$n5 != -$n3*$n6 && $n2*$n5 != -$n4*$n6);
$a = random(2,5,1);
do {$b = random(2,5,1)} until (gcd($a,$b)==1);

$frac = "($a x^{$n1} y^{$n2})^{$n5}($b x^{$n3} y^{$n4})^{$n6}";
$numconst = ($n5 < 0) ? $b**$n6 : $a**$n5;
$denconst = ($n5 < 0) ? $a**(-$n5) : $b**(-$n6);
$xexp = $n1*$n5 + $n3*$n6;
$yexp = $n2*$n5 + $n4*$n6;

if ($xexp < 0 && $yexp < 0) {
  $num = Formula("$numconst")->reduce;
  $den = Formula("$denconst x^(-$xexp) y^(-$yexp)")->reduce;
} elsif ($xexp > 0 && $yexp <0) {
  $num = Formula("$numconst x^$xexp")->reduce;
  $den = Formula("$denconst y^(-$yexp)")->reduce;
} elsif ($xexp < 0 && $yexp > 0) {
  $num = Formula("$numconst y^$yexp")->reduce;
  $den = Formula("$denconst x^(-$xexp)")->reduce;
} else {
  $num = Formula("$numconst x^($xexp) y^($yexp)")->reduce;
  $den = Formula("$denconst ")->reduce;
}
LimitedPowers::OnlyPositiveIntegers();
$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify further");
          $self->setMessage(2,"Simplify further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);
Context()->operators->undefine("+","-","/");


if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER


END_TEXT

###################################
# Answers

$showPartialCorrectAnswers = 1;
Context()->operators->undefine("+","-");
ANS($multians->cmp() );

$dispans = "\frac{$num}{$den}";
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle $dispans\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
$PAR
  ${BBOLD}Note: $EBOLD  If your answer is not a fraction, write "1" in the denominator.
$BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization. From NPL. Edited and updated in 2013.');
ENDDOCUMENT();
