##DESCRIPTION
##  College Algebra Review: Computing Absolute Values
##ENDDESCRIPTION

##KEYWORDS('algebra', 'absolute value')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Real Numbers')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.2')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

###################################
# Setup

Context("Numeric");

do {
$b = random(2,6,1);
$a = random(3,7,1)*$b;
$c = random(1,7,1)*$b;
} until ($a/$c != 1);
$m1 = random(2,20,1);
do {$n1 = random(4,20,1)} until ($n1 != $m1);
do {$p1 = random(5,20,1)} until ($p1 != $m1 && $p1 != $n1);

Context("Numeric")->variables->are(x=>"Real", y=>"Real", z=>"Real");
@vars = ("x","y","z");
$choice = random(0,2,1);
$var = $vars[$choice];

LimitedPowers::OnlyPositiveIntegers();
$frac = "\frac{\left($a $var^{$m1}\right)\left($var^{$n1}\right)}{$c $var^{-$p1}}";

#Context(LimitedPowers::OnlyPositiveIntegers());
$numconst = $a/gcd($a,$c);
$denconst = $c/gcd($a,$c);
$numexp1 = ($m1+$n1+$p1);
$den = Formula("$denconst ")->reduce;
$num = Formula("$numconst $var^{$numexp1}")->reduce;
Context()->texStrings;
if ($denconst == 1) {
  $texanswer = "\( $num  } \)";
} else {
  $texanswer = "\(\displaystyle \frac{$num }{$denconst  } \)";
}

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);
Context()->operators->undefine("+","-","/");


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER


END_TEXT
Context()->normalStrings;

###################################
# Answers


Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";
$showPartialCorrectAnswers = 1;
Context()->operators->undefine("+","-");
ANS($multians->cmp() );

Context()->texStrings;
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} $texanswer
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
$PAR
  ${BBOLD}Note: $EBOLD  If your answer is not a fraction, write "1" in the denominator.
$BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
