## DESCRIPTION
## Simplify a rational expression
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Simplify rational expressions)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################

#Product rule for exponents, negative exponent and a positive power of 1 (student must enter in the denominator)

Context("RationalFunction");
Context()->variables->add(y =>'Real');

#common factors are: a, x^$nd, and y^$mn
$a = non_zero_random(4, 9);

$nn = non_zero_random(3, 12);
do{
$nd = random(2, 8);
}until($nd < $nn);

$n = $nn - $nd;

$md = non_zero_random(3, 12);
do{
$mn = random(2, 8);
}until($mn < $md);

$m = $md - $mn;

$b = list_random(2, 4, 8);
$c = list_random(3, 9);

#Question
$g = Formula("($a*$b*x^$nn*y^$mn)/($a*$c*x^$nd*y^$md)");

#answer
$f = Formula("($b*x^$n)/($c*y^$m)")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify. Use only positive exponents in your answer.

[``[$g] = ``]  [________]{$cmp}

 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
There is a common factor of [`[$a]x^[$nd]y^[$mn]`].  We can divide that out of the numerator and denomiantor to get:

[``\frac{[$a] \cdot [$b] \cdot x^[$nd] \cdot x^[$n] \cdot y^[$mn] }{[$a] \cdot [$c] \cdot x^[$nd] \cdot y^[$mn] \cdot y^[$m]}``]

[``[$f]``]

END_PGML_SOLUTION

ENDDOCUMENT();




