## DESCRIPTION
## Multiply Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Simplify rational expressions)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a','b','c','d','m','n','p','r','s','t','w','y');
Context()->variables->are($var =>'Real');

#Problem is of the form:
#(c*x^n)/((x + a)(x + b)) TIMES (x^2 - a^2)/(c*d*x^m)
#
$a = random(2,10);
do{
$b = non_zero_random(-8, 8, 1);
}until($b != $a and $b != -$a);

$c = random(2, 8);
$d = random(2, 8);

$n = random(2, 8);

do{
$m = random(1, 5);
}until($n > $m);

#Question
$g1 = Formula("($c*$var^$n)/($var^2 + ($a+$b)*$var + $a*$b)")->reduce;
$g2 = Formula("($var^2 - $a**2)/($c*$d*$var^$m)")->reduce;

$g1new = Formula("($c*$var^$n)/(($var + $a)($var + $b))")->reduce;
$g2new = Formula("(($var + $a)($var - $a))/($c*$d*$var^$m)")->reduce;
$gcf = Formula("$c*$var^$m*($var + $a)")->reduce;

#answer
$f = Formula("($var^($n - $m)*($var - $a))/($d*($var + $b))")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Multiply and simplify.


[``[$g1]\cdot[$g2] = ``]  [________]{$cmp}

 
END_PGML

##############################################

BEGIN_PGML_SOLUTION
Factor both expressions before multiplying.

The denominator of the first expression can be factored as: 

[`([$var]+[$a])([$var] + [$b])`]

The numerator of the second expression can be factored as:

[`([$var] + [$a])([$var] - [$a])`]

In factored form, the problem is now:

[`\displaystyle [$g1new] \cdot [$g2new]`]

[`[$gcf]`] is a common factor.  If we reduce this and multiply straight across, we are left with [`\displaystyle [$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();