## DESCRIPTION
## Multiply Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Simplify rational expressions)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
Context()->variables->are(a =>'Real', b => 'Real');

#constants
$a = random(2, 8);
$b = random(2, 8);
$c = random(2, 8);

#exponents
$n1 = random(2, 10);
$n2 = random(2, 10);
$m1 = random(4, 12);
$m2 = random(2, 12);
$m3 = random(1, 10);

do{
$m4 = random(2, 25);
}until($m4 > ($m1 + $m2 - $m3));

#Question
$g1 = Formula("($a*$b*a^$n1*b^$m1)/(b^$m3)");
$g2 = Formula("($c*a^$n2*b^$m2)/($b*$c*b^$m4)");

$g1new = Formula("($a*a^$n1*b^$m1)/(b^$m3)");
$g2new = Formula("(a^$n2*b^$m2)/(b^$m4)");

$mult = Formula("($a*a^($n1+$n2)*b^($m1+$m2))/(b^($m3+$m4))");

#answer
$f = Formula("($a*a^($n1+$n2))/(b^($m3+$m4-$m1-$m2))")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Multiply and simplify. Use only positive exponents in your answer.

[``[$g1]\cdot[$g2] = ``]  [________]{$cmp}

 
END_PGML

##############################################
BEGIN_PGML_SOLUTION
Try reducing before you multiply.

We have common factors of [$c] and [$b]:

[`\displaystyle [$g1new] \cdot [$g2new]`]

Multply straight across:

[`\displaystyle [$mult]`]

And reduce:

[`\displaystyle [$f]`]

END_PGML_SOLUTION



ENDDOCUMENT();