## DESCRIPTION
## Multiply Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a','b','c','d','m','n','p','r','s','t','w','y');
Context()->variables->are($var =>'Real');

#Problem is of the form:
#(x-a)^2/(x^2 - a^2) TIMES (x+a)^2/((x-a)(bx+c))
#
$a = random(2,10);
$b = random(2, 6);
$c = non_zero_random(-7, 7);

#Question
$g1 = Formula("($var^2 - 2*$a*$var + $a**2)/($var^2 - $a**2)")->reduce;
$g2 = Formula("($var^2 + 2*$a*$var + $a**2)/($b*$var^2 + ($c-$a*$b)*$var - $a*$c)")->reduce;

#FOR SOLUTION
$g1new = Formula("(($var-$a)($var-$a))/(($var+$a)($var -$a))")->reduce;
$g2new = Formula("(($var + $a)($var + $a))/(($var - $a)($b*$var + $c))")->reduce;

$first = $b;
$middle = $c - $a*$b;
$last = -$a*$c;
$prod = $first*$last;
$num1 = $c;
$num2 = -$a*$b;

$gcf = Formula("($var - $a)($var - $a)($var + $a)");

#answer
$f = Formula("($var + $a)/($b*$var + $c)");

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Multiply and simplify.

[``[$g1]\cdot[$g2] = ``]  [________]{$cmp}

 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
Factor both expressions before multiplying.

To factor the denominator of the second expression, find two numbers with a product of [`[$first]\cdot[$last] = [$prod]`] and a sum of [`[$middle]`].  Those numbers are [`[$num1]`] and [`[$num2]`].

[`[$first][$var]^2 + [$num1][$var] + [$num2][$var] + [$last]`]

[`[$var]([$first][$var] + [$c]) - [$a]([$first][$var] + [$c])`]

[`([$var] - [$a])([$first][$var] + [$c])`]

In factored form, the problem is now:

[`\displaystyle [$g1new] \cdot [$g2new]`]

[`[$gcf]`] is a common factor.  If we reduce this and multiply straight across, we are left with [`\displaystyle [$f]`].
END_PGML_SOLUTION
ENDDOCUMENT();