## DESCRIPTION
## Simplify a rational expression
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Simplify rational expressions)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################


Context("RationalFunction");
$var = list_random('y', 'z', 'a', 'b', 'c', 'p', 'r', 't', 'm');
Context()->variables->add($var =>'Real');

#problem is of the form:
#[(x^2 + a)(x + b)]/[(x + c)(x + b)]
#factor by grouping in the numerator
$a = non_zero_random(1, 9);
$b = non_zero_random(-9, 9);
$c = non_zero_random(-9, 9);

#Question
$g = Formula("($var^3 + $b*$var^2 + $a*$var + $a*$b)/($var^2 + ($c + $b)*$var + $b*$c)")->reduce;
#answer
$f = Formula("($var^2 + $a)/($var + $c)")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the rational expression.

[``[$g] = ``]  [________]{$cmp}
 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
We can only cancel common factors, so you must factor before you can simplify.

Factor the numerator by grouping.

[`[$var]^2([$var] + [$b]) + [$a]([$var] + [$b])`]

[`([$var] + [$b])([$var]^2 + [$a])`]

[`\frac{([$var]^2 + [$a])([$var] + [$b])}{([$var] + [$b])([$var] + [$c])}`]

[`([$var] + [$b])`] is a common factor.  If we reduce this we are left with [`[$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();