##DESCRIPTION
##  Vector fields
##ENDDESCRIPTION

##KEYWORDS('Vector Field')

## DBsubject('Calculus')
## DBchapter('Vector Calculus')
## DBsection('Vector Fields')
## Date('')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')

###################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT( beginproblem() );

###################################
#  Setup

Context("Vector2D");
Context()->flags->set(
reduceConstants => 0,
reduceConstantFunctions => 0,
);

$a = random(2,5,1);
$aneg = -1*$a;

$s = random(-1,1,2);

$b = $s*$a;
$bneg = $s*$aneg;

$F = Compute("<$bneg*y,$b*x>");

foreach my $j (0..2) {
  foreach my $i (0..2) {
    ($x,$y) = ($i-1, $j-1);
    $G[$i][$j] = $F->eval(x=>$x,y=>$y);
  }
}

$ma = MultiAnswer(
  $G[0][2], $G[1][2], $G[2][2], 
  $G[0][1], $G[1][1], $G[2][1], 
  $G[0][0], $G[1][0], $G[2][0], 
  )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\begin{array}{rrr} %s & %s & %s \\ %s & %s & %s \\ %s & %s & %s \end{array}",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..8) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$rw = 10; # answer rule width

##################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose \( \vec{F} = $F \).  Complete the following table of values of \( \vec{F} \).
$BR
$BR
$BCENTER
Values of \( \vec{F} \)
$PAR
\{
BeginTable().
  AlignedRow(["\(y=1\)",$ma->ans_rule($rw),$ma->ans_rule($rw),$ma->ans_rule($rw)],align=>'LEFT').
  TableSpace(5,0).
  AlignedRow(["\(y=0\)",$ma->ans_rule($rw),$ma->ans_rule($rw),$ma->ans_rule($rw)],align=>'LEFT').
  TableSpace(5,0).
  AlignedRow(["\(y=-1\)",$ma->ans_rule($rw),$ma->ans_rule($rw),$ma->ans_rule($rw)],align=>'LEFT').
  TableSpace(5,0).
  AlignedRow([" ","\(x=-1\)","\(x=0\)","\(x=1\)"],align=>'CENTER').
  TableSpace(5,0).
EndTable();
\}
$ECENTER
$PAR
Using your table of values as a starting point, sketch this vector field on a piece of paper for \( -2 \leq x \leq 2 \) and \( -2 \leq y \leq 2 \).
END_TEXT
Context()->normalStrings;


##################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $ma->cmp() );

COMMENT('MathObject version');

ENDDOCUMENT();

