## DESCRIPTION
## College Algebra
## ENDDESCRIPTION
 
 
## DBsubject()
## DBchapter()
## DBsection()
## Date(6/21/2017)
## Institution(Community College of Denver, Colorado Community College System)
## Author(Michael Sulock)
## MO(1)
## KEYWORDS('college algebra')
 
 
###########################
#  Initialization
 
DOCUMENT();
 
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"contextLimitedPoint.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
);
 
TEXT(beginproblem());

 
###########################
#  Setup
Context("Numeric");
 

$distance=non_zero_random(40,200,20);
$speed = random(10,30,2);


# Adding the two letters as variables 
Context()->variables->add(t => "Real");
Context()->variables->add(d => "Real");


# The two sides of the correct answer
$ans1=d;
$ans2=Formula("$distance- t*$speed");
 

# Allows for the correct equation to be given in either order
$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2) = @{$student};
      my ( $f1, $f2) = @{$correct};
      if (   ( $f1 == $stu1 && $f2 == $stu2)     ||
             ( $f1 == $stu2 && $f2 == $stu1)      )
{return [1,1];} 
 
else { return [0,0];  }
 
} );
 
 
###########################
#  Main text
 
BEGIN_PGML


A boat is [$distance] miles away from the marina, sailing directly toward it at [$speed] miles per hour. 

Write an equation for the distance of the boat from the marina, [`d`], after [`t`] hours. 
 
 [____]{$multians} [`=`] [____]{$multians} [@ AnswerFormatHelp("numbers") @]*
 

 
END_PGML
 
 
 
############################
#  Solution
 
#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION
 
COMMENT('MathObject version. Uses PGML.');
 
ENDDOCUMENT();