## DBsubject('Statistics')
## DBchapter('Confidence intervals')
## DBsection('One sample proportion')
## level(4)
## Date('10/09/2012')
## Author('Djun Kim')
## Institution('UBC')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());


##############################################################
#
#  Setup
#
#
Context("Numeric");
$num = random(20,25,1);
$n = random(120, 130, 1);
$phat = $num/$n;

$m = random(0.03, 0.05, 0.01);

$ans1 = Real(ceil(1.96**2 * ($phat) * (1 - $phat) / $m**2));
$ans2 = Real(ceil(2.00**2 * ($phat) * (1 - $phat) / $m**2));


##############################################################
#
#  Text
#
#

Context()->texStrings;

BEGIN_TEXT
Refer to the following scenario.
$PAR
A government official is in charge of allocating social programs throughout the city of Vancouver.  
He will decide where these social outreach programs should be located based on the percentage 
of residents living below the poverty line in each region of the city.  He takes a simple random 
sample of $n people living in Gastown and finds that $num have an annual income that is below 
the poverty line. 

$PAR
Suppose that the government official wants to re-estimate the population proportion and wishes
for his 95$PERCENT confidence interval to have a margin of error no larger than $m.  How large 
a sample should he take to achieve this? Please carry answers to at least six decimal places in intermediate steps.

$PAR
Sample size = \{ ans_rule(6) \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 0;

ANS( $ans2->cmp( 
   correct_ans => "$ans1 (using a critical value of 1.96)  or $ans2 (using a critical value of 2.0)",
   checker => sub {
       my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
       return (($student == $ans1 || $student == $ans2)? 1 : 0);      # return 1 if correct, 0 otherwise
   }
));



ENDDOCUMENT();
