# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Solve the following quadratic equation
#
#       a x^2 + bx + c = 0
#
# If you need to use the square root symbol, 
# as in x=sqrt{17}, type it like: *sqrt(17)*
#
# We choose a, b, c such that b^2-4ac<0 which means
# that there are no real solutions to these equations.
#
# Last updated: Hughes 8/4/13
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','square root','prime','subtract')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################
DOCUMENT();

loadMacros(
      "PGstandard.pl",
      "MathObjects.pl",
      "answerHints.pl",
      "PGML.pl",
      "contextLimitedRadical.pl",
    ,
  "PGcourse.pl",
);

##############################################
    
Context("LimitedRadical");
#Context()->flags->set(reduceConstants=>0, reduceConstantFunctions=>0, formatStudentAnswer=>parsed);
                               
Context()->operators->redefine(',',using=>',',from=>'Numeric');
Context()->operators->redefine('or',using=>',',from=>'Numeric');
Context()->operators->set(
      ','=>{string=>' or ',TeX=>'\hbox{ or }'},
      'or'=>{string=>' or ',TeX=>'\hbox{ or }'}
    );
Context()->lists->set(List => {separator => " or "});

# add solution strings to context- this means that if
# students enter these (and they are not correct), then 
# WW will not give a Context warning
Context()->strings->add("no real solutions"=>{},
  "no real solution"=>{alias=>'no real solutions'}, 
  "none"=>{alias=>'no real solutions'}, 
  );

$mymsg = "Are you sure your solution satisfies the equation?";
Context()->{error}{msg}{"Unexpected character '='"} 
  = "$mymsg";

$var = "x";
$a = -1;
$b = 2;
$c = 1;

while( ($b**2-4*$a*$c)>=0 or abs($b**2-4*$a*$c)>180 )
{
    $a = random(2,5,1);
    $b = non_zero_random(-10,10,1);
    $c = non_zero_random(-10,10,1);
}
$questionFormula = Formula("$a*$var^2+$b*$var+$c")->reduce;

# for the answer
$ans = Compute("no real solutions");

##############################################

TEXT(beginproblem());
BEGIN_PGML
Solve the quadratic equation

     [` [$questionFormula]= 0`]

     [_______________________]

* Enter multiple answers separated by commas, such as [`x=1,x=-1`]
* If you need to use the square root symbol, 
as in [`x=\sqrt{17}`], type it using: *sqrt(17)*
* If there are no real solutions, enter *no real solutions*

END_PGML

##############################################

ANS($ans->cmp(checker=>sub {
     my ($correct,$student,$ansHash) = @_;
     if(!Value::String->typeMatch($student))
     {
        Value->Error("$mymsg");
        return 0;
     }
        $check = ($student->value eq "no real solutions");
        Value->Error("$mymsg") unless($check);
        return $check;
    },
));


##############################################

BEGIN_PGML_SOLUTION
This equation does 
not easily factor, so we will use the Quadratic Formula:

     [`
        \begin{aligned}
            [$var] & = \frac{-b\pm\sqrt{b^2-4ac}}{2a}\\
                   & = \frac{[$b*-1]\pm\sqrt{[$b**2]-[$c*$a*4]}}{[$a*2]}\\
                   & = \frac{[$b*-1]\pm\sqrt{[$b**2-$c*$a*4]}}{[$a*2]}\\
        \end{aligned}
     `]

The square root of a negative number is _not_ real, so we conclude that there are *no real solutions* to this
equation.  
END_PGML_SOLUTION

##############################################

ENDDOCUMENT();
