# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Solve the following quadratic equation
#
#       b x^2  + a = 0
#
# If you need to use the square root symbol, 
# as in x=sqrt{17}, type it like: *sqrt(17)*
#
# a and b are both prime numbers
#
# There are no real solutions to these equations.
#
# Last updated: Hughes 8/4/13
# ENDDESCRIPTION

## DBCCSS('A-REI.4.b')
## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('equations', 'Simplification')
## KEYWORDS('solve','quadratic','equation','square root','prime','subtract')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################
DOCUMENT();

loadMacros(
      "PGstandard.pl",
      "MathObjects.pl",
      "answerHints.pl",
      "PGML.pl",
      "contextLimitedRadical.pl",
    ,
  "PGcourse.pl",
);

##############################################
    
Context("LimitedRadical");
#Context()->flags->set(reduceConstants=>0, reduceConstantFunctions=>0, formatStudentAnswer=>parsed);
                               
Context()->operators->redefine(',',using=>',',from=>'Numeric');
Context()->operators->redefine('or',using=>',',from=>'Numeric');
Context()->operators->set(
      ','=>{string=>' or ',TeX=>'\hbox{ or }'},
      'or'=>{string=>' or ',TeX=>'\hbox{ or }'}
    );
Context()->lists->set(List => {separator => " or "});

# add solution strings to context- this means that if
# students enter these (and they are not correct), then 
# WW will not give a Context warning
Context()->strings->add("no real solutions"=>{},
  "no real solution"=>{alias=>'no real solutions'}, 
  "none"=>{alias=>'no real solutions'}, 
  );

$mymsg = "Are you sure your solution satisfies the equation?";
Context()->{error}{msg}{"Unexpected character '='"} 
  = "$mymsg";

$var = "x";
@primes = (2,3,5,7,11,13,17,19,23,29,31,37,41,43,47,53,59,61,67);
$index1 = random(0,$#primes,1);
$index2 = $index1;
$index2 = random(0,$#primes,1) while($index2==$index1);

# choose $a and $b
$a = $primes[$index1];
$b = $primes[$index2];

# for the answer
$ans = Compute("no real solutions");

##############################################

TEXT(beginproblem());
BEGIN_PGML
Solve the quadratic equation

     [` [$b][$var]^2  + [$a]= 0`]

     [_______________________]

* Enter multiple answers separated by commas, such as [`x=1,x=-1`]
* If you need to use the square root symbol, 
as in [`x=\sqrt{17}`], type it using: *sqrt(17)*
* If there are no real solutions, enter *no real solutions*

END_PGML

##############################################

ANS($ans->cmp(checker=>sub {
     my ($correct,$student,$ansHash) = @_;
     if(!Value::String->typeMatch($student))
     {
        Value->Error("$mymsg");
        return 0;
     }
        $check = ($student->value eq "no real solutions");
        Value->Error("$mymsg") unless($check);
        return $check;
    },
));


##############################################

BEGIN_PGML_SOLUTION
Let's try to rearrange the equation so that [`[$var]`] is isolated

    [`
       \begin{aligned}
                 [$b][$var]^2+[$a] = 0& \Rightarrow [$b][$var]^2=-[$a] \\
                                    & \Rightarrow [$var]^2=-\frac{[$a]}{[$b]} \\
       \end{aligned}
    `]

The square root of a negative number is _not_ real, so we conclude that there are *no real solutions* to this
equation.  Note that _all_ quadratic equations have *two* (possibly repeated) solutions; in this problem both solutions are _not_ real. 
END_PGML_SOLUTION

##############################################

ENDDOCUMENT();

