##DESCRIPTION
##  sample matching problem
##ENDDESCRIPTION
## DBsubject(Calculus - single variable)
## DBchapter(Limits and continuity)
## DBsection(Definitions and existence (conceptual))
## Institution(Univeristy of Utah)
## Author(Utah ww group)
## Level(4)
## Static(1)
## MO(1)
## TitleText1('Calculus I')
## AuthorText1('Jerrold Marsden and Alan Weinstein')
## EditionText1('2')
## Section1('.')
## Problem1('')
## TitleText2('Calculus')
## AuthorText2('Dale Varberg, Edwin J. Purcell, and Steven E. Rigdon')
## EditionText2('9')
## Section2('.')
## Problem2('')
## KEYWORDS('calculus')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());
$showPartialCorrectAnswers = 0;


Context("Numeric");
$ans=MultiAnswer(Compute(2),Compute(3),Compute(1))->with(
  singleResult=>1,
  checker => sub {
    my ($correct, $student) = @_;
    if ($student->[0]==$correct->[0] && $student->[1]==$correct->[1] && 
        $student->[2]==$correct->[2]){
       return 1;
    }
    elsif ($student->[0]==$student->[1] || $student->[0] == $student->[2] ||
       $student->[1]==$student->[2]){
      Value->Error("Don't use the same choice twice.");
    }
  }
);

Context()->texStrings;
BEGIN_TEXT
In this problem we consider three functions
\(f\).  The first two are continuous at \(x=0\), i.e.,
\[\lim_{x\to 0}f(x) = f(0).\]

The third function is continuous from the right at \(x=0\), 
\[ \lim_{x \to 0^+} f(x) = f(0).\] 

$BR In order use the \(\epsilon/\delta\) definition to prove the continuity statements, one must give a definition of \(\delta\) in terms of \(\epsilon\)
such that 
\[0<|x-0| < \delta \quad \Longrightarrow \quad |f(x)-f(0)| <
\epsilon. \]
To prove the right-continuity statement requires a definition of \(\delta\) in terms of \(\epsilon\) such that 
$BR
 \[0 < x-0 < \delta \quad \Longrightarrow \quad |f(x)-f(0)| <
\epsilon. \]
$PAR 
For each function in the list below, enter the number (1,2, or 3) of one of these choices 
$BR $BBOLD 1.$EBOLD \(\delta = \epsilon^2 \)
$BR $BBOLD 2.$EBOLD \(\delta = \epsilon\)
$BR $BBOLD 3.$EBOLD \(\delta = \sqrt{\epsilon}\)
$BR 
so that your choices establish
continuity of the first two functions, and right-continuity of the third function,  at \(x = 0\).    
${BBOLD}You may use each choice only once. $EBOLD  
$PAR
$BR \(f(x) = x \): \{ $ans->ans_rule(5) \}
$BR \(f(x) = x^2 \): \{ $ans->ans_rule(5) \}
$BR \(f(x) = \sqrt{x}\): \{ $ans->ans_rule(5) \}
END_TEXT
Context()->normalStrings;

ANS($ans->cmp());

ENDDOCUMENT();        # This should be the last executable line in the problem.
