## DESCRIPTION
## Construct a basis for a given plane in R^3.
##   All coefficients will be small integers, and there will be a basis with small integer entries.
##   There are infinitely many possible solutions. Any of them are accepted.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Euclidean spaces')
## DBsection('Basis and dimension')
## Date('March 3, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('2.4: Bases and Rank')
## Problem1('')
## KEYWORDS('linear algebra','plane','basis')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Matrix");

do {
    $c1 = random(-2,3);
    $c2 = random(-2,3);
    $c3 = random(-2,3);
    $c4 = random(-2,3);
    $C = Matrix([[$c1,$c2],[$c3,$c4]]);
        
    $r1 = list_random(-5,-4,-3,-2,-1,1,2,3,4);
    $r2 = list_random(-5,-4,-3,-2,-1,1,2,3,4);
    $R = Matrix([[1,0,$r1],[0,1,$r2]]);
        
    $A = $C * $R;
} until (abs($C->det) > 0.1 && (abs($r1) > 0 || abs($r2) > 0));

Context("Vector");
Context()->lists->set(Vector => {open => "(", close => ")"});
Context()->parens->set("(" => {type => "Vector", close => ")"});
$v1 = Vector($A->row(1));
$v2 = Vector($A->row(2));

$x = -$r1;
$y = -$r2;
$covec = Vector([$x,$y,1]);

$multians = MultiAnswer($v1,$v2)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if(abs($s0 . $covec) < 0.001 && abs($s1 . $covec) < 0.001 && $s0->norm > 0.001 && $s1->norm > 0.001) {
            return 1 - ($s0->isParallel($s1));
        } else {
            return 0;
        }
    }
);

Context("Numeric");
Context()->variables->add(y => 'Real');
Context()->variables->add(z => 'Real');
Context()->noreduce('(-x)-y','(-x)+y');

$eq = Formula("$x x + $y y + z")->reduce;


##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Compute a basis for the plane in \( \mathbb{R}^3 \) with equation \( $eq = 0 \).
$BR
$BR
\( \Bigg\lbrace \)\{$multians->ans_array(3)\}, \{$multians->ans_array(3)\}\( \Bigg\rbrace \).
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
