# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 02/06/09

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear, constant coefficients, homogeneous (repeated roots))
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('11.10')
## Problem1('19')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('11.10')
## Problem2('19')
## KEYWORDS('calculus','differential equations','trigonometric functions')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real', A=>'Real', B=>'Real', C=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,6,1);
$asq = $a*$a;

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

${BBOLD}(a)$EBOLD
Find the general solution of the differential equation
\[ y''(t) + $asq y(t) = 0. \]
general solution = \{ ans_rule(35) \}
$BR
${BITALIC}(Use the letters A and B for any constants you have
in your solution.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
For each of the following initial conditions, find a particular
solution.
$BR
(i)   \(y(0)=0\),   \(y'(0)=1\):  \( y = \) \{ ans_rule(35) \}
$BR
(ii)  \(y(0)=1\),   \(y'(0)=0\):  \( y = \) \{ ans_rule(35) \}
$BR
(iii) \(y(0)=1\),   \(y(1)=0\):   \( y = \) \{ ans_rule(35) \}
$BR
(iv)  \(y(0)=0\),   \(y(1)=1\):   \( y = \) \{ ans_rule(35) \}
$BR
${BITALIC}(On a sheet of paper, sketch a graph of the solutions you
found in part (b) to see how they differ.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS(Compute("A*cos($a*t) + B*sin($a*t)")->cmp( checker => sub {
    my ( $cor, $stu, $ansHash ) = @_;
    my $opt = $cor->substitute( A=>'C' );
    $opt = $opt->substitute( B=>'A' );
    $opt = $opt->substitute( C=>'B' );
    return ($cor == $stu) || ($opt == $stu);
} ) );
ANS(Compute("(1/$a)*sin($a*t)")->cmp() );
ANS(Compute("cos($a*t)")->cmp() );
ANS(Compute("cos($a*t) - (cos($a)/sin($a))*sin($a*t)")->cmp() );
ANS(Compute("(1/sin($a))*sin($a*t)")->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
Since \(\omega^2=$asq\), \(\omega=$a\), and so the general
solution is of the form 
\[ y(t)=A \cos($a t) + B \sin($a t).\]

$PAR
${BBOLD}(b)$EBOLD
(i)
\(y(0)=0\)
gives \(A \cos(0) + B \sin(0)=0\) so that \(A=0\); then
\(y'(t)=$a B \cos($b t)\), so \(y'(0) = 1\) requires 
\($a B = 1\), so
\[
y(t)={{1} \over {$a}} \sin($a t).
\]
$PAR
(ii)
\(y(0)=1\),
gives \(A \cos(0) + B \sin(0)=1\) so that \(A=1\).  Then
\( y'(t)=$a B \cos($a t) - $a \sin($a t) \), so
\(y'(0)=0\) gives \($a B =0\) and so
\[ y(t)= \cos($a t). \]
$PAR
(iii) 
\(y(0)=1\)
gives \(A \cos(0) + B \sin(0)=1\) so that \(A=1\).
Then \(y(1)=0\) gives \(A \sin $a + \cos $a = 0\) and so
\(B = {{-\cos $a}\over{\sin $a}}\), so
\[
y(t)= \cos($a t) - {{\cos $a}\over{\sin $a}} \sin($a t).
\]
$PAR
(iv)
\(y(0)=0\),
gives \(A \cos(0) + B \sin(0)=0\) so that \(A=0\).
Then \(y(1)=1\) gives \(B \sin($a) =1\) and so
\(B = {{1} \over {\sin($a)}}\), so
\[
y(t)={{1} \over {\sin($a)}} \sin($a t).
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
