## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## KEYWORDS('derivative' 'extrema' 'optimization')
## Tagged by tda2d

## DBsubject('Calculus')
## DBchapter('Applications of Differentiation')
## DBsection('Optimization Problems')
## Date('8/23/07')
## Author('')
## Institution('Union College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "Parser.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",            # Union College utilities
  "PGcourse.pl",           # Customization file for the course
);

TEXT(beginproblem);

###################################
# Setup

$a = random(200,320,10);
$b = random(3,6,1);
$c = random(12,16,1);

$length = sqrt($a*($b+$c)/(2*$b));
$width = sqrt(2*$b*$a/($b+$c));

$ma = MultiAnswer(Real($length),Real($width))->with(
  singleResult => 1, separator => " x ", tex_separator => '\times',
  checker => sub {
    my $correct = shift; my $student = shift;
    my ($a,$b) = @{$correct}; my ($A,$B) = @{$student};
    return ($a == $A && $b == $B) || ($a == $B && $b == $A);
  }
);

###################################
# Main text

BEGIN_TEXT
A fence is to be built to enclose a rectangular area of $a square
feet.  The fence along three sides is to be made of material that
costs $b dollars per foot, and the material for the fourth side
costs $c dollars per foot.  Find the dimensions of the enclosure
that is most economical to construct.
$PAR
Dimensions: \{$ma->ans_rule(10)\} x \{$ma->ans_rule(10)\}
END_TEXT

###################################
# Answers

$showPartialCorrectAnswers = 1;
ANS($ma->cmp);

###################################

ENDDOCUMENT();        # This should be the last executable line in the problem.
