## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## KEYWORDS('derivative' 'extrema' 'optimization')
## Tagged by tda2d

## DBsubject('Calculus')
## DBchapter('Applications of Differentiation')
## DBsection('Optimization Problems')
## Date('8/23/07')
## Author('')
## Institution('Union College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",            # Union College utilities
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl",           # Customization file for the course
);

TEXT(beginproblem);

###################################
# Setup

$a = random(6,12,2);
$b = random(14,20,2);
$x = (($a + $b)-sqrt(($a+$b)**2 - 3*$a*$b))/6;

###################################
#Compute answers

$length = $b - 2*$x;
$width = $a - 2*$x;

$ma = MultiAnswer(Real($length),Real($width))->with(
  singleResult => 1, separator => " x ", tex_separator => '\times',
  checker => sub {
    my $correct = shift; my $student = shift;
    my ($a,$b) = @{$correct}; my ($A,$B) = @{$student};
    return ($a == $A && $b == $B) || ($a == $B && $b == $A);
  }
);

###################################
# Main text

BEGIN_TEXT
An open box is to be made out of a $a-inch by $b-inch piece of
cardboard by cutting out squares of equal size from the four corners
and bending up the sides.  Find the dimensions of the resulting box
that has the largest volume.
$PAR
Dimensions of the bottom of the box: \{$ma->ans_rule(10)\} x \{$ma->ans_rule(10)\}$BR
Height of the box: \{ans_rule(10)\}
END_TEXT

###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($ma->cmp);
ANS(Real($x)->cmp);

###################################

ENDDOCUMENT();        # This should be the last executable line in the problem.
