## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Trigonometric integrals)
## Institution(John Wiley & Sons)
## Author(Brian Burns)
## MLT(tanodd_sec)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Anton')
## EditionText1('9')
## Section1('7.3')
## Problem1('30')
## KEYWORDS('derivative')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "PGcourse.pl"
);

###################################
# Setup
Context("Numeric");
$a=random(2,9,1);
Context()->variables->are(t=>'Real');
Context()->flags->set(limits=>[-5*pi/(12*$a),5*pi/(12*$a)]);
$g=Compute("(tan($a*t))^5*(sec($a*t))^4");
$F=Compute("(tan($a*t))^8/(8*$a)+(tan($a*t))^6/(6*$a)");
$ans=Compute("$F")->cmp(
    checker => sub {
        my ($correct, $student, $ansHash) = @_;
        my $dstud = $student->D('t');
        return ($dstud == $g);
    }
);
###################################
BEGIN_PGML
Evaluate the integral

>>[`` \int  [$g] \;dt ``] = [_]{$ans}{15} +C <<

on the interval [`-\pi/2 < [$a]t < \pi/2 `].
END_PGML

###################################
BEGIN_PGML_SOLUTION
Substitute [`u=\tan([$a]t)`], [`du = [$a]\sec^2([$a]t)\;dt`], and use the trig identity [`\sec^2([$a]t) = 1+\tan^2([$a]t) = 1+u^2`] to obtain

    [``\begin{aligned} \int \tan^5([$a]t)\sec^4([$a]t)\;dt &= \int \tan^5([$a]t)\sec^2([$a]t)\sec^2([$a]t)\;dt\\
     = \int u^5(1+u^2) \frac{du}{[$a]}&= \frac{1}{[$a]}\int u^5+u^7\; du +C\\
     = \frac{1}{[$a]}\left(\frac{u^6}{6}+\frac{u^8}{8}\right) +C &= \frac{1}{[$a]}\left(\frac{\tan^6([$a]t)}{6}+\frac{\tan^8([$a]t)}{8}\right)+C \end{aligned} ``]

Another approach uses the substitution [`w=\sec([$a]t)`], [`dw = [$a]\sec([$a]t)\tan([$a]t)\;dt`] and the above trig identity in the form [`\tan^2([$a]t)=\sec^2([$a]t)-1`] to transform the integral into [``\int (w^2-1)^2 w^3\;\frac{dw}{[$a]}``].  The result using this approach looks a little different from the previous result but using the above trig identity one can show that the difference between the two results is a constant, so the difference doesn't matter because the indefinite integral is defined only up to a constant anyhow.

END_PGML_SOLUTION

###################################
ENDDOCUMENT();
