# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 11/20/08

## Author('Gavin LaRose');
## DBsubject(Calculus - single variable)
## DBchapter(Limits and continuity)
## DBsection(Applications - instantaneous rate of change)
## Institution(University of Michigan)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('2.2')
## Problem1('17')
## KEYWORDS('calculus','derivative','instantaneous velocity','rate of change')

## Textbook tags
## HHChapter('Key Concept: The Derivative')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

Context("Numeric");

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

# we plot one of the functions f(x) = ln(x+.5)+1-ln(.5) or 
# g(x) = e^x

$whichf = random(0,1,1);
$y0 = -0.5;
if ( $whichf ) {
    $f = "log(x+0.5)+1-log(.5) for x in <0,6> using color:blue and weight:2";
#    $f = "sqrt(x+0.5)+1 for x in <0,6> with color:black and weight:2";
    $grDescr = "a graph of a function with a positive, decreasing slope.";
    $x0 = 0;  $x1 = 5;
    $y1 = 4;
    @xlabels = ( 1,2,3,4,5 );
    $labyloc = -0.25;

    ## in this case f(x1)/x1 > f(x0)/x0
    $pop = PopUp( [ "?", "<", "=", ">" ], ">" );
} else {
    $f = "e^x for x in <0,3> using color blue and weight 2";
    $grDescr = "a graph of a function with a positive, increasing slope.";
    $x0 = 0;  $x1 = 3;
    $y1 = 25;
    @xlabels = ( 0.5,1,1.5,2,2.5 );
    $labyloc = 0.25;

    ## in this case f(x1)/x1 < f(x0)/x0 (for the x0, x1 we choose)
    $pop = PopUp( [ "?", "<", "=", ">" ], "<" );
}

$gr = init_graph($x0,$y0,$x1,$y1,axes=>[0,0],size=>[250,250]);
$gr->lb('reset');
$solgr = init_graph($x0,$y0,$x1,$y1,axes=>[0,0],size=>[250,250]);
$solgr->lb('reset');
add_functions($gr,$f);
add_functions($solgr,$f);

# now add some arbitrary x-coordinates
$dx = random(1,3);
@xc = ($x0+$dx, $x0+2*$dx, $x0+3*$dx, $x0+4*$dx, $x0+5*$dx);
@fstrings = ();
for( my $i=0; $i<@xc; $i++ ) {
    $just = ( $i == @xc - 1 ) ? 'right' : 'center';
    $gr->lb(new Label($xlabels[$i],$labyloc,"$xc[$i]",'black',$just,'bottom'));
    $solgr->lb(new Label($xlabels[$i],$labyloc,"$xc[$i]",'black',$just,
			 'bottom'));
    push( @fstrings, "f($xc[$i])", {} );
}
Context()->strings->add( @fstrings );

## pick a random coordinate to ask about
$j = random(2,4,1);
$j++ if ( ! $whichf && $j == 2 );  ## make slope inequality more obvious
$xa = $xc[$j];
$xb = $xc[$j-1];

## now add the line to the solution graph
if ( $whichf ) {
    add_functions( $solgr, "x*(log($xlabels[$j]+0.5)+1-log(0.5))/" .
		   "$xlabels[$j] " .
		   "for x in <$x0, $x1> using color:black and weight:2" );
} else {
    add_functions( $solgr, "x*(e^($xlabels[$j]))/$xlabels[$j] " .
		   "for x in <$x0, $x1> using color:black and weight:2" );
}
$solgrDescr = $grDescr . "  a line is also graphed through (0,0) and (" .
    "$xa,f($xa)).";

## then we allow entry of the points in either order, so we use
##    a multianswer grader to mark them
$mp = MultiAnswer( Compute(0), Compute(0), Compute($xa),
		   String("f($xa)") )->with(
    singleResult => 0,
    checkTypes => 0,
    checker => sub {
	my ( $correct, $student, $ansHash ) = @_;
	my @x = ();  my @stux = ();
	my @y = ();  my @stuy = ();
	( $x[0], $y[0], $x[1], $y[1] ) = @{$correct};
	( $stux[0], $stuy[0], $stux[1], $stuy[1] ) = @{$student};

	my @ret = (0, 0, 0, 0);
	if ( ref($x[0]) == ref($stux[0]) && $x[0] == $stux[0] ) {
	    $ret[0] = 1;
	    $ret[1] = ((ref($y[0]) == ref($stuy[0])) && ($y[0] == $stuy[0]));
	    $ret[2] = ((ref($x[1]) == ref($stux[1])) && ($x[1] == $stux[1]));
	    $ret[3] = ((ref($y[1]) == ref($stuy[1])) && ($y[1] == $stuy[1]));
	} elsif ( ref($x[1]) == ref($stux[1]) && $x[1] == $stux[1] ) {
	    $ret[0] = 0;
	    $ret[2] = 1;
	    $ret[1] = ((ref($y[0]) == ref($stuy[0])) && ($y[0] == $stuy[0]));
	    $ret[3] = ((ref($y[1]) == ref($stuy[1])) && ($y[1] == $stuy[1]));
	} elsif ( ref($x[0]) == ref($stux[1]) && $x[0] == $stux[1] ) {
	    $ret[2] = 1;
	    $ret[0] = ((ref($x[1]) == ref($stux[0])) && ($x[1] == $stux[0]));
	    $ret[1] = ((ref($y[1]) == ref($stuy[0])) && ($y[1] == $stuy[0]));
	    $ret[3] = ((ref($y[0]) == ref($stuy[1])) && ($y[0] == $stuy[1]));
	} elsif ( ref($x[1]) == ref($stux[0]) && $x[1] == $stux[0] ) {
	    $ret[0] = 1;
	    $ret[2] = 0;
	    $ret[1] = ((ref($y[1]) == ref($stuy[0])) && ($y[1] == $stuy[0]));
	    $ret[3] = ((ref($y[0]) == ref($stuy[1])) && ($y[0] == $stuy[1]));
	} else {
	    ## none of the x-values matched, so we have just the y-values
	    ##    to check
	    $ret[0] = 0;
	    $ret[2] = 0;
	    if ( ref($y[0]) == ref($stuy[0]) && $y[0] == $stuy[0] ) {
		$ret[1] = 1;
		$ret[3] = ( (ref($y[1]) == ref($stuy[1])) &&
			    ($y[1] == $stuy[1]) );
	    } else {
		$ret[1] = ( (ref($y[0]) == ref($stuy[1])) &&
			    ($y[0] == $stuy[1]) );
		$ret[3] = ( (ref($y[1]) == ref($stuy[0])) &&
			    ($y[1] == $stuy[0]) );
	    }
	}
	return [ @ret ];
    }
);

Context()->texStrings;
BEGIN_TEXT

Suppose \( y = f(x) \) graphed in the figure below
represents the cost of manufacturing \( x \) kilograms of a chemical.
$PAR
$BCENTER
\{ image( insertGraph($gr), tex_size=>500, height=>250, width=>250,
	  extra_html_tags=>'alt="' . $grDescr . '"' ) \}
$ECENTER
$PAR
\( f(x)/x \) represents the average cost of
producing 1 kilogram of the chemical when \( x \) kilograms are made.
This problem asks you to visualize these averages graphically.

$PAR
${BBOLD}(a)$EBOLD
We can represent \( f($xa)/$xa \) as the slope of a line.  Through what
points does this line extend?
$BR
( \{$mp->ans_rule(10)\}, \{$mp->ans_rule(10)\} ) and
( \{$mp->ans_rule(10)\}, \{$mp->ans_rule(10)\} )
$BR
${BITALIC}(Enter your answers as (x,y) coordinates, where the x
coordinate in each case is an integer value and the y coordinate
is either an integer or the y-value on the curve.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
Complete the inequality:
$BR
\( f($xb)/$xb \) \{ $pop->menu() \} \( f($xa)/$xa \)

END_TEXT
Context()->normalStrings;

ANS($mp->cmp() );
ANS($pop->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

\(f($xa)/$xa\) is the slope of the line connecting
\((0,0)\) to \(($xa,f($xa))\), as shown below.
$PAR
$BCENTER
\{ image( insertGraph($solgr), tex_size=>500, height=>250, width=>250,
	  extra_html_tags=>'alt="' . $solgrDescr . '"' ) \}
$ECENTER
$PAR
It is clear from the figure that
\( f($xb)/$xb \) \{ $pop->correct_ans() \} \(f($xa)/$xa\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
