## DESCRIPTION
##   Domain and Range of Multivariable Function



##  Test version to try to standardize student input of function part on left of >= 1



## ENDDESCRIPTION

## KEYWORDS('Multivariable', 'Continuity', 'Domain', 'Range')

## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Limits and Continuity')
## Date('2/26/2013')
## Author('Paul Seeburger')
## Institution('Monroe Community College')

DOCUMENT();
loadMacros("PG.pl",
           "PGbasicmacros.pl",
           "PGchoicemacros.pl",
           "PGanswermacros.pl",
           "PGauxiliaryFunctions.pl",
           "PGgraphmacros.pl",
           "MathObjects.pl",
           "contextArbitraryString.pl",
           "answerHints.pl",
            );
 

## Show partial correct answers
$showPartialCorrectAnswers = 1;

Context()->variables->add(y=>"Real");

#$a = random(2, 25);
$b = random(2, 17, 1);
$c = random(3,7,2);
$a=$b*$c;

$f = Formula("sqrt(x^2 + $b*y^2 - $a)")->reduce();


## Ok, we are ready to begin the problem...
##
TEXT(beginproblem());

Context()->texStrings;
BEGIN_TEXT
$BR

State the $BBOLD Domain$EBOLD of \(f\), i.e., the largest set on which the function \(\displaystyle{f(x,y) = $f}\)
is continuous.$BR
Insert the correct condition in the set-builder notation below.
$PAR Be sure to use the clearest possible form for the condition.
$PAR\(D_{f}: $LBRACE(x, y):\) \{ans_rule(20)\} \($RBRACE\)
$PAR
$BR
Now enter the $BBOLD Range$EBOLD of \(f\) using interval notation. $BR
(Note: You can enter \(\infty\) as \(\text{inf}\) and \(-\infty\) as \(\text{inf}\).)
$PAR
\(R_{f}:\) \{ans_rule()\}

    
END_TEXT

$showHint = 5;
BEGIN_TEXT
$PAR
If you don't get the $BBOLD domain $EBOLD correct in $showHint tries, you can get a hint.
END_TEXT
BEGIN_HINT
The condition for the $BBOLD domain $EBOLD in this problem should be in Ellipse Form, i.e., x^2/a + y^2/b = 1.  Of course there may also be an inequality involved.  =)
END_HINT
Context()->normalStrings;

$domainA = "x^2/$a + y^2/$c >= 1";

$domain2 = "y^2/$c + x^2/$a >= 1";
$domain3 = "1 <= x^2/$a + y^2/$c";
$domain4 = "1 <= y^2/$c + x^2/$a";

$domain = $domainA;
$domain =~ tr/ *//d; # remove *
$domain2 =~ tr/ *//d; # remove *
$domain3 =~ tr/ *//d; # remove *
$domain4 =~ tr/ *//d; # remove *

Context("ArbitraryString");

    
    ANS(Compute($domainA)->cmp(checker => sub {
      my ($correct,$student,$ans) = @_;
      $student = $student->value; # get student answer string
      $student =~ s/ //g; # remove whitespace

      my $m1 = index($student, ">=");
      my $m2 = index($student, "<=");
      my $m3 = $m2 + 2;
      my $sa = $student;
      my $e = 0;
      if (($m1 > -1 && $m2 > -1) || ($m1 == -1 && $m2 == -1)) {$e = 1;}
      else {
         if ($m1 > -1) {
            $sa = substr $student, 0, $m1; # the part before >=
            $ste = substr $student, $m1;    # the rest of the answer
            if ($ste ne ">=1") { $e = 1; }
         }
         if ($m2 > -1) {
            $sa = substr $student, $m3; # the part after <=
            $ste = substr $student, 0, $m2;    # the rest of the answer
            if ($ste ne "1") { $e = 1; }
         }
         Context("Numeric");
         Context()->variables->add(y=>"Real");

         $s = Formula($sa)->reduce();
         $an = Formula("x^2/$a+y^2/$c")->reduce();
         $an2 = Formula("y^2/$c+x^2/$a")->reduce();
      }

#      $student =~ tr/ *//d; # remove * and whitespace  
   
       if ($e == 1) {
           return (0 == 1);
       }
       else {
          return $s eq $an2 ||
                 $s eq $an; # $s == $an   allows more variation
       }
    })->withPostFilter(AnswerHints(
         sub {
            my ($correct, $student, $ans) = @_;
            my $y = $student->value;
            $y =~ tr/ *//d; # remove *
            my $m1 = index($y,"<=");
            my $m2 = index($y,">=");

            return ($m1 == -1 && $m2 == -1) ||
                   ($m1 > -1 && $m2 > -1)
        } =>"Consider your inequality condition more carefully.",
         sub {
            my ($correct, $student, $ans) = @_;
            my $x = "x^2 + $b*y^2 - $a>=0";
            my $x2 = "x^2 + $b*y^2 >= $a";
            my $x3 = "x^2/$b + y^2 >= $c";
            my $y = $student->value;
            $y =~ tr/ *//d; # remove *
            $x =~ tr/ *//d; # remove *
            $x2 =~ tr/ *//d; # remove *
            return $y eq $x ||
                   $y eq $x2;
        } =>"This is a correct condition, but not the simplest form.",
         sub {
            my ($correct, $student, $ans) = @_;
            my $y = $student->value;
            $y =~ tr/ *//d; # remove *
            my $b = substr $y, 0, 2;  #index($y,"1<");
            my $e = substr $y, -2, 2; #index($y,"=1");
            return $b ne "1<" &&
                   $e ne "=1";
        } =>"Consider the simplest form of the boundary of this region.",
)));


#ANS( str_cmp($domain,'remove_whitespace')); 

Context("Interval");
Parser::Number::NoDecimals;

Context()->flags->set(
  reduceConstants=>0, # no decimals
  reduceConstantFunctions=>0, # combine 4+5*2?
  formatStudentAnswer=>'parsed', # no decimals
);

ANS( Compute("[0, inf)")->cmp()->withPostFilter(AnswerHints( 
  Compute("(0, inf)") => "Think more carefully about your left endpoint."
)));


ENDDOCUMENT();




