## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Applications)
## Date(03/19/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-A-mixture)
## Level(5)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('')
## EditionText1('December 9, 2010')
## Section1('.')
## Problem1('')


##################################
#  Initialization


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##################################
#  Setup

Context("Matrix");

Context("Numeric");
Context()->variables->are(t=>"Real");

$r01 = random(5,30,5);
do { $r21 = random(5,30,5); } until (($r21 != $r12) && ($r21 != $r01));
$r23 = $r01;
$r12 = $r21 + $r23;

$V1 = random(30,100,10);
$V2 = random(110,200,10);


# matrix A = M^(-1) K listed by rows (not columns)
@A = ();
$A[0] = Compute("-$r12 / $V1"); # A11
$A[1] = Compute("$r21 / $V2");  # A12
$A[2] = Compute("$r12 / $V1");  # A21
$A[3] = Compute("-($r21 + $r23) / $V2"); # A22

$multians1 = MultiAnswer( $A[0], $A[1], $A[2], $A[3] )->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='5'>
    <tr><td> %s </td><td> %s </td></tr>
    <tr><td> %s </td><td> %s </td></tr>
    </table>",
    tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\[2ex] %s & %s \end{array}\right\rbrack" ,

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
        my @s = @{$student};
        my @score = ();

        foreach my $i (0..3) {
          if ($c[$i]->typeMatch($s[$i]) && $c[$i] == $s[$i]) 
          { $score[$i] = 1; } else { $score[$i] = 0; } 
        }
        return [ @score ];

    }
);



#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{
ColumnTable(
"Consider two brine tanks connected as shown in the figure.
Pure water flows into the top of tank 1 at a rate of 
\( $r01 \ \mathrm{L/min} \).  The brine solution is pumped 
from tank 1 into tank 2 at a rate of \( $r12 \ \mathrm{L/min} \), 
and from tank 2 into tank 1 at a rate of \( $r21 \ \mathrm{L/min} \).  
A brine solution flows out the bottom of tank 2 at a rate of 
\( $r23 \ \mathrm{L/min} \).  Suppose there are \( $V1 \ \mathrm{L} \)
of brine in tank 1 and \( $V2 \ \mathrm{L} \) of brine in tank 2.
Let \( x \) be the amount of salt, in kilograms, in tank 1 after \( t \)
minutes, and \( y \) the amount of salt, in kilograms, in tank 2 after
\( t \) minutes.  
Assume that each tank is mixed perfectly.".
$BR.
$BR.
"Set up a system of first-order differential equations that models this
situation.".
mbox(
[
display_matrix([['\(x^{\,\prime}\)'],['\(y^{\,\prime}\)']]),
' = ', 
display_matrix([[$multians1->ans_rule(10),$multians1->ans_rule(10)],[$multians1->ans_rule(10),$multians1->ans_rule(10)]]),
display_matrix([['\(x\)'],['\(y\)']])
],
allowbreaks=>'yes'
) 
,
$BCENTER.
image("mixing-solutions-01-new.png", width=>246, height=>181, tex_size=>500).
$ECENTER
,
indent => 0, separation => 30, valign => "TOP"
);
\}

END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

;

ENDDOCUMENT();



