# DESCRIPTION
# Skills building problem
# ENDDESCRIPTION

## DBsubject(Precalculus)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: point-slope form)
## Level(2)
## KEYWORDS('point-slope form', 'line')

DOCUMENT();
loadMacros("PGstandard.pl",
           "MathObjects.pl",
           "PGML.pl",
	   "parserPopUp.pl",
# "PGchoicemacros.pl",
# "PGgraphmacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Point");

@alph = qw( a b c d m n p q r s t u v w );
$start_i = random(1,scalar(@alph),1) - 3;
$a = $alph[$start_i];
$b = $alph[$start_i+1];
$c = $alph[$start_i+2];

$v1 = $a;
$v2 = $b;
$v3 = $c;

Context()->constants->add( $a => 2*pi-1, $b => 1.71828, $c => sqrt(17) );
Context()->flags->set( formatStudentAnswer=>'parsed' );

$pt = "($c,$a)";

TEXT(beginproblem());
BEGIN_PGML
Let [`[$a]`], [`[$b]`], and [`[$c]`] be unknown real numbers.  For the line
whose formula is given by
>> [`` y = [$b](x-[$c]) + [$a], ``] <<
determine the slope of the line, and a point that lies on the line.
Your answers will involve [`[$a]`], [`[$b]`], and/or [`[$c]`].

    slope = [___]{$b}{5}  
    a point = [___]{Point("$pt")->cmp(checker=>sub {
	my( $cor, $stu, $ans ) = @_;
	my ( $cx, $cy ) = $cor->value;
	my ( $sx, $sy ) = $stu->value;
	if ( $b*($sx - $c) + $a == $sy ) {
	    Context()->texStrings;
	    $ans->{correct_ans_latex_string} = "$stu->{equation}";
	    Context()->normalStrings;
	    return 1;
	} else {
	    $ans->{correct_ans_latex_string} = "($v3,$v1)";
	    return 0;
	}
    })}{10}

END_PGML

BEGIN_PGML_SOLUTION

Note that this is equivalent to the point-slope form of the line,
>> [`` y - [$a] = [$b](x - [$c]), ``] <<
from which we can see that the slope is [`[$b]`] and a point on
the line is [`([$c], [$a])`].

END_PGML_SOLUTION

ENDDOCUMENT();
