## DESCRIPTION
## Higher order ODEs: constant coefficient second order homogeneous linear ODEs
## ENDDESCRIPTION

## KEYWORDS('differential equations','second order linear ODE','constant coefficients')

## DBsubject('Differential Equations')
## DBchapter('Higher Order Differential Equations')
## DBsection('Second Order Linear')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Notes on Diffy Qs')
## EditionText1('December 9, 2010')
## AuthorText1('Jiri Lebl')
## Section1('2.2')
## Problem1('8')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"parserAssignment.pl",
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->add(
A=>"Real",B=>"Real",y=>"Real",
);
parser::Assignment->Allow;

$a  = random(3,7,1);
do { $y0 = random(2,7,1); } until ($y0 != $a);

# char poly (r-$a)^2
$a2 = 2*$a;
$aa = $a*$a;

$A = $y0;
$B = -($a)*$y0;

$diffeq = "y'' - $a2 y' + $aa y = 0"; # tex

$answera = Compute("y = (A+Bx) e^($a x)");

$answerb = Compute("y = ($A + $B x) e^($a x)")->reduce;

######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
(a) Find the general solution to \( $diffeq \).
In your answer, use \( A \) and \( B \) to 
denote arbitrary constants and \( x \) 
the independent variable.
$BR
\{ ans_rule(40) \}
\{ AnswerFormatHelp("equations") \}
$BR
$BR
(b) Find the particular solution that satisfies
\( y(0) = $y0 \) and \( y'(0) = 0 \).
$BR
\{ ans_rule(40) \}
\{ AnswerFormatHelp("equations") \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $answera->cmp( checker => sub {
    my ( $correct, $student, $self ) = @_;
    my $stu   = Formula($student->{tree}{rop});

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('A'))==Formula(0) || 
      Formula($stu->D('B'))==Formula(0) 
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(1.24);

    my $a11 = $stu->eval(A=>1,B=>0,x=>$x,y=>0);
    my $a12 = $stu->eval(A=>0,B=>1,x=>$x,y=>0);
 
    my $a21 = $stu->D('x')->eval(A=>1,B=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval(A=>0,B=>1,x=>$x,y=>0);

    my $wronskian = $a11*$a22 - $a21*$a12;

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($wronskian==Real(0));

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stux  = Formula($stu->D('x'));
    my $stuxx = Formula($stu->D('x','x'));
    my $stuDE = Formula($stuxx - $a2*$stux + $aa*$stu)
    ->with(test_points=>[[1,1,0.1,0],[2,1,0,0],[1,2,-0.1,0],[1,1,0,1]]);

    return ($stuDE==Formula(0));

}));


ANS( $answerb->cmp() );

COMMENT("MathObject version.");

ENDDOCUMENT();