## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/29/06

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear independence)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## MLT(indep2Q)
## Level(2)
## Static(1)
## MO(1)

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixReduce.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$showPartialCorrectAnswers = 0;
install_problem_grader(~~&std_problem_grader);
SRAND($psvn);

Context('Matrix');

$Z = Matrix([0,0,0]);

$ind = random(0,1,1); # dependent = 0, independent = 1

if ($ind == 0) {

  $popup = PopUp(["Choose","linearly dependent", "linearly independent"],"linearly dependent");

  $A = Matrix([
  [random(-5,5,1),random(-5,5,1),random(-5,5,1)],
  [random(-5,5,1),random(-5,5,1),random(-5,5,1)],
  [random(-5,5,1),random(-5,5,1),random(-5,5,1)],
  ]);
    
  $s1 = random(2,5,1);
  $s2 = random(2,5,1);
  foreach my $i (1..3) {
    change_matrix_entry($A,[3,$i], $s1 * $A->element(1,$i) + $s2 * $A->element(2,$i));
  }
  $w[0] = Real(-1*$s1);
  $w[1] = Real(-1*$s2);
  $w[2] = Real(1);

} else { # independent

  $popup = PopUp(["Choose","linearly dependent", "linearly independent"],"linearly independent");

  do {

  $A = Matrix([
  [random(-5,5,1),random(-5,5,1),random(-5,5,1)],
  [random(-5,5,1),random(-5,5,1),random(-5,5,1)],
  [random(-5,5,1),random(-5,5,1),random(-5,5,1)],
  ]);

  } until ($A->det != 0);

  foreach my $i (0..2) {
    $w[$i] = Real(0);
  }

}

foreach my $i (1..3) {
  $a[$i - 1] = $A->row($i);
}


$multians = MultiAnswer($w[0],$w[1],$w[2])->with(
  singleResult => 1,
  allowBlankAnswers => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      if ($ind == 0) {
        return 0 if ($s[0] == Real(0) and $s[1] == Real(0) and $s[2] == Real(0));
        return $s[0] * $a[0] + $s[1] * $a[1] + $s[2] * $a[2] == $Z;
      } else {
        return ($s[0] == Real(0) and $s[1] == Real(0) and $s[2] == Real(0));
      }
  }
);

Context()->texStrings;
BEGIN_TEXT
Are the vectors \( \vec{u} = $a[0] \), \( \vec{v} = $a[1] \) and \( \vec{w} = $a[2] \) linearly independent?  
$BR
\{ $popup->menu \}
$BR
$BR
If they are linearly dependent, find scalars that are not all zero such that the equation below is true.
If they are linearly independent, find the only scalars that will make the equation below true.
$BR
$BR
\{ $multians->ans_rule(5) \} \( \vec{u} + \)
\{ $multians->ans_rule(5) \} \( \vec{v} + \)
\{ $multians->ans_rule(5) \} \( \vec{w} = \)
\( \vec{0} \).
END_TEXT
Context()->normalStrings;

ANS( $popup->cmp );
ANS( $multians->cmp );

COMMENT('Independent half of the time, dependent half the time.');
ENDDOCUMENT();
