# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# 
# Give examples of types of numbers
#
# 
#
# ENDDESCRIPTION

## DBsubject('Numbers')
## DBchapter('Real Numbers')
## DBsection('Misc.')
## KEYWORDS('real numbers')
## DBCCSS('')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserOneOf.pl",
  "answerHints.pl",
  "contextFraction.pl",
  "PCCmacros.pl",
  "PGcourse.pl",
);

##############################################
TEXT(beginproblem());

Context("Numeric");


##############################################

BEGIN_PGML
Give examples of each type of number. If no such number exists, enter [|DNE|]* or [|NONE|]*.

a. Give an example of a whole number that is not an integer.

    [_________]

a. Give an example of an integer that is not a whole number.

    [_________]

a. Give an example of a rational number that is not an integer.

    [_________]


a. Give an example of a irrational number.

    [_________]


a. Give an example of a irrational number that is also an integer.

    [_________]


END_PGML



##############################################

$showPartialCorrectAnswers = 1;

ANS(OneOf(DNE,NONE)->cmp);
ANS(Compute(-1)->cmp(correct_ans => "any of -1, -2, -3, ...", 
                     correct_ans_latex_string => "\text{any of }-1, -2, -3,\ldots",
                     checker=>sub{
  my ( $correct, $student, $ansHash ) = @_;
  return ((int($student) == $student) and ($student < 0));
}
)->withPostFilter(AnswerHints( 
  sub{
  my ( $correct, $student, $ansHash ) = @_;
  return (int($student) == $student) and ($student > 0);
  } => "This is an integer, but it is also a whole number",
))
);


ANS(Real("0.5")->cmp(correct_ans => "many possibilities: 1/2, 2.3,...", 
                     correct_ans_latex_string => "\text{many possibilities: }\frac{1}{2}, 2.3, \ldots",
                     checker=>sub{
  my ( $correct, $student, $ansHash ) = @_;
  $student = abs($student);
  $student = $student->value;
  Value->Error("Your answer is an integer") if $student == int($student);
  my $OK = 0;
  #in theory this should detect all rational numbers with denominators < the 30th Fibonacci number (and many with larger denominators too). Machine rounding error may come into play though.
  for my $i (0..30)
    {if (abs($student - int($student)) < 10**(-6)) {$OK = 1; last;}
     $student = 1/($student - int($student));
    } 
  return $OK;
}
)
);


ANS(Compute("sqrt(2)")->cmp(correct_ans => "many possibilities: sqrt(2), pi,...", 
                     correct_ans_latex_string => "\text{many possibilities: }\sqrt{2}, \pi, \ldots",
                     checker=>sub{
  my ( $correct, $student, $ansHash ) = @_;
  $student = abs($student);
  $student = $student->value;
  my $OK = 0;
  #in theory this should detect all rational numbers with denominators < the 30th Fibonacci number (and many with larger denominators too). Machine rounding error may come into play though.
  for my $i (0..30)
    {if (abs($student - int($student)) < 10**(-6)) {$OK = 1; last;}
     $student = 1/($student - int($student));
    } 
  return !$OK;
}
)
);

ANS(OneOf(DNE,NONE)->cmp()->withPostFilter(AnswerHints( 
  sub{
  my ( $correct, $student, $ansHash ) = @_;
  return (int($student) == $student) and ($student > 0);
  } => "This is an integer",
)));

##############################################

BEGIN_PGML_SOLUTION

a. Give an example of a whole number that is not an integer.

    Since all whole numbers belong to integers, we cannot write any whole number which is not an integer. Type [|DNE|]* (does not exist) for this question.

a. Give an example of an integer that is not a whole number. If no such number exists, enter [|DNE|]* or [|NONE|]*.

    Any negative integer, like [`-1`], is not a whole number, but is an integer.

a. Give an example of a rational number that is not an integer. If no such number exists, enter [|DNE|]* or [|NONE|]*.

    Any terminating decimal, like [`1.2`], is a rational number, but is not an integer.

a. Give an example of a irrational number.

    [`\pi`] is the easiest number to remember as an irrational number. Another constant worth knowing is [`e\approx2.718`]. Finally, the square root of most integers are irrational, like [`\sqrt{2}`] and [`\sqrt{3}`].

a. Give an example of a irrational number that is also an integer.  If no such number exists, enter [|DNE|]* or [|NONE|]*.

    All irrational numbers are non-repeating and non-terminating decimals. No irrational numbers are integers. Type [|DNE|]* (does not exist) for this question.


END_PGML_SOLUTION

ENDDOCUMENT();
