## DESCRIPTION
## Calculus One, approximating areas.
## ENDDESCRIPTION


## DBsubject(Calculus)
## DBchapter()
## DBsection()
## Date(09/11/2018)
## Institution(Community College of Denver, Colorado Community College System)
## Author(Percy Makita, Updated Brittni Lorton February 2020)
## KEYWORDS('Calculus 2', 'Chapter 3', 'Section 3', 'Trigonometric Substitution')



###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"PGML.pl",
"PGcourse.pl",
"parserFormulaUpToConstant.pl",
);



TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

###########################
#  Setup
$a=non_zero_random(-12,12);
$f=Formula("1/(x^2+($a)x)")->reduce();
#this is the answer a student would have gotten if they did partial fractions and did not simplify
$ans=FormulaUpToConstant("(1/$a)ln(abs(x/(x+$a)))+C")->reduce();

#this is the answer a student would have gotten if they did completing the square and trig substitution and did not simplify
$ans1 = FormulaUpToConstant("-2/$a*ln|(x+$a/2)/(sqrt(x^2+$a*x))+($a/2)/(sqrt(x^2+$a*x))|+C");



###########################
#  Main text


BEGIN_PGML
Use the technique of completing the square or the method of partial fractions to evaluate the integral

[``\int [$f]dx=``][___________________]{$ans->cmp( checker => sub {
    my ( $correct, $student, $self ) = @_;
    return ($correct == $student || $ans1 == $student);
}), } [@ AnswerFormatHelp("formulas") @]*

Use a capital C for any needed constant in your answer


END_PGML


############################
#  Solution

#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION

COMMENT('MathObject version. Uses PGML.');

ENDDOCUMENT();