## DESCRIPTION
## Calculus Two, Trigonometric Substitution
## ENDDESCRIPTION


## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Trigonometric substitution)
## Date(09/10/2018)
## Institution(Community College of Denver, Colorado Community College System)
## Author(Percy Makita, Updated Brittni Lorton February 2020)
## KEYWORDS('Calculus 2', 'Chapter 3', 'Section 3', 'Trigonometric Substitution')



###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
  "answerHints.pl",
"AnswerFormatHelp.pl",
"PGML.pl",
"PGcourse.pl",
"parserFormulaUpToConstant.pl",
);


TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

###########################
#  Setup
$a = random(1,6,1);
$f=Formula("1/sqrt(x^2-$a^2)")->reduce();

#answer with test points to include both pieces of the domain to ensure students are marked wrong if they forget the absolute value.
$ans1=FormulaUpToConstant("ln(abs(x+sqrt(x**2-$a**2)))+C")->with(limits => [[$a+1,$a+7]], test_at => [-$a-2])->reduce();
#if students do forget the absolute value but their answer is correct besides that, this is the message they get.
$ans1 = $ans1->cmp()->withPostFilter(AnswerHints(
FormulaUpToConstant("ln(x+sqrt(x**2-$a**2))+C")->with(limits=>[[4,10]]) => "Are you missing an absolute value somewhere?"
));

Context()->variables->add(a => "Real"); #adding the variable a to the context for the second piece of this problem
$fa=Formula("1/sqrt(x^2-a^2)")->reduce();
#answer with test points to include both pieces of the domain to ensure students are marked wrong if they forget the absolute value.
$ans2=FormulaUpToConstant("ln(abs(x+sqrt(x**2-a**2)))+C")->with(limits => [[1,-6],[1,-2]], test_at => [[1,2],[1,6]])->reduce(); 
#$ans3 is the another form of the correct answer that we need to define because WW is seeing a as a variable but we are asking students to treat a as a constant.
$ans3=FormulaUpToConstant("ln(abs(x/a+sqrt(x**2-a**2)/a))+C")->with(limits => [[1,-6],[1,-2]], test_at => [[1,2],[1,6]])->reduce(); 
#test_at => [[1,-2],[1,-6]])->reduce(); 

#$ans2 = $ans2->cmp()->withPostFilter(AnswerHints( FormulaUpToConstant("ln(x+sqrt(x**2-a**2))+C")->with(limits=>[[1,6],[4,10]]) => "Are you missing an absolute value somewhere?" ));

#$ans2->cmp( checker => sub {
#    my ( $correct, $student, $self ) = @_;
#    return ($correct == $student || $ans3 == $student);
#}), 

###########################
#  Main text


BEGIN_PGML
Integrate the following using the method of trigonometric substitution. 
Express the final answer in terms of the original variable.

a. [``\int [$f] dx=``][________________________________]{$ans1}[@ AnswerFormatHelp("formulas") @]*

b. [``\int [$fa] dx=``][________________________________]{$ans2->cmp( checker => sub {
    my ( $correct, $student, $self ) = @_;
    return ($correct == $student || $ans3 == $student);
}), }(for [`a`] being any real number) [@ AnswerFormatHelp("formulas") @]*

Use a capital C for any needed constant in your answer


END_PGML



############################
#  Solution

#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION

COMMENT('MathObject version. Uses PGML.');

ENDDOCUMENT();