## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear independence)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix');

$A1 = Matrix([non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)])->transpose;
$A2 = Matrix([non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)])->transpose;

$c1 = random(2,4,1);
$c2 = random(2,4,1);

$b1 = non_zero_random(-3,3,1);
$b2 = non_zero_random(-3,3,1);

$v1 = $c1 * $A1;
$v2 = $c2 * $A2;
$v3 = $b1 * $A1 + $b2 * $A2;

$V = Matrix([
[$v1->element(1,1),$v2->element(1,1),$v3->element(1,1)],
[$v1->element(2,1),$v2->element(2,1),$v3->element(2,1)],
[$v1->element(3,1),$v2->element(3,1),$v3->element(3,1)],
]);

$v1vec = Vector($v1);
$v2vec = Vector($v2);
$v3vec = Vector($v3);

$popup = PopUp(['choose','linearly dependent','linearly independent'],'linearly dependent');

$multians = MultiAnswer($b1/$c1, $b2/$c2, -1)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my $S = Matrix([$s[0],$s[1],$s[2]])->transpose;
      my $Z = Matrix([0,0,0])->transpose;
      return 0 if $S == $Z;
      return $V * $S == $Z ;
  }
);

Context()->texStrings;
BEGIN_TEXT
Let \[ \vec{v}_1 = $v1vec, \ \vec{v}_2 = $v2vec, \ \vec{v}_3 = $v3vec. \]
Are the vectors \( \vec{v}_1, \vec{v}_2\) and \( \vec{v}_3 \) linearly independent?
$BR
\{ $popup->menu \}
$BR
$BR
If the vectors are independent, enter zero in every answer blank since those are only the values that make the equation below true.  
If they are dependent, find numbers, not all zero, that make the equation below true.    You should be able to explain and justify your answer.
$BR
$BR
\( \langle 0,0,0 \rangle = \) 
\{ $multians->ans_rule(5) \} \( $v1vec + \)
\{ $multians->ans_rule(5) \} \( $v2vec + \)
\{ $multians->ans_rule(5) \} \( $v3vec. \)
END_TEXT
Context()->normalStrings;

ANS($popup->cmp);
ANS($multians->cmp);

COMMENT('MathObject version.  Always dependent.');
ENDDOCUMENT();
