##DESCRIPTION
##  A problem to the expectation of the square of a normal random variable.
##ENDDESCRIPTION


## DBsubject(Statistics)
## DBchapter(Confidence intervals)
## DBsection(One sample mean - z)
## Date(10/19/2012)
## Institution(The University of British Columbia)
## Author(Djun Kim)
## Level(3)
## MO(1)
## KEYWORDS('statistics', 'continuous random variables', 'expectation')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "unionTables.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

#######

TEXT(beginproblem());

#######
# Setup

Context("Numeric");

@data = (183, 222, 303, 262,  178,  232,  268,  201,  244,  183,  201,  140); 

foreach $item (@data) {$sum += $item};
$mean = $sum/@data;

######
# Part a


$a1 = Real ( $mean  - (1.96*50/sqrt(12)) );
$a2 = Real ( $mean + (1.96*50/sqrt(12)) );


$part_a_multians = MultiAnswer($a1, $a2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);

######
# Part b

@b_hours = (25, 26, 27, 29, 30, 31);
$b_ind = random(0,5,1);
$b_hours = @b_hours[$b_ind];

$part_b = ceil( ((2 * 1.96 * 50)/ $b_hours)**2  );
$part_b_answer = Real( $part_b );

### 
# Sol

$smean = sprintf("%.2f", $mean);
$sa1 = sprintf("%.1f", $a1);
$sa2 = sprintf("%.1f", $a2);

$snsqt = sprintf("%.2f", 2*1.96*50/$b_hours);
$sn = sprintf("%.2f", (2*1.96*50/$b_hours)**2 );


#######

Context()->texStrings;

BEGIN_TEXT
$PAR

The time in hours for a worker to repair an electrical instrument is a Normally distributed random variable with a mean of \(\mu\) and a standard deviation of 50. The repair times for 12 such instruments chosen at random are as follows:

\{ 
BeginTable(border=>1, tex_border=>"1pt", spacing=>0, padding=>4).
AlignedRow( [@data], align=>LEFT, separation=>0).

EndTable()
\}

$PAR
$BBOLD Part a) $EBOLD
 Find a 95% confidence interval for \( \mu \). For both sides of the bound, leave your answer with 1 decimal place.  
( \{$part_a_multians->ans_rule(6)\} , \{$part_a_multians->ans_rule(6)\} ).

$PAR
$BBOLD Part b) $EBOLD
Find the least number of repair times needed to be sampled in order to reduce the width of the confidence interval to below $b_hours hours.  \{ ans_rule(6) \}


END_TEXT

#######

$showPartialCorrectAnswers = 1;

ANS(  $part_a_multians->cmp(tol=>'0.001' , tolType=>'absolute') );
ANS(  $part_b_answer->cmp(tol=>'0.001' , tolType=>'absolute') );




##############################################################
#
#  Solution
#

Context()->texStrings;

BEGIN_SOLUTION
$BR
$BR
$BBOLD Part a) $EBOLD
$BR
$BR
The answer is \( \underbrace{\bar{x}}_{\text{averaged from data}} \pm 1.96 \times \frac{\sigma}{\sqrt{n}} \)
$BR
\( = $smean \pm 1.96 \times \frac{50}{\sqrt{12}} \)
$BR
\( = [$sa1, $sa2]. \)

$BR
$BR
$BBOLD Part b) $EBOLD
$BR
$BR
We wish for the width \( = 2 \times 1.96 \times \frac{50}{\sqrt{n}} < $b_hours \)
$BR $BR
so we find
$BR $BR
\( 2 \times 1.96 \times 50 < $b_hours \sqrt{n} \)
$BR
\( \implies \sqrt{n} > $snsqt \)
$BR
\( \implies n > $sn \)
$BR
\( \implies n = $part_b_answer \)

END_SOLUTION

Context()->normalStrings;


ENDDOCUMENT();
