 #DESCRIPTION
##  A problem to compute probability associated with a continuous uniform random variable.
##ENDDESCRIPTION


## DBsubject(Statistics)
## DBchapter(Hypothesis tests)
## DBsection(One sample mean - t)
## Date(11/15/2012)
## Institution(UBC)
## Author(Djun Kim)
## Level(3)
## MO(1)
## KEYWORDS('statistics', 'variance', 'ANOVA', '2-sample t')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "PGcourse.pl"
);

#######

TEXT(beginproblem());

Context("Numeric");


$ver = random(0, 3, 1);

#---
# Part a)

$ind = random(0,2,1);
@narr = (27, 31, 41);
@muarr = (91.29, 82.72, 81.35);
@starr = (4.94, 4.61, 4.52);
@tco = (2.055529, 2.042272, 2.021075);

$n = @narr[$ind];
$xbar = @muarr[$ind];
$xsd = @starr[$ind];
$tco = @tco[$ind];

$df = $n-1;

#--
# Part b)

$tcoshow = sprintf("%.2f", $tco);

$blow = sprintf("%.2f", $xbar - $tco*$xsd/($n**0.5));
$bhi = sprintf("%.2f", $xbar + $tco*$xsd/($n**0.5));

@bar1arr = (24.99, 23.69, 31.86, 25.84);
@bar2arr = (32.92, 23.92, 24.62, 24.64);
@ss1arr = (108.47, 89.60, 139.75, 52.69);
@ss2arr = (77.68, 66.35, 109.97, 105.63);

$var1 = sprintf("%.2f", $ss1arr[$ver]/14.);
$var2 = sprintf("%.2f", $ss2arr[$ver]/12.);

$ans[0] = Real($df); # Real((108.47 + 77.68) /(15+13-2));
$ans[1] = Real($df); #Real((89.60  + 66.35) /(15+13-2));
$ans[2] = Real($df); # Real((139.75 + 109.97)/(15+13-2));
$ans[3] = Real($df); # Real((52.69  + 105.63)/(15+13-2));


$dispans = sprintf("%.3f", $ans[$ver]);

#---
# Part b)


$a[0] = $blow;
$b[0] = $bhi;

$a[1] = $blow;
$b[1] = $bhi;

$a[2] = $blow;
$b[2] = $bhi;

$a[3] = $blow;
$b[3] = $bhi;

#$a[1] = Real((23.69-23.92) - 2.056 * sqrt($ans[1] * (1/13 + 1/15)));
#$b[1] = Real((23.69-23.92) + 2.056 * sqrt($ans[1] * (1/13 + 1/15)));

#$a[2] = Real((31.86-24.62) - 2.056 * sqrt($ans[2] * (1/13 + 1/15)));
#$b[2] = Real((31.86-24.62) + 2.056 * sqrt($ans[2] * (1/13 + 1/15)));

#$a[3] = Real((25.84-24.64) - 2.056 * sqrt($ans[3] * (1/13 + 1/15)));
#$b[3] = Real((25.84-24.64) + 2.056 * sqrt($ans[3] * (1/13 + 1/15)));

$ashow = sprintf("%.2f", $a[1]);
$bshow = sprintf("%.2f", $b[1]);

$multians = MultiAnswer($a[$ver], $b[$ver])->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $a_stu, $b_stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( (abs($a_stu - $f1) < 0.005 && abs($b_stu - $f2) < 0.005 ))
      {
         return [1,1];
      }
      else {
         return [0,0];  
      }

  }
);

# $multians = [$blow,$bi];

#---
# Part c)

if ($a[$ver] <= 80 && $b[$ver] >= 80) { # 0 is in the interval
   $mc3_answer      = "We should not reject the hypothesis since 80 km/h is in the interval found in (b).";
   $mc3_distractor4 = "We should reject the hypothesis since 80 km/h is not in the interval found in (b).";
}
else {
   $mc3_answer      = "We should reject the hypothesis since 80 km/h is not in the interval found in (b).";
   $mc3_distractor4 = "We should not reject the hypothesis since 80 km/h is in the interval found in (b).";
}

$mc3_distractor1 = "We should reject the hypothesis since the sample mean was not 80 km/h.";
$mc3_distractor2 = "We should not reject the hypothesis since the sample mean is in the interval found in (b).";
$mc3_distractor3 = "We should reject the hypothesis since 80 km/h is in the interval found in (b).";

$mc3 = new_multiple_choice();
$mc3->qa(
  "The police hypothesized that the mean speed of vehicles over the bridge would be the speed limit, 80 km/h. Taking a significance level of 5 $PERCENT, what should infer about this hypothesis?",
  $mc3_answer
);
$mc3->extra(
  $mc3_distractor1, 
  $mc3_distractor2,
  $mc3_distractor3,
  $mc3_distractor4,
);

$mc4_answer = "decrease the Type I error probability.";
$mc4_distractor1 = "not change the Type I error probability.";
$mc4_distractor2 = "not change the Type II error probability.";
$mc4_distractor3 = "increase the Type I error probability.";
$mc4_distractor4 = "either increase or decrease the Type I error probability.";

$mc4 = new_checkbox_multiple_choice();
$mc4 -> qa (
"Decreasing the significance level of the hypothesis test above would (select all that apply)", 
$mc4_answer
);
$mc4 -> extra(
$mc4_distractor1,
$mc4_distractor2,
$mc4_distractor3,
$mc4_distractor4,
);

$st = $xsd;

#######
# Problem Text

Context()->texStrings;

BEGIN_TEXT

Traffic police monitor the speed of vehicles as they travel over a new bridge. The average speed for a sample of \( $n \) vehicles was \( $xbar \) km/h, with the sample standard deviation being \( $st \) km/h. We will assume that the speeds are Normally distributed, and the police are interested in the mean speed.


$PAR
$BBOLD Part a) $EBOLD
Since the variance of the underlying Normal distribution is not known, inference here would involve the t distribution. How many degrees of freedom would the relevant t distribution have? \{ ans_rule(6) \}


$PAR
$BBOLD Part b) $EBOLD
Create a 95 $PERCENT confidence interval for the mean speed of vehicles crossing the bridge. Give the upper and lower bounds to your interval, each to 2 decimal places.   ( \{ $multians->ans_rule(6) \},\{ $multians->ans_rule(6)\}) 
$BR
	

$PAR
$BBOLD Part c) $EBOLD


\{ $mc3->print_q() \}
$BR
\{ $mc3->print_a() \}


$PAR
$BBOLD Part d) $EBOLD


\{ $mc4->print_q() \}
$BR
\{ $mc4->print_a() \}
END_TEXT

#######

$showPartialCorrectAnswers = 1;

#ANS(  $ans[$ver]->cmp(tol=>0.0005, tolType=>'absolute') );
ANS(  Real($df)->cmp(tol=>0.000, tolType=>'absolute') );
ANS( $multians->cmp() ); 
ANS(  radio_cmp($mc3->correct_ans()) );
ANS(  radio_cmp($mc4->correct_ans()) );

##############################################################
#
#  Solution
#

Context()->texStrings;

BEGIN_SOLUTION
$BR
$BBOLD Part a) $EBOLD
$BR
$BR
The answer is \( $n - 1 \) degrees of freedom.

$BR
$BR
$BBOLD Part b) $EBOLD
$BR
$BR
	The confidence interval is of the form \( \bar{x} \pm t_{n-1}(0.975) \frac{s}{\sqrt{n}} \) where the 97.5 $PERCENT point of \( t_{n-1} \) is denoted \( t_{n-1}(0.975). \)
$BR
Here, the interval is \( $xbar \pm $tcoshow \times \frac{$xsd}{\sqrt{$n}} = ($ashow, $bshow). \)
$BR
$BR
$BBOLD Part c) $EBOLD
$BR
$BR
We check to see if \( 80 \) is contained within the confidence interval above. If it is, then we do not reject the hypothesis; if it is not, then we reject the hypothesis. In this case, the correct answer is:
$BR $BR $mc3_answer

$BR
$BR
$BBOLD Part d) $EBOLD
$BR
$BR
Typically decreasing the significance level will increase the Type II error probability.

END_SOLUTION


ENDDOCUMENT();
