 #DESCRIPTION
##  A problem to compute probability associated with a continuous uniform random variable.
##ENDDESCRIPTION


## DBsubject(Statistics)
## DBchapter(Confidence intervals)
## DBsection(Independent samples - t)
## Date(11/15/2012)
## Institution(UBC)
## Author(Djun Kim)
## Level(4)
## MO(1)
## KEYWORDS('statistics', 'variance', 'ANOVA', '2-sample t')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

#######

TEXT(beginproblem());

Context("Numeric");


$ver = random(0, 3, 1);

# $ver=3;

#---
# Part a)

@bar1arr = (24.99, 23.69, 31.86, 25.84);
@bar2arr = (32.92, 23.92, 24.62, 24.64);
@ss1arr = (108.47, 89.60, 139.75, 52.69);
@ss2arr = (77.68, 66.35, 109.97, 105.63);

$var1 = sprintf("%.2f", $ss1arr[$ver]/14.);
$var2 = sprintf("%.2f", $ss2arr[$ver]/12.);

$data[0] = "Given: $BR$BCENTER \(".
"n_1 = 15, \bar{x} = 24.99, \sum{(x_i - \bar{x})^2} = 108.47".
"\),$ECENTER$BR and $BR$BCENTER \(".
"n_2 = 13, \bar{y} = 32.92, \sum{ (y_i - \bar{y})^2} = 77.68".
"\).$ECENTER";

$data[1] = "Given: $BR$BCENTER \(".
"n_1 = 15, \bar{x} = 23.69, \sum{(x_i - \bar{x})^2} = 89.60".
"\), $ECENTER$BR and $BR$BCENTER \(".
"n_2 = 13, \bar{y} = 23.92, \sum{ (y_i - \bar{y})^2} = 66.35".
"\).$ECENTER";

$data[2] = "Given: $BR$BCENTER \(".
"n_1 = 15, \bar{x} = 31.86, \sum{(x_i - \bar{x})^2} = 139.75".
"\), $ECENTER$BR and $BR$BCENTER \(".
"n_2 = 13, \bar{y} =  24.62, \sum{ (y_i - \bar{y})^2} =  109.97".
"\).$ECENTER";

$data[3] = "Given: $BR$BCENTER \(".
"n_1 = 15, \bar{x} = 25.84, \sum{(x_i - \bar{x})^2} = 52.69".
"\), $ECENTER$BR and $BR$BCENTER \(".
"n_2 = 13, \bar{y} = 24.64, \sum{ (y_i - \bar{y})^2} =  105.63".
"\).$ECENTER";


$ans[0] = Real((14*$var1 + 12*$var2) /(15+13-2));
$ans[1] = Real((14*$var1  + 12*$var2) /(15+13-2));
$ans[2] = Real((14*$var1 + 12*$var2)/(15+13-2));
$ans[3] = Real((14*$var1  + 12*$var2)/(15+13-2));

$thediff = sprintf("%.2f", $bar1arr[$ver] - $bar2arr[$ver]);

$dispans = sprintf("%.3f", $ans[$ver]);

#---
# Part b)


$a[0] = Real((24.99-32.92) - 2.056 * sqrt($ans[0] * (1/13 + 1/15)) );
$b[0] = Real((24.99-32.92) + 2.056 * sqrt($ans[0] * (1/13 + 1/15)));

$a[1] = Real((23.69-23.92) - 2.056 * sqrt($ans[1] * (1/13 + 1/15)));
$b[1] = Real((23.69-23.92) + 2.056 * sqrt($ans[1] * (1/13 + 1/15)));

$a[2] = Real((31.86-24.62) - 2.056 * sqrt($ans[2] * (1/13 + 1/15)));
$b[2] = Real((31.86-24.62) + 2.056 * sqrt($ans[2] * (1/13 + 1/15)));

$a[3] = Real((25.84-24.64) - 2.056 * sqrt($ans[3] * (1/13 + 1/15)));
$b[3] = Real((25.84-24.64) + 2.056 * sqrt($ans[3] * (1/13 + 1/15)));

$ashow = sprintf("%.2f", $a[$ver]);
$bshow = sprintf("%.2f", $b[$ver]);

$multians = MultiAnswer($a[$ver], $b[$ver])->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $a_stu, $b_stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( (abs($a_stu - $f1) < 0.005 && abs($b_stu - $f2) < 0.005 ))
      {
         return [1,1];
      }
      else {
         return [0,0];  
      }

  }
);

#---
# Part c)

if ($a[$ver] <= 0 && $b[$ver] >= 0) { # 0 is in the interval
   $mc3_answer      = "Do not reject \( H_0 \) since 0 is within the interval found in part (b).";
   $mc3_distractor4 = "Reject \( H_0 \) since 0 is not within the interval found in part (b).";
}
else {
   $mc3_answer      = "Reject \( H_0 \) since 0 is not within the interval found in part (b).";
   $mc3_distractor4 = "Do not reject \( H_0 \) since 0 is within the interval found in part (b).";

}


$mc3_distractor1 = "Reject \( H_0 \) since 0 is in the interval found in part (b).";
$mc3_distractor2 = "Do not reject \( H_0 \) since 0 is not in the interval found in part (b).";
$mc3_distractor3 = "Do not reject \( H_0 \) since \( $thediff \) is within the interval found in part (b).";

$mc3 = new_multiple_choice();
$mc3->qa(
  "Based on the 95$PERCENT confidence interval constructed in the previous part, which of the following conclusions can be drawn when we test \( H_0 :  \mu_A = \mu_B\)  vs. \( H_a:  \mu_A \ne \mu_B \) with \( \alpha = 0.05 \). ",
  $mc3_answer
);
$mc3->extra(
  $mc3_distractor1, 
  $mc3_distractor2,
  $mc3_distractor3,
  $mc3_distractor4,
);

# $data[$ver] removed

#######
# Problem Text

Context()->texStrings;

BEGIN_TEXT
$PAR 


Some car tires can develop what is known as "heel and toe" wear if not rotated after a certain mileage. To assess this issue, a consumer group investigated the tire wear on two brands of tire, A and B, say. Fifteen cars were fitted with new brand A tires and thirteen with brand B tires, the cars assigned to brand at random. (Two cars initially assigned to brand B suffered serious tire faults other than heel and toe wear, and were excluded from the study.) The cars were driven in regular driving conditions, and the mileage at which heal and toe wear could be observed was recorded on each car. For the cars with brand A tires, the mean mileage observed was \( $bar1arr[$ver] \) (in \( 10^3 \) miles ) and the variance was \( $var1 \) (in \( 10^6 \) miles\(^2\)). For the cars with brand B, the corresponding statistics were \( $bar2arr[$ver] \) (in \( 10^3 \) miles) and \( $var2 \) (in \( 10^6 \) miles\(^2\) ) respectively. The mileage before heal and toe wear is detectable is assumed to be Normally distributed for both brands.



$PAR
$BBOLD Part a) $EBOLD
Calculate the pooled variance \( s^2 \) to 3 decimal places.  During intermediate steps to arrive at the answer, make sure you keep as many decimal places as possible so that you can achieve the precision required in this question. \{ ans_rule(6) \} \( \times 10^6 \) miles \(^2 \)


$PAR
$BBOLD Part b) $EBOLD
 Determine a 95$PERCENT confidence interval for \( \mu_A - \mu_B \), the difference in the mean \( 10^3 \) mileages before heal and toe wear for the two brands of tire.  Leave your answer to 2 decimal places.  ( \{ $multians->ans_rule(6) \},\{ $multians->ans_rule(6)\}) 


$PAR
$BBOLD Part c) $EBOLD


\{ $mc3->print_q() \}
$BR
\{ $mc3->print_a() \}


END_TEXT

#######

$showPartialCorrectAnswers = 1;

ANS(  $ans[$ver]->cmp(tol=>0.005, tolType=>'absolute') );
ANS( $multians->cmp() ); 
ANS(  radio_cmp($mc3->correct_ans()) );



##############################################################
#
#  Solution
#

Context()->texStrings;

BEGIN_SOLUTION
$BR
$BBOLD Part a) $EBOLD
$BR
$BR
\( s^2 = \frac{(n_A-1)s_A^2 + (n_B-1)s_B^2}{n_A + n_B - 2} \)
$BR
\( = \frac{(15-1)$var1 + (13-1)$var2}{15+13-2} \)
$BR
\( = $dispans \) miles \(^2 \).

$BR
$BR
$BBOLD Part b) $EBOLD
$BR
$BR
The 95% confidence interval for \( \mu_A - \mu_B \) is:
$BR
\( (\bar{x}_A - \bar{x}_B) \pm \underbrace{t_{0.025, df=15+13-2}}_{=2.056} \times \text{SE} \)
$BR
\( = ($bar1arr[$ver] - $bar2arr[$ver]) \pm 2.056 (s \sqrt{1/n_A + 1/n_B} ) \)
$BR
\( = ($ashow, $bshow). \)

$BR
$BR
$BBOLD Part c) $EBOLD
$BR
$BR
We check to see if \( 0 \) is contained within the confidence interval above. If it is, then we do not reject the null hypothesis; if it is not, then we reject the null hypothesis. In this case, the correct answer is:
$BR $BR $mc3_answer

END_SOLUTION


ENDDOCUMENT();
