# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 03/25/10


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Double integrals over rectangles)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(RiemannSumUsingContours)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('16.1')
## Section2('16.1')
## Section3('16.1')
## Problem1('3')
## Problem2('3')
## Problem3('3')
## KEYWORDS(''double integral', definite integrals', 'functions', 'multivariable')

## Textbook tags
## HHChapter1('Integrating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
$showPartialCorrectAnswers = 1;

$x0 = random(3,9,2);
$y0 = random(2,6,2);
$x1 = $x0 + 6;
$y1 = $y0 + 6;
$func = "sqrt((x-$x0)*(y-$y0))";

$gr = init_graph( $x0-.5, $y0-.5, $x0+6.5, $y0+6.5, size=>[300,300] );
$gr->lb( 'reset' );
## add a grid and labels
for ( my $x=$x0; $x<=$x0+6; $x+=2 ) {
    $gr->moveTo( $x, $y0-0.5 );
    $gr->lineTo( $x, $y0+6.5, 'gray' );
    $gr->lb( new Label( $x, $y0, $x, 'black', 'center', 'top' ) );
}
for ( my $y=$y0; $y<=$y0+6; $y+=2 ) {
    $gr->moveTo( $x0-0.5, $y );
    $gr->lineTo( $x0+6.5, $y, 'gray' );
    $gr->lb( new Label( $x0, $y, $y, 'black', 'right', 'middle' ) );
}

## then we graph the contours 1-5 and label them in some
##    appropriate manner
$z0   = random(1,5,1);
$clab = $z0;
$dz   = list_random(-1,1);
@fn = ();
for ( my $c=1; $c<=5; $c++ ) {
    $xmin = ($c*$c/6.5 + $x0);
    $f = "$c*$c/(x - $x0) + $y0 for x in <$xmin, $x1> using " .
        "color:blue and thickness:2";
    push( @fn, $f );

    ## add the label; try to put a white background in
    $gr->lb( new Label( $x0+$c+.2, $y0+$c+.2, $clab, 'black', 'center', 'middle' ) );
    $clab += $dz;
}
add_functions( $gr, @fn );

$z1 = $z0 + $dz;  $z2 = $z1 + $dz;  $z3 = $z2 + $dz;  $z4 = $z3 + $dz;
$sum1 = "$z1+$z1+$z2+$z1+$z3+$z3+$z2+$z3+$z4";
$sum2 = "$z0+$z0+$z0+$z0+$z1+$z2+$z0+$z2+$z3";

$ze1 = $z0 - $dz;
$ze2 = $z4 + $dz;
$sum1 = "$ze1 + $ze1 + $ze1 + $ze1 + $z1 + $z1 + $ze1 + $z1 + $z3";
$sum2 = "$z1 + $z2 + $z3 + $z2 + $z3 + $z4 + $z3 + $z4 + $ze2";

if ( $dz < 0 ) {
    $oversum = $sum1;
    $undersum = $sum2;
} else {
    $oversum = $sum2;
    $undersum = $sum1;
}
$overval = Compute( $oversum );
$underval = Compute( $undersum );
$over = Compute( "($oversum)*2*2" );
$under = Compute( "($undersum)*2*2");

## allow answers between this and the over/under.
$tol = ($over - $under)/2;
$overTest = Compute( "($oversum)*2*2 - $tol" )->with( tolType=>'absolute', tolerance=>($tol+1) );
$underTest = Compute( "($undersum)*2*2 + $tol" )->with( tolType=>'absolute', tolerance=>($tol+1) );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

The figure below shows contours of \( g(x,y) \) on the region
\( R \), with \( $x0 \le x\le $x1 \) and \( $y0 \le y\le $y1 \).
$PAR
$BCENTER
\{ image( insertGraph( $gr ), height=>300, width=>300, tex_size=>300,
	  extra_html_tags=>'alt="contour graph showing the contours ' .
	  'of the function."' ) \}
$ECENTER
$PAR
Using \( \Delta x = \Delta y =2 \), find an overestimate and an
underestimate for \( \int_R g(x,y)\, dA \).

$PAR
Overestimate = \{ $over->ans_rule( 55 ) \}
$BR
Underestimate = \{ $under->ans_rule( 55 ) \}

END_TEXT
Context()->normalStrings;

# ANS($over->cmp( tolerance=>.075 ) );
# ANS($under->cmp( tolerance=>.21 ) );
ANS( $over->cmp( checker=>sub {
    my ( $c, $s, $a ) = @_;
    return ( $c == $s || ( $c > $s && ( $c - $tol - 1) <= $s ) );
} ) );
ANS( $under->cmp( checker=>sub {
    my ( $c, $s, $a ) = @_;
    return ( $c == $s || ( $c < $s && ( $c + $tol - 1) >= $s ) );
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

There are nine squares. Using the largest value of \(g\) in each square,
assuming that the largest value may be one larger than the largest contour
value seen (if the contour doesn't seem to cross a corner of the square),
we have
\[
\hbox{Over } \approx ($oversum)\Delta x \Delta y = ($overval)(2)(2) = $over,
\]
and similarly, for an underestimate,
\[
\hbox{Under }\approx ($undersum)\Delta x \Delta y = ($underval)(2)(2) = $under.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
