#Problem APEX 3.2.18

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Mean value theorem)
## Institution(Valdosta State University)
## Author(S. V. Ault)
## Level(2)
## MO(1)
## TitleText1('APEX Calculus')
## AuthorText1('Hartman')
## EditionText1('3.0')
## Section1('3.2')
## Problem1('18')

DOCUMENT();

loadMacros(
    'PGstandard.pl',             'PGML.pl',
    'parserRadioMultiAnswer.pl', 'PGcourse.pl'
);

$a1 = random(1,  3);
$b1 = random(4,  6);
$c1 = random(1,  10);
$e1 = random(-6, -4);
$f1 = random(7,  10);

$coef1 = 6 * $a1 * $b1;
$coef2 = 3 * ($b1 - $a1);

$slope =
    ((2 * $f1**3 - $coef2 * $f1**2 - $coef1 * $f1) -
        (2 * $e1**3 - $coef2 * $e1**2 - $coef1 * $e1)) /
    ($f1 - $e1);

$p1 = $e1;
$p2 = $f1;

$f = "2 x^3 - $coef2 x^2 - $coef1 x + $c1";

$crit1 =
    Formula(
        ($b1 - $a1 -
            sqrt($b1 * $b1 + 2 * $a1 * $b1 + $a1 * $a1 + 4 * $slope / 6)) /
        2);
$crit2 =
    Formula(
        ($b1 - $a1 +
            sqrt($b1 * $b1 + 2 * $a1 * $b1 + $a1 * $a1 + 4 * $slope / 6)) /
        2);

$RMA = RadioMultiAnswer(
    [
        [
            'The Mean Value Theorem can be applied, and \(c=\) %s. '
                . '(Enter all values separated by commas)',
            List($crit1, $crit2)
        ],
        ["The Mean Value Theorem can not be applied."]
    ],
    0,
    checker => sub {
        my ($correct, $student, $self, $ans) = @_;
        return $student->[0] == $correct->[0]
            && ($correct->[1][0] == $student->[1][0]
                || List($crit2, $crit1) == $student->[1][0]) ? 1 : 0;
    },
    checkTypes => 0
);

BEGIN_PGML
A function [`f(x)`] and interval [`[a, b]`] are given.
Check if the Mean Value Theorem can be applied to [`f`] on [`[a, b]`].
If so, find all values [`c`] in [`[a, b]`] guaranteed by the Mean Value Theorem.

[``f(x) = [$f] \quad \text{on} \quad [[$p1], [$p2]]``]

[_]{$RMA}
END_PGML

$b2 = 2 * $coef2;
$y1 = (2 * $e1**3 - $coef2 * $e1**2 - $coef1 * $e1 + $c1);
$y2 = (2 * $f1**3 - $coef2 * $f1**2 - $coef1 * $f1 + $c1);
$c2 = -$coef1 - $slope;

BEGIN_PGML_SOLUTION
[`f(x) = [$f]`] is a polynomial, therefore it is continuous
and differentiable on every interval.  So the Mean Value
Theorem applies.

Set [`f'(c) = \dfrac{f(b) - f(a)}{b-a}`] and solve for [`c`].

[``
  \begin{aligned}
    6c^2 - [$b2] c - [$coef1] &= \dfrac{[$y2] - [$y1]}{[$p2] - [$p1]} = [$slope]\\
    6c^2 - [$b2] c + [$c2] &= 0 \\
    c &= [$crit1], [$crit2]
  \end{aligned}
``]

(The solutions can be found using the quadratic formula.)
END_PGML_SOLUTION

ENDDOCUMENT();
