#Problem APEX 3.2.20

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Mean value theorem)
## Institution(Valdosta State University)
## Author(S. V. Ault)
## Level(2)
## MO(1)
## TitleText1('APEX Calculus')
## AuthorText1('Hartman')
## EditionText1('3.0')
## Section1('3.2')
## Problem1('20')

DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
    'PGstandard.pl',             'PGML.pl',
    'parserRadioMultiAnswer.pl', 'PGcourse.pl'
);

$a = random(1, 5, 1);

$p1 = -1;
$p2 = 1;

$f = Compute("$a*sin^-1(x)")->reduce();

$crit1 = Compute("sqrt(1 - 4/(pi**2))");
$crit2 = -$crit1;
$crit  = Compute("sqrt(1 - 4/(pi^2)),-sqrt(1 - 4/(pi^2))");

$RMA = RadioMultiAnswer(
    [
        [
            'The Mean Value Theorem can be applied, and \(c=\) %s. '
                . '(Enter all values separated by commas)',
            $crit
        ],
        ["The Mean Value Theorem can not be applied."]
    ],
    0,
    checker => sub {
        my ($correct, $student, $self, $ans) = @_;
        return $student->[0] == $correct->[0]
            && ($correct->[1][0] == $student->[1][0]
                || List($crit2, $crit1) == $student->[1][0]) ? 1 : 0;
    },
    checkTypes => 0
);

BEGIN_PGML
A function [`f(x)`] and interval [`[a, b]`] are given.
Check if the Mean Value Theorem can be applied to [`f`] on [`[a, b]`].
If so, find all values [`c`] in [`[a, b]`] guaranteed by the Mean Value Theorem

[`` f(x) = [$f] \qquad \textrm{on} \; [[$p1], [$p2]] ``]

[_]{$RMA}{20}
END_PGML

BEGIN_PGML_SOLUTION
[`f(x) = [$f]`] is continuous on [`[-1, 1]`] and differentiable
on [` (-1, 1) `], so the Mean Value Theorem applies.

[``
  \begin{array}{rcl}
    f'(c) &=& \dfrac{f(b) - f(a)}{b - a} \\
    \\
    \dfrac{[$a]}{\sqrt{1 - c^2}} &=&
    \dfrac{[$a] \sin^{-1}(1) -  [$a] \sin^{-1}(-1)}
    {1 - (-1)}\\
    \\
    \dfrac{[$a]}{\sqrt{1 - c^2}}  &=& \dfrac{[$a] \pi}{2} \\
    \\
    \sqrt{1 - c^2}  &=& \dfrac{2}{\pi} \\
    \\
    c &=& \pm\sqrt{1 - \dfrac{4}{\pi^2}}
  \end{array}
``]

Both values of [`c`] are in the interval [`(-1, 1)`].
END_PGML_SOLUTION

ENDDOCUMENT();
