# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/20/17
# ENDDESCRIPTION
 
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## DBsubject(Algebra)
## DBchapter(Basic Algebra)
## DBsection(Algebraic Expressions)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('solve','linear','equation','application','subtract','divide')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextCurrency.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(p=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$original = random(50,300,10);
$pNum = random(5,30,5);
$p = $pNum/100;
$afterDiscount = $original*(1-$p);

$ansString = "p-$p*p=$afterDiscount";
$ans = Compute("$ansString");

Context("Currency");
$afterDiscountC = Currency($afterDiscount);


##############################################

TEXT(beginproblem());

BEGIN_PGML

The price of a washing machine after [`[$pNum]%`] discount is [`[$afterDiscountC]`]. What was the original price of the washing machine (before the discount was applied)?

Assume the washing machine's price before the discount is [`p`] dollars. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

Assume the washing machine's original price before the discount is [`p`] dollars. With the discount rate of [`[$pNum]%`], the discount part is [`[$p]p`] dollars. The washing machine's new price is its original price minus the discount, so we can write the equation:

    [``p-[$p]p=[$afterDiscount]``]

END_PGML_SOLUTION

ENDDOCUMENT();
