# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/20/17
# ENDDESCRIPTION
 
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## DBsubject(Algebra)
## DBchapter(Basic Algebra)
## DBsection(Algebraic Expressions)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('solve','linear','equation','application','subtract','divide')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(h=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$base1 = random(1,2,0.5)*12;
$base2 = random(1,2,0.5)*12;
$height = random(1,2,0.5)*12;
$volume = $base1*$base2*$height/1728;
$volumeInch = $volume*1728;

$ansString = "$base1*$base2*h=$volumeInch";
$ans = Compute("$ansString");


##############################################

TEXT(beginproblem());

BEGIN_PGML

A cat litter box has a rectangular base that is [`[$base1]`] inches by [`[$base2]`] inches. What will the height of the cat litter be if [`[$volume]`] cubic feet of cat litter is poured? (Hint: [`1 \text{ ft}^3 = 1728 \text{ in}^3`])

Assume [`h`] inches will be the height of the cat litter if [`[$volume]`] cubic feet of cat litter is poured. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

The volume of a rectangular prism can be calculated by the formula:

    [``\text{Volume}=\text{base1}\cdot\text{base2}\cdot\text{height}``]

In this problem, we must first convert all units to the same. It's easier to convert [`[$volume]`] cubic feet to cubic inches:

    [``[$volume]\text{ ft}^3\cdot\frac{1728\text{ in}^3}{1\text{ ft}^3}=[$volume]\cdot1728\text{ in}^3=[$volumeInch]\text{ in}^3``]

Next, we can write an equation based on the volume formula:

    [``[$base1]\cdot[$base2]h=[$volumeInch]``]

END_PGML_SOLUTION

ENDDOCUMENT();
