# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Modeling with a linear equation without solving it.
#
# Last updated: Yao 4/20/17
# ENDDESCRIPTION
 
## DBCCSS('6.EE.2.a','7.EE.4','A-CED.1','A-REI.3')
## DBsubject(Algebra)
## DBchapter(Basic Algebra)
## DBsection(Algebraic Expressions)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MO(1)
## KEYWORDS('solve','linear','equation','application','subtract','divide')
 
##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserImplicitPlane.pl",
  "contextCurrency.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("ImplicitPlane");
Context()->variables->are(p=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(showExtraParens=>0);

$original = random(800,1500,100);
$pNum = random(0.2,1,0.1);
$p = $pNum/100;
$new = $original*(1+$p);

$ansString = "p+$p*p=$new";
$ans = Compute("$ansString");

Context("Currency");
$newC = Currency($new);


##############################################

TEXT(beginproblem());

BEGIN_PGML

In May 2016, the median rent price for a one-bedroom apartment in a city was reported to be [`[$newC]`] per month. This was reported to be an increase of [`[$pNum]%`] over the previous month. Based on this reporting, what was the median price of a one-bedroom apartment in April 2016?

Assume the median price of a one-bedroom apartment in April 2016 was [`p`] dollars. Write an equation to model this scenario. There is no need to solve it.

    [_________________________]

END_PGML


##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ansString)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ansString)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));

##############################################

BEGIN_PGML_SOLUTION

Assume the median price of a one-bedroom apartment in April 2016 was [`p`] dollars. A [`[$pNum]%`] increase is [`[$p]p`] dollars. Since April's price plus the increase is May's price, we can write the equation:

    [``p+[$p]p=[$new]``]

END_PGML_SOLUTION

ENDDOCUMENT();
