# DESCRIPTION
# 
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter()
## DBsection()
## KEYWORDS()
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "contextLimitedPowers.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;

#
# Subclass the numeric functions
#
package my::Function::numeric;
our @ISA = ('Parser::Function::numeric');

#
# Override sqrt() to return a special value (usually 1) when evaluated
# effectively eliminating it from the product.
#
sub sqrt {
  my $self = shift;
  my $value = $self->context->flag("setSqrt");
  return $value+1 if $value && $_[0] == 1; # force sqrt(1) to be incorrect
  return $value if $value;
  return $self->SUPER::sqrt(@_);
}

#
# end of subclass
#
package main;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
Context()->flags->set(reduceConstantFunctions=>0);
Context()->variables->are(x=>"Real", y=>"Real");
LimitedPowers::OnlyPositiveIntegers();

$r1 = random(2,4,2);
if ($r1 == 2) {
  $a = list_random(3,4,5,6,7,8,9,10,11,12,13,14,15,20);
} else {
  $a = random(2,3);
}
$ar = ($a)**$r1;

$r2 = random(3,5,2);
if ($r2 == 3) {
  $b = random(2,5)*random(-1,1,2);
} else {
  $b = random(2,3)*random(-1,1,2);
}
$br = ($b)**$r2;

do {
  $p11 = random(2,13)*$r1;
  do {$p12 = random(2,13)*$r1} until ($p12 != $p11);
  $k11 = $p11/$r1;
  $k12 = $p12/$r1;

  $p21 = random(2,13)*$r2;
  do {$p22 = random(2,13)*$r2} until ($p22 != $p21); 
  $k21 = $p21/$r2;
  $k22 = $p22/$r2;
} until ($k21 != $k11 && $k22 != $k12);

if ($r1 == 2) {
  $dispa = "\(\sqrt{$ar x^{$p11} y^{$p12}} = \)";
} else {
  $dispa = "\(\sqrt[$r1]{$ar x^{$p11} y^{$p12}} = \)";
}

if ($r2 == 2) {
  $dispb = "\(\displaystyle \sqrt{$br x^{$p21} y^{$p22}} = \)";
} else {
  $dispb = "\(\displaystyle  \sqrt[$r2]{$br x^{$p21} y^{$p22}} = \)";
}

$ans1 = Formula("$a x^{$k11} y^{$k12}");
$ans2 = Formula("$b x^{$k21} y^{$k22}");


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Simplify. Assume that all variables represent positive real numbers.
$PAR
(a) $dispa \{ans_rule(20)\}
$PAR
(b) $dispb \{ans_rule(20)\}
$PAR
Write your answer using radical notation if necessary.
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS( $ans1-> cmp( checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 if $ans->{isPreview} || $correct != $student;
  #
  #  Get parsed formula for student and correct answers
  #
  $student = $ans->{student_formula};
  $correct = $correct->{original_formula} if defined $correct->{original_formula};
  #
  # check if equal when sqrt's are replaced by 1
  #
  Context()->flags->set(setSqrt => 1);
  delete $correct->{test_values};
  delete $student->{test_values};
  my $OK = ($correct == $student);
  Context()->flags->set(setSqrt => 0);
  #
  Value::Error("Check to see if your answer is simplified.") unless $OK;
  return $OK;
  }, formatStudentAnswer=>"reduced"
  )
);

ANS( $ans2-> cmp( checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 if $ans->{isPreview} || $correct != $student;
  #
  #  Get parsed formula for student and correct answers
  #
  $student = $ans->{student_formula};
  $correct = $correct->{original_formula} if defined $correct->{original_formula};
  #
  # check if equal when sqrt's are replaced by 1
  #
  Context()->flags->set(setSqrt => 1);
  delete $correct->{test_values};
  delete $student->{test_values};
  my $OK = ($correct == $student);
  Context()->flags->set(setSqrt => 0);
  #
  Value::Error("Check to see if your answer is simplified.") unless $OK;
  return $OK;
  }, formatStudentAnswer=>"reduced"
  )
);


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} (a) \($a x^{$k11} y^{$k12}\), $SPACE (b) \($b x^{$k21} y^{$k22}\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Help:${EBOLD} 
  \{ AnswerFormatHelp("exponents", "Click here for help writing exponents and square roots.") \}
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
