## DBsubject(Calculus - multivariable)
## DBchapter(Fundamental theorems)
## DBsection(Stokes' theorem)
## Institution(W.H.Freeman)
## Author(JustAsk - Kobi Fonarov)
## Level(2)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Rogawski')
## EditionText1('2')
## Section1('17.2')
## Problem1('13')
## KEYWORDS('calculus')
## RESOURCES('image_17_2_13.png')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "freemanMacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context('Vector');


$a=random(2,9);
$b=random(2,9);

$fy=Formula("$b*x^($b-1)*e^(x^($b))")->reduce();
$az=Formula("$a*e^y - e^(x^($b))")->reduce();


$curve="\mathcal{C}";
$domain="\mathcal{D}";
$surf="\mathcal{S}";
$FF="\mathbf{F}";
$curl="\text{curl}";
$ii="\mathbf{i}";
$jj="\mathbf{j}";
$kk="\mathbf{k}";
$BB="\mathbf{B}";
$AA="\mathbf{A}";
$rr="\mathbf{r}";
$cc="\mathbf{c}";
$GG="\mathbf{G}";

TEXT('<SCRIPT>jsMath.Extension.Require("AMSmath");</SCRIPT>')
       if $displayMode eq 'HTML_jsMath';

Context()->texStrings;

BEGIN_TEXT
\{ textbook_ref_exact("Rogawski ET 2e", "17.2","13") \}
$PAR
Let \(I\) be the flux of \($GG = \left<$a e^y,$fy,0\right>\)
through the upper hemisphere \($surf\) of the unit sphere. 
$BR$BR
$BBOLD (a) $EBOLD Find a vector field \(\mathbf{A}\) such that \($curl(\mathbf{A})=$GG\).
$BR
$BBOLD (b) $EBOLD Calculate the circulation of \(\mathbf{A}\) around \(\partial $surf\).
$BR
$BBOLD (c) $EBOLD Compute \(I\), the flux of \($GG\) through \($surf\).
$BR$BR
$BBOLD (a) $EBOLD \(\mathbf{A}=\) \{ans_rule(50)\}
$BR
$BBOLD (b) $EBOLD \(\int_{$curve} \mathbf{A}\cdot d\mathbf{s}=\) \{ans_rule()\} $BR
$BBOLD (c) $EBOLD \(I =\) \{ans_rule()\}
$PAR

END_TEXT 

Context("Vector");
ANS(Vector(Real(0),Real(0),Formula($az))->cmp(
  checker => sub {
    my ($correct,$student,$ansHash) = @_; 
	Value->Error("Your answer has the wrong number of components.") unless $student->length == 3;
	my $s1= $student . i;
	my $s2 = $student . j;
	my $s3 = $student . k;
	my $curl = Vector($s3->D('y')-$s2->D('z'),$s1->D('z')-$s3->D('x'),0);
	return $curl == Vector(Formula("$a e^y"), $fy, 0);
  }
));
#Context("Numeric");
ANS(Real(0)->cmp);
ANS(Real(0)->cmp);


Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$SOL $PAR
$BBOLD (a) $EBOLD We search for a vector field \(\mathbf{A}\) so that \($GG=$curl  (\mathbf{A})\). That is,
\[\left< \frac{\partial {\mathbf{A}}_3}{\partial y}-\frac{\partial {\mathbf{A}}_2}{\partial z},\frac{\partial {\mathbf{A}}_1}{\partial z}
-\frac{\partial {\mathbf{A}}_3}{\partial x},\frac{\partial {\mathbf{A}}_2}{\partial x}-\frac{\partial {\mathbf{A}}_1}{\partial y} \right> =
\]\[
 \left<  $a e^y ,$fy,0 \right> \]
We note that the third coordinate of this curl vector must be zero;
this can be satisfied if \(\mathbf A_1 = 0\)
and \(\mathbf A_2 = 0\). $BR
With this in mind, we let \(\mathbf{A}= \left< 0,0, $a e^y - e^{x^{$b}} \right>\). $BR
The vector field \(\mathbf{A}= \left< 0,0,$a e^y - e^{x^{$b}} \right>\) satisfies this equality. Indeed,
\[ \frac{\partial {\mathbf{A}}_3}{\partial y}-\frac{\partial {\mathbf{A}}_2}{\partial z}=$a e^y ,\quad 
   \frac{\partial {\mathbf{A}}_1}{\partial z}-\frac{\partial {\mathbf{A}}_3}{\partial x}=$fy,\]\[
   \frac{\partial {\mathbf{A}}_2}{\partial x}-\frac{\partial {\mathbf{A}}_1}{\partial y}=0\]
Note, the solution to this part is not unique.  One could add the gradient of
a scalar field to get another solution.

$BBOLD (b) $EBOLD 
The boundary \($curve\) is the circle \( x^2+ y^2=1\), parametrized by
\[\gamma (t)=(\cos t, \sin t,0),\quad 0\le t\le 2\pi \]
\{image("image_17_2_13.png", width=>208, height=>186)\} $BR
We compute the following values:
\[\mathbf{A}\left(\gamma (t)\right)=
  \left< 0,0,$a e^y - e^{x^{$b}} \right> \bigg|_{x=\cos t\text{,}y=\sin t}=
  \left< 0,0,$a e^{\sin t}- e^{\cos^{$b}t} \right>\]
\[{\gamma }'(t)=\left< - \sin t, \cos t,0 \right>\]
\[\mathbf{A}\left(\gamma (t)\right)\cdot {\gamma }'(t) =\]\[
  \left< 0,0, $a e^{\sin t}- e^{{ \cos}^{$b}t} \right> \cdot  \left< - \sin t, \cos t,0 \right> =
  0\]
Therefore,
\[ \int_{$curve} \mathbf{A}\cdot d\mathbf{s} =\int_0^{2\pi } 0 \,dt=0 \quad \mathbf{(1)} \]
$BBOLD (c) $EBOLD We found that \($GG=$curl  (\mathbf{A})\), where \(\mathbf{A}= \left< 0,0,$a e^y - e^{x^{$b}} \right>\). $PAR
We compute the flux of \(\mathbf{G}\) through \($surf\). By Stokes' Theorem,
\[
\iint_{$surf} $GG\cdot d\mathbf{S} =
\iint_{$surf} $curl  (\mathbf{A})\cdot d\mathbf{S} =
\int_{$curve} \mathbf{A}\cdot d\mathbf{s}
\]
Combining with \(\mathbf{(1)}\) we get
\[\iint_{$surf} $GG\cdot d\mathbf{S} =0\]




END_SOLUTION

ENDDOCUMENT();
