## DESCRIPTION
## Equations for lines in the plane
## ENDDESCRIPTION

## KEYWORDS('linear equations')

## DBsubject('Algebra') 
## BookTitle('Algebra: Form and Function')
## DBchapter('Linear Functions') 
## BookChapter('Linear functions, expressions, and equations')
## DBsection('Equations for Lines') 
## BookSection('Equations for lines in the plane')
## Date('01/01/10')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Algebra: Form and Function')
## EditionText1('1')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## Section1('5.4')
## Problem1('20')

##############################################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGcourse.pl",
"AnswerFormatHelp.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

##############################################
#  Setup

Context("Numeric");

($x,$y) = (random(-9,-2,1),random(-9,-2,1));
$m = random(2,9,1);
$b = random(2,12,1);

$fac1 = Formula("$y");
$fac2 = Formula("$m * (x - $x)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);



##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Write an equation for the line through the point
\( ($x,$y) \) and parallel to the line \( y = $m x - $b \) 
in point-slope form \( y = y_0 + m(x-x_0) \).
$BR
$BR
\( y = \) \{ $multians->ans_rule(10) \} + \{ $multians->ans_rule(20) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;

ANS( $multians->cmp() );

ENDDOCUMENT();  