## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/30/06

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Associated matrices)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(4)
## MO(1)

DOCUMENT() ;        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "PGgraphmacros.pl",
  "unionTables.pl",
  "PGmatrixmacros.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 1 ;

TEXT(beginproblem());

Context('Matrix');


# First we define a graph with x and y in the range -1 to 10, axes (strong lines) 
# defined at the point [0,0] and
# with 11 gridlines horizontally and 11 grid lines veritically.
# $graph is a graph object (or more appropriately, a pointer to a graph object).

# We  will define a function and it's first and second derivatives defined 
# on the domain [-1,10]

$graph = init_graph(-1,-1,10,10,'axes'=>[0,0],'grid'=>[11,11],pixels=>[300,300]);

@colors = ("blue", "red" ); 

@labels = ('A', 'B' );

$r1=random( 8, 9, .5 );
$r2 = random(5,6.5, .5 );

$t1=random( 3, 5, 1 );
$t2=random( $t1+2, 9, 1 );
$dt = $t2 - $t1;

$x1 = Compute("$r1 * cos($t1 * pi / 19)");
$y1 = Compute("$r1 * sin($t1 * pi / 19)");

$x2 = Compute("$r2 * cos($t2 * pi / 19)");
$y2 = Compute("$r2 * sin($t2 * pi / 19)");

$graph->moveTo(0,0);
$graph->arrowTo($x1,$y1,'red',2);

$graph->moveTo(0,0);
$graph->arrowTo($x2,$y2,'blue',2);

$graph->lb( new Label($x1,$y1,'A', 'red','center','bottom'));
$graph->lb( new Label($x2,$y2,'B', 'blue','center','bottom'));
		       
$A = Matrix([[$x1],[$y1]]);
$B = Matrix([[$x2],[$y2]]);

$A1 = "\left\lbrack\begin{array}{r} $r1 \cos($t1\pi/19) \\ $r1 \sin($t1\pi/19) \end{array}\right\rbrack";
$B1 = "\left\lbrack\begin{array}{r} $r2 \cos($t2\pi/19) \\ $r2 \sin($t2\pi/19) \end{array}\right\rbrack";

$answer = Matrix([
[$r2/$r1*cos($dt * pi / 19),-1*$r2/$r1*sin($dt * pi / 19)],
[$r2/$r1*sin($dt * pi / 19),$r2/$r1*cos($dt * pi / 19)]
]);

Context()->texStrings;
BEGIN_TEXT
\{
ColumnTable(
"The vector \(\vec{A}\) (red) has length \($r1\), and makes an angle of 
  \($t1\pi/19\) with the \(x\)-axis.
$BR$BR
The vector \(\vec{B}\) (blue) has length \($r2\), and makes an angle of 
\($t2\pi/19\) with the \(x\)-axis.  "
, 
image(insertGraph($graph),width=>300,height=>300,tex_size=>800)
,
separation=>20, indent=>0, valign=>"TOP"
);
\}
$BR$BR
Find the matrix which rotates and dilates vector \(A\) into vector \(B\):
$BR
$BR
\{mbox($answer->ans_array()," \( $A1 = $B1 \)")\}
END_TEXT
Context()->normalStrings;


ANS( $answer->cmp(
     checker => sub {
       my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
       return ($student * $A == $B);      # return 1 if correct, 0 otherwise
     }
) );

ENDDOCUMENT();
