##DESCRIPTION
## lcao tagged and PAID on 3-22-2004

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Inverses)
## Date(6/3/2002)
## Level(3)
## MO(1)
## TitleText1('Elementary Linear Algebra')
## AuthorText1('Larson, Edwards, Falvo')
## EditionText1('5')
## Section1('2.3')
## Problem1('38')
## KEYWORDS('matrix', 'inverse')
##
##ENDDESCRIPTION

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "PGmatrixmacros.pl",
  "PGnumericalmacros.pl",
  "PGmorematrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->are(t=>'Real');

$a[1] = random(2,6,1) * random(-1,1,2);
$a[2] = $a[1];
$a[3] = random(2,6,1) * random(-1,1,2);
$a[4] = $a[3]; 

$tag = random(1,4,1);
$a[$tag] = - $a[$tag];

$m = random(2,5,1);
$n = random(2,5,1);
if ($m == $n) {$n = $n+1;} 

$p = random(2,9,1);

# det = $a[1] * $a[4] * e^(($m+$n)*t); 

$a142 = $a[1]*$a[4]/$a[2];
$a143 = $a[1]*$a[4]/$a[3];

$ans11 = Compute("sin($p*t)/$a[1]/e^{$m t}");
$ans12 = Compute("- cos($p t)/$a142/e^{$m t}");
$ans21 = Compute("- cos($p t)/$a143/e^{$n t}");
$ans22 = Compute("sin($p*t)/$a[4]/e^{$n t}");


$multians = MultiAnswer($ans11, $ans12, $ans21, $ans22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);




BEGIN_TEXT

\{ 
mbox( 
'If \(A=\)', 
display_matrix([
["\($a[1] e^{$m t} \sin($p t)\)", "\($a[2] e^{$n t} \cos($p t)\)"], 
["\($a[3] e^{$m t} \cos($p t)\)", "\($a[4] e^{$n t} \sin($p t)\)"]
]),
', then ' 
) 
\}
$BR
\{ 
mbox(
  "\( A^{-1} = \)",
  display_matrix([
  [$multians->ans_rule(15),$multians->ans_rule(15)], 
  [$multians->ans_rule(15),$multians->ans_rule(15)]
  ], align=>'cc')
);
\}

END_TEXT

ANS( $multians->cmp );

;
ENDDOCUMENT();       # This should be the last executable line in the problem.

