##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

## KEYWORDS ('differential' 'equation' 'model' 'lr' 'circuit' 'series')
## Tagged by cmd6a 4/26/06

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Modeling with Differential Equations')
## Date('1/25/10')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations: with Boundary Value Problems')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('1.3')
## Problem1('')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add("t"=>"Real","q"=>"Real","q'"=>"Real","q''"=>"Real","q'''"=>"Real","E"=>"Real");
Parser::BOP::equality->Allow;
##############################################################
#
#  Setup
#
#
$L = Compute(random(1,3));
$C = Compute(random(5,10));
$R = Compute(random(4,9));
$diffeqn = Formula("$L*q''+$R*q'+q/$C-E");

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
A radio uses an LRC loop series circuit consisting of an inductor, resistor, and capacitor to tune radio stations. For this problem \(L = $L, R = $R, C = $C\). 
\{image("RLC_series_circuit.png",width=>140, height=>170);\}
$BR
Write down the differential equation modeling the circuit using E for the impressed voltage.
$BR
\{$diffeqn->ans_rule(25)\} \(=0\)
$BR
$BR
$BR
$BBOLD Note: $EBOLD Use the prime notation for your derivatives, not Leibniz. For example, the derivative of a function f would be written as \(f^\prime\). Use q for the dependent variable, t for the independent variable.
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($diffeqn->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0($L*q''+$R*q'+q/$C-E)");
                $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));

ENDDOCUMENT();        
