## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## KEYWORDS ('differential equation model cooling')
## Tagged by cmd6a 4/26/06

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Modeling with Differential Equations')
## Date('12/17/09')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations: with Boundary Value Problems')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('1.3')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
#########################################################
##
## Set up the context to allow ' in variable names
##
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("t"=>"Real","T"=>"Real","T'"=>"Real","T''"=>"Real","T'''"=>"Real");

$r = Compute(random(2,7));
$T = Compute(random(350,450,25));
$iv = Compute(random(20,120,5));
$DE = Formula("T'-$r($T-T)");

Context()->texStrings;
BEGIN_TEXT
A liquid at temperature \($iv\)F is placed in an oven at temperature \($T\). The temperature of the liquid increases at a rate \($r\) times the difference between the temperature of the liquid and that of the oven. Write a differential equation for the temperature T(t) of the liquid.
$BR
\{$DE->ans_rule()\} \(= 0\)
$BR
\(T(0) = \) \{$iv->ans_rule\}

$BR
Note: use T,T', etc instead of T(t), T'(t) in your answers. After you set up your differential equation you will have to set it equal to zero so that WeBWorK will understand your answer, do this in a way so that the highest order derivative has a positive coefficient.
END_TEXT
#ANS($DE->cmp(checker=>sub {
#my ( $correct, $student, $ansHash ) = @_;
# $other= Formula("T'+$r(T-$T)");
#return (($correct == $student)||($other==$student));
#} ) );
ANS($DE->cmp());

ANS($iv->cmp);

Context()->normalStrings;
ENDDOCUMENT();        # This should be the last executable line in the problem.
