## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## KEYWORDS ('differential equation model tank flow')
## Tagged by cmd6a 4/26/06

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Modeling with Differential Equations')
## Date('12/17/09')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations: with Boundary Value Problems')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('1.3')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

#########################################################
##
## Set up the context to allow ' in variable names
##
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("t"=>"Real","A"=>"Real","A'"=>"Real","A''"=>"Real","A'''"=>"Real");

$iv = Compute(random(5,20,1));
$rin = Compute(random(2,10,.5));
$rout = Compute(random(2,10,.5));
if ($rin == $rout) {$rin = $rin + 1.5}
$DE = Formula("A'-$rin+$rout")->reduce;
$A = Formula("($rin-$rout)*t+$iv");
Context()->texStrings;
BEGIN_TEXT
A tank contains \($iv\) gallons of water to start, \($rin\) gallons of water flow into the tank while \($rout\) gallons of water flow out of the tank per minute. Write a differential for the amount of water A(t) (in gallons) in the tank at time t in minutes.
$BR
\{$DE->ans_rule()\} \(= 0\)
$BR
\(A(0) = \) \{$iv->ans_rule\}
$BR
Solve the differential equation: \(A(t) = \) \{$A->ans_rule\}
$BR
Note: use A,A', etc instead of A(t), A'(t) in your answers. After you set up your differential equation you will have to set it equal to zero so that WeBWorK will understand your answer, do this in a way so that the highest order derivative has a positive coefficient.
END_TEXT
ANS($DE->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0(A'-$rin+$rout)");
                $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));
ANS($iv->cmp);
ANS($A->cmp);
Context()->normalStrings;
ENDDOCUMENT();        # This should be the last executable line in the problem.
