## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## KEYWORDS ('differential equation model exponential growth')
## Tagged by cmd6a 4/26/06

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Modeling with Differential Equations')
## Date('12/16/09')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations: with Boundary Value Problems')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('1.3')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
#########################################################
##
## Set up the context to allow ' in variable names
##
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("K"=>"Real","K'"=>"Real");

$iv = Compute(random(10000,100000,5000));
$DE = Formula("K'-0.1*K");

Context()->texStrings;
BEGIN_TEXT

$PAR
The knowledge that a civilization can amass in a year is proportional to the knowledge it currently posses. If K(t) is the knowledge that a certain civilization possesses at time t, this civilization can add 10% to its knowledge each year, and this civilization possesses \($iv\) volumes of knowledge at time 0, write an initial volume problem for the amount of knowledge the civilization possesses:
$BR
\{$DE->ans_rule()\} \(= 0\)
$BR
\(K(0) =\) \{$iv->ans_rule\}
$BR
Note: use K,K' instead of K(t), K'(t) in your answers. After you set up your differential equation you will have to set it equal to zero so that WeBWorK will understand your answer, do this in a way so that the highest order derivative has a positive coefficient.
END_TEXT
ANS($DE->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0(K'-0.1*K)");
                 $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));
ANS($iv->cmp);
Context()->normalStrings;
ENDDOCUMENT();        # This should be the last executable line in the problem.
