## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear combinations)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix')->variables->are('v1'=>'Real','v2'=>'Real');
Context()->variables->set(v1 => {TeX => "\vec{v}_1"}, v2 => {TeX => "\vec{v}_2"}, );
Context()->flags->set(
      formatStudentAnswer=>'parsed'
);

do {

$Mred = Matrix([
[1,0,random(-5,5,1)],
[0,1,random(-5,5,1)],
[0,0,0],
]);

$U = GL3Z();

$M = $U * $Mred;

$b = Matrix([
[$M->element(1,3)],
[$M->element(2,3)],
[$M->element(3,3)],
]);

$A1 = $M->column(1);
$A2 = $M->column(2);

$A = Matrix([
[$M->element(1,1),$M->element(1,2)],
[$M->element(2,1),$M->element(2,2)],
[$M->element(3,1),$M->element(3,2)],
]);

$x1 = $Mred->element(1,3);
$x2 = $Mred->element(2,3);

} until (($x1 != 0) and ($x2 != 0));

$answer = Compute("$x1 v1 + $x2 v2");

Context()->texStrings;
BEGIN_TEXT
Let \[ \vec{w} = $b, \ \vec{v}_1 = $A1, \ \vec{v}_2 = $A2. \]
If possible, express \( \vec{w} \) as a linear combination of the vectors \( \vec{v}_1 \) and \( \vec{v}_2 \).
Otherwise, enter ${BITALIC}DNE${EITALIC}.
For example, the answer \( \vec{w} = 4 \vec{v}_1 + 5 \vec{v}_2 \) would
be entered \( \verb!4v1 + 5v2! \).
$BR
$BR
\( \vec{w} = \) \{ ans_rule(30) \} 
END_TEXT
Context()->normalStrings;

ANS($answer->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  my $s1 = $student->eval(v1=>1,v2=>0);
  my $s2 = $student->eval(v1=>0,v2=>1);
  my $S = Matrix([$s1,$s2])->transpose;
  return $A * $S == $b;
}));

COMMENT('MathObject version.  Answer is always unique.');
ENDDOCUMENT();
