## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear combinations)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix')->variables->are('v1'=>'Real','v2'=>'Real','v3'=>'Real',);
Context()->variables->set(v1 => {TeX => "\vec{v}_1"}, v2 => {TeX => "\vec{v}_2"}, v3 => {TeX => "\vec{v}_3"}, );
Context()->flags->set(
      formatStudentAnswer=>'parsed'
);

do {

$Mred = Matrix([
[1,0,0,random(-5,5,1)],
[0,1,0,random(-5,5,1)],
[0,0,1,random(-5,5,1)],
]);

$U = GL3Z();

$M = $U * $Mred;

$b = Matrix([
[$M->element(1,4)],
[$M->element(2,4)],
[$M->element(3,4)],
]);

$bvec = Vector($b);

$A1 = $M->column(1);
$A2 = $M->column(2);
$A3 = $M->column(3);

$A1vec = Vector($A1);
$A2vec = Vector($A2);
$A3vec = Vector($A3);

$A = Matrix([
[$M->element(1,1),$M->element(1,2),$M->element(1,3)],
[$M->element(2,1),$M->element(2,2),$M->element(2,3)],
[$M->element(3,1),$M->element(3,2),$M->element(3,3)],
]);

$x1 = $Mred->element(1,4);
$x2 = $Mred->element(2,4);
$x3 = $Mred->element(3,4);

} until (($x1 != 0) and ($x2 != 0) and ($x3 != 0));

$answer = Compute("$x1 v1 + $x2 v2 + $x3 v3");

Context()->texStrings;
BEGIN_TEXT
Let \[ \vec{w} = $bvec, \ \vec{v}_1 = $A1vec, \ \vec{v}_2 = $A2vec, \ \vec{v}_3 = $A3vec. \]
If possible, express \( \vec{w} \) as a linear combination of the vectors \( \vec{v}_1, \vec{v}_2 \) and \( \vec{v}_3 \).
Otherwise, enter ${BITALIC}DNE${EITALIC}.
For example, the answer \( \vec{w} = 4 \vec{v}_1 + 5 \vec{v}_2 + 6 \vec{v}_3 \) would
be entered \( \verb!4v1 + 5v2 + 6v3! \).
$BR
$BR
\( \vec{w} = \) \{ ans_rule(30) \} 
END_TEXT
Context()->normalStrings;

ANS($answer->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  my $s1 = $student->eval(v1=>1,v2=>0,v3=>0);
  my $s2 = $student->eval(v1=>0,v2=>1,v3=>0);
  my $s3 = $student->eval(v1=>0,v2=>0,v3=>1);
  my $S = Matrix([$s1,$s2,$s3])->transpose;
  return $A * $S == $b;
}));

COMMENT('MathObject version.  Answer is always unique.');
ENDDOCUMENT();
