## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear combinations)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixCheckers.pl",
  "MatrixReduce.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Fraction');
Context()->parens->set("[" => {formMatrix => 1});

do {

# Ax = b

$A = Matrix([
[random(-5,5,1),random(-5,5,1),random(-5,5,1)],
[random(-5,5,1),random(-5,5,1),random(-5,5,1)],
[random(-5,5,1),random(-5,5,1),random(-5,5,1)],
]);

$A1 = $A->column(1);
$A2 = $A->column(2);
$A3 = $A->column(3);

$b = Matrix([
[non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1)],
]);

# M = [A | b]
$M = concatenate_columns_into_matrix( ($A1,$A2,$A3,$b) );
$M = apply_fraction_to_matrix_entries($M);
$Mred = rref($M);

$x1 = $Mred->element(1,4);
$x2 = $Mred->element(2,4);
$x3 = $Mred->element(3,4);

} until (($A->det != 0) and (round(Real($x1)) != Real($x1)));

$multians = MultiAnswer($x1,$x2,$x3)->with(
  singleResult => 1,
  checker => sub {
    my ($correct,$student,$ansHash) = @_;  # get the parameters
    my @c = @{$correct};      # extract the correct answers
    my @s = @{$student};      # extract the student answers
    return (($s[0] == $c[0]) and ($s[1] == $c[1]) and ($s[2] == $c[2]));
  },
);

Context()->texStrings;
BEGIN_TEXT
Write \( $b \) as a linear combination of the vectors \( $A1, $A2, $A3. \)
$BR
$BR
\( $b = \) 
\{ $multians->ans_rule(10) \} \( $A1 + \)
\{ $multians->ans_rule(10) \} \( $A2 + \)
\{ $multians->ans_rule(10) \} \( $A3.\)
END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS($multians->cmp);

COMMENT('MathObject version.  Unique answer has at least one weight that is not an integer');
ENDDOCUMENT();
