## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear combinations)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Numeric')->variables->are('x'=>'Real');
Context()->noreduce('(-x)-y');

do {

$A = Matrix([
[non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
]);

} until ($A->det != 0);

foreach my $i (1..3) {
  foreach my $j (1..3) {
    $a[$i][$j] = $A->element($i,$j);
  }
}

$p1 = Formula("$a[1][1] + $a[2][1] * x + $a[3][1] * x^2")->reduce;
$p2 = Formula("$a[1][2] + $a[2][2] * x + $a[3][2] * x^2")->reduce;

$q = Formula("$a[1][3] + $a[2][3] * x + $a[3][3] * x^2")->reduce;

$x1 = Compute("DNE");
$x2 = Compute("DNE");

$multians = MultiAnswer($x1, $x2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      return (($s[0] == $c[0]) and ($s[1] == $c[1]));
  }
);

Context()->texStrings;
BEGIN_TEXT
If possible, write \( $q \) as a linear combination of \( $p1 \) and \( $p2 \).
Otherwise, enter ${BITALIC}DNE${EITALIC} in all answer blanks.
$BR
$BR
\( $q = \) 
\{ $multians->ans_rule(5) \} \( ($p1) + \)
\{ $multians->ans_rule(5) \} \( ($p2). \)
END_TEXT
Context()->normalStrings;

ANS($multians->cmp);

COMMENT('MathObject version.  Answer is always DNE.');
ENDDOCUMENT();
