##DESCRIPTION
##Calculus: Limits and Derivatives
##ENDDESCRIPTION

##Tagged by YJ


## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Definition of the derivative)
## Date(5/26/2005)
## Institution(UVA)
## Author(Jeff Holt)
## MLT(defDerFormulaGivenRepresent)
## Level(4)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('5')
## Section1('2.8')
## Problem1('')
## TitleText2('Calculus: Early Transcendentals')
## AuthorText2('Stewart')
## EditionText2('6')
## Section2('2.7')
## Problem2('')
## KEYWORDS('calculus', 'limits', 'derivatives')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = non_zero_random(2,9,1);

Context('Numeric');
$func = Formula("$a^x");
$mp = MultiAnswer($func, 0)->with(
  singleResult=>1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      my $ff1 = Formula($f1);
      $ff1 = $ff1->substitute(x => "x+$f2");
      $ff1 = $ff1 - $f1->eval(x=> $f2);
      my $sf1 = Formula($f1stu);
      $sf1 = $sf1->substitute(x => "x+$f2stu");
      $sf1 = $sf1 - $f1stu->eval(x=> $f2stu);
      return [1,1] if($ff1 == $sf1);
      return 0;
    }
);

BEGIN_TEXT
\[
    \lim_{x\rightarrow 0}\frac{$a^x - 1}{x}
\]
represents the derivative of some function \(f(x)\) at some 
number \(a\).
Find \(f(x)\) and \(a\).
$BR
$BR
\(f(x)\) = \{ $mp->ans_rule(25) \}
$BR
$BR
\(a\) = \{ $mp->ans_rule(25) \}
END_TEXT

#@ans = (fun_cmp("$a^x", vars=>"x"),num_cmp(0));
#ANS(@ans);
ANS($mp->cmp);


ENDDOCUMENT();        # This should be the last executable line in the problem.
