##DESCRIPTION
##Calculus: Limits and Derivatives
##ENDDESCRIPTION

##Tagged by YJ


## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Definition of the derivative)
## Date(5/26/2005)
## Institution(UVA)
## Author(Jeff Holt)
## MLT(defDerFormulaGivenRepresent)
## Level(4)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('5')
## Section1('2.8')
## Problem1('')
## TitleText2('Calculus: Early Transcendentals')
## AuthorText2('Stewart')
## EditionText2('6')
## Section2('2.7')
## Problem2('')
## KEYWORDS('calculus', 'limits', 'derivatives')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Numeric');
Context()->variables->add(h=>'Real');

$b = random(2,9);
$bsqu = Compute("($b)^2");

$Funct = Formula("sqrt(x)");

$mp = MultiAnswer($Funct, Compute("$bsqu"))->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};

      #Correct given limit.
      $CorrDeriv = Formula("(sqrt($bsqu + h) - ($b))/ h");

      $f_b_h_stu = Formula($f1stu)->substitute(x=>Formula("$f2stu + h"));
      $f_b_stu = $f1stu->eval(x=>$f2stu);

      #Construct limit the student is claiming.
      $DerivStu = Formula("(($f_b_h_stu) - ($f_b_stu))/ h");
      
      #Check if the given function evaluated at the given point fits the
      #  desired limit. There technically could be many correct answers:
      #e.g. Limit ((4 + h)^3 - 64) / h.
      #
      # f(x) = x^3, a = 4 is correct.
      # f(x) = (x + 2)^3, a = 2 is also correct.
      # f(x) = x^3 + constant, a = 4 should also be correct.
      if ($CorrDeriv == $DerivStu) {
          return 1; } 
      else {
          return 0; }
      }
);

BEGIN_TEXT
The limit
\[
    \lim_{h\rightarrow 0}\frac{\sqrt{$bsqu+h} - $b}{h}
\]
represents the derivative of some function \(f(x)\) at some number \(a\).
Find \(f\) and \(a\).
$BR
$BR
\(f(x)\) = \{ $mp->ans_rule(25) \}
$BR
$BR
\(a\) = \{ $mp->ans_rule(25) \}
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();        # This should be the last executable line in the problem.
