## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Partial fractions)
## Institution(Rochester)
## MLT(ParFracMissingDenomMult)
## Level(2)
## MO(1)
## KEYWORDS('integral' 'partial fraction')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

#because I have to use "sort" later, I cant have an "$a" or "$b" var.
$a1 = non_zero_random(1,12,1);
$b1 = non_zero_random(1,12,1);
$c1 = non_zero_random(1,5,1);
$d1 = non_zero_random(6,12,1);

$poly1 = "x^2 +4 x + 4" ;
$poly2 = "x^2 +5 x + 6" ;
$poly3 = "x^2 +6 x + 9" ;
$rand = random(0,2,1);
@polys = ($poly1,$poly2,$poly3) ;
$randpoly= $polys[$rand] ;

# old
#we need to make sure that A is less than B.
#@1 = ($a1 ,$b1 - 2*$a1 );
#@sort1 = num_sort(@1) ;
#@2 = (3*$a1 - $b1 ,$b1 - 2*$a1 );
#@sort2 = num_sort(@2) ;
#@3 = ($a1 ,$b1 - 3*$a1 );
#@sort3 = num_sort(@3 );
#@ansA = ( $sort1[0], $sort2[0] ,$sort3[0] );
#@ansB = ( $sort1[1], $sort2[1] ,$sort3[1] );

@ansA = ($a1,3*$a1 - $b1,$a1);
@ansB = ($b1 - 2*$a1,$b1 - 2*$a1,$b1 - 3*$a1);

# These are the actual answers
$ans1 = $ansA[$rand];
$ans2 = $ansB[$rand];

  $multipart = MultiAnswer($ans1, $ans2)->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


BEGIN_TEXT

The form of the partial fraction decomposition for the integrand of $BR

\[ \int_{$c1}^{$d1} \frac{$a1 x + $b1 }{$randpoly} dx \] $BR

is

\[ \frac{A}{denominator} + \frac{B}{denominator} \] $BR

Find the numbers A and B. $BR

A = \{$multipart->ans_rule(20)\}
$SPACE 
B = \{$multipart->ans_rule(20)\}

END_TEXT

ANS($multipart->cmp() ); 

# old 
#A = \{ans_rule(20)\}
#B = \{ans_rule(20)\}
#END_TEXT

# old
#ANS(num_cmp($ansA[$rand]));
#ANS(num_cmp($ansB[$rand]));

ENDDOCUMENT();        # This should be the last executable line in the problem.
