## DESCRIPTION
## Elementary Algebra
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(07/24/2018)
## Institution(Red Rocks Community College)
## Author(Craig Faulhaber)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
 
);

##############################################

#Integer exponents two variables

Context("Numeric");
Context()->variables->add(y=>"Real");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);

$a = random(2,4);
$b = random(-8,-2);
$c = random(2,7);
$absb = abs($b);

$ab = $a - $b;
$cd = $c - $d;

$e = random(2,3);
$f = random(-3,-2);
$ef = $e**abs($f);
$efa = $ef*$a;
$cf = abs($c*$f);
$absf = abs($f);

do {$d = random(2,6)} until $d < $cf;

$cfd = $cf + $d;

$ans = Formula("y^$absb/($efa*x^{[$cfd]})");


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify. Use only positive exponents in your answer.

[` \dfrac{([$e]x^{[$c]})^{[$f]}}{[$a]x^{[$d]}y^{[$b]}} = `]  [________]

 
END_PGML

##############################################
ANS($ans -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));
##########################################

BEGIN_PGML_SOLUTION
*Laws of Exponents*

Power Rule:  [`(x^n)^m = x^{n \cdot m}`]

Product Rule:  [`x^n \cdot x^m = x^{n + m}`]

Quotient Rule:  [``\frac{x^n}{x^m} = x^{n-m}``]

[` \dfrac{([$e]x^{[$c]})^{[$f]}}{[$a]x^{[$d]}y^{[$b]}} = \dfrac{([$e])^{[$f]}x^{-[$cf]}}{[$a]x^{[$d]}y^{[$b]}}`]

We can write the expression using positive exponents as:

[``\frac{y^{[$absb]}}{[$a]\cdot [$e]^{[$absf]} x^{[$cf]} x^{[$d]}} = \frac{y^{[$absb]}}{[$a]\cdot [$ef] x^{[$cf]} x^{[$d]}}``]

Then use the product rule and add exponents:

[``[$ans]``]
END_PGML_SOLUTION

ENDDOCUMENT();
