## DESCRIPTION
## Elementary Algebra
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(07/24/2018)
## Institution(Red Rocks Community College)
## Author(Craig Faulhaber)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
 
);

##############################################

#Integer exponents, exponent of 1 in numerator, fractions to simplify

Context("Numeric");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);

$b = random(3,12);
$ab=$b-1;

$c = random(2,6);
$d = random(2,4);
$cd=$c*$d;

$ans = Formula("1/($d*x^$ab)");


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify. Use only positive exponents in your answer.

[` \dfrac{[$c]x}{[$cd]x^{[$b]}} = `]  [________]

 
END_PGML

##############################################
ANS($ans -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

##############################
BEGIN_PGML_SOLUTION
We can use the quotient rule (and subtract the exponents) or we can look for common factors to simplify.

[` \dfrac{[$c]x}{[$cd]x^{[$b]}} = \dfrac{[$c]x}{[$c] \cdot [$d] x \cdot x^{[$ab]}}`]

There is a common factor of [`[$c]x`].  The expression reduces to:

[``[$ans]``]

END_PGML_SOLUTION

ENDDOCUMENT();




