## DESCRIPTION
## Simplify a rational expression
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################

#Product rule for exponents, negative exponent and a positive power of 1 (student must enter in the denominator)

Context("RationalFunction");
Context()->variables->add(y =>'Real');

#common factors are: x^$nn, and y^$mn

$nd = non_zero_random(3, 12);
do{
$n = random(2, 8);
}until($n < $nd);
$nn = $nd - $n;

$md = non_zero_random(3, 12);
do{
$m = random(2, 8);
}until($m < $md);

$mn = $md - $m;


#Question
$g = Formula("(x^$nn*y^$mn)/(x^$nd*y^$md)");

#answer
$f = Formula("1/(x^$n*y^$m)")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify. Use only positive exponents in your answer.

[``[$g] = ``]  [________]{$cmp}

 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
The quotient rule is

[``\frac{x^n}{x^m} = x^{n-m}``]

To simplify, we start by subtracting the exponents:

[`x^{[$nn] - [$nd]} = x^{-[$n]}`]

[`y^{[$mn] - [$md]} = y^{-[$m]}`]

Then we use the definition of negative exponents,

[``x^{-n} = \frac{1}{x^n}``]

to rewrite the expression with a positive exponent.

[``x^{-[$n]}y^{-[$m]}=[$f]``]
END_PGML_SOLUTION

ENDDOCUMENT();




