# DESCRIPTION
# Sigma notation
# ENDDESCRIPTION


## DBsubject(Calculus - single variable)
## DBchapter(Integrals)
## DBsection(Sigma notation)
## Date(11/01/2012)
## Institution(Hope College)
## Author(Paul Pearson)
## KEYWORDS('calculus', 'sigma notation')

################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


################################
#  Setup

Context("ArbitraryString");

$n = random(5,8,1);

$sum_string = "1";
foreach my $i (2..$n) {
  $i = $i**2;
  $sum_string = $sum_string . "+$i";
}


$answer = String($sum_string);


################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Write out the following sum, one term for each value of \( k \).
Simplify each term as much as possible.  
For example, enter \( 1+4+27 \) instead of \( 1^1 + 2^2 + 3^3 \) or \( 32 \).  
$PAR
\( \displaystyle \sum_{k=1}^{$n} k^2 = \)
\{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;


#################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp( 
checker => sub {
  my ( $correct, $student, $ansHash ) = @_;
  $student =~ s/ //g; # compress all blank space
  $student =~ s/~~+~~-/~~-/g; # convert +- to -
  $student =~ s/~~-~~+/~~-/g; # convert -+ to -
  $student =~ s/~~+~~+/~~+/g; # convert ++ to +
  $student =~ s/~~-~~-/~~+/g; # convert -- to +
  return $correct == String($student);
}
));

COMMENT('MathObject version.');

ENDDOCUMENT();
