##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Trigonometry')
## DBchapter('Polar coordinates & vectors')
## DBsection('Polar and rectangular coordinates')
## Date('7/31/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "AppletObjects.pl",
   "PGML.pl",
   "PGgraphmacros.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Vector");
Context()->flags->set(tolType=>'absolute', tolerance=>0.05);

##############################################
#  Begin trig functions in degrees

Context()->functions->remove("cos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cos {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub cos {Parser::Function->call('cos',@_)} 
#  Add the new functions to the Context
Context()->functions->add( cos => {class => 'NewFunc', TeX => '\cos'}, );


Context()->functions->remove("sin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sin {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub sin {Parser::Function->call('sin',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sin => {class => 'NewFunc', TeX => '\sin'}, );


Context()->functions->remove("tan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub tan {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180)/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub tan {Parser::Function->call('tan',@_)}
#  Add the new functions to the Context
Context()->functions->add( tan => {class => 'NewFunc', TeX => '\tan'}, );


Context()->functions->remove("sec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sec {
  shift; my $x = shift;
  return 1.0/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub sec {Parser::Function->call('sec',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sec => {class => 'NewFunc', TeX => '\sec'}, );


Context()->functions->remove("csc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub csc {
  shift; my $x = shift;
  return 1.0/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub csc {Parser::Function->call('csc',@_)} 
#  Add the new functions to the Context
Context()->functions->add( csc => {class => 'NewFunc', TeX => '\csc'}, );


Context()->functions->remove("cot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cot {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180)/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub cot {Parser::Function->call('cot',@_)}
#  Add the new functions to the Context
Context()->functions->add( cot => {class => 'NewFunc', TeX => '\cot'}, );



#sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))}
#sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])}
#sub atan {CORE::atan2($_[1],1)}
#sub acot {CORE::atan2(1,$_[1])}
#sub asec {acos($_[0],1.0/$_[1])}
#sub acsc {asin($_[0],1.0/$_[1])}


Context()->functions->remove("acos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acos {Parser::Function->call('acos',@_)}
#  Add the new functions to the Context
Context()->functions->add( acos => {class => 'NewFunc', TeX => '\arccos'}, );


Context()->functions->remove("asin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub asin {Parser::Function->call('asin',@_)}
#  Add the new functions to the Context
Context()->functions->add( asin => {class => 'NewFunc', TeX => '\arcsin'}, );


Context()->functions->remove("atan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub atan {CORE::atan2($_[1],1)*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub atan {Parser::Function->call('atan',@_)}
#  Add the new functions to the Context
Context()->functions->add( atan => {class => 'NewFunc', TeX => '\arctan'}, );


Context()->functions->remove("asec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asec {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
#sub asec {acos($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub asec {Parser::Function->call('asec',@_)}
#  Add the new functions to the Context
Context()->functions->add( asec => {class => 'NewFunc', TeX => '\arcsec'}, );


Context()->functions->remove("acsc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acsc {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
#sub acsc {asin($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub acsc {Parser::Function->call('acsc',@_)}
#  Add the new functions to the Context
Context()->functions->add( acsc => {class => 'NewFunc', TeX => '\arccsc'}, );


Context()->functions->remove("acot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acot {CORE::atan2(1,$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acot {Parser::Function->call('acot',@_)}
#  Add the new functions to the Context
Context()->functions->add( acot => {class => 'NewFunc', TeX => '\arccot'}, );


#  End trig functions in degrees
###################################################

###################################################
#  Make a subclass of MultiAnswer
#

package myMultiAnswer;
our @ISA = ('parser::MultiAnswer');

sub new {
  my $self = shift;
  my $ma = $self->SUPER::new(@_);
  $ma->{part} = 1;
  $ma->{answerName} = 'answerBox';
  $ma->{answerNames}{0} = $ma->{answerName};
  $ma->{id} = $MultiAnswer::answerPrefix.$ma->{answerName};
  $ma->{singleResult} = 1;
  $ma->{namedRules}  = 1;
  return $ma;
}

sub ANS_NAME {
  my $self = shift;
  my $i = shift;
  $self->{answerNames}{$i} = ($i == 0 ? $self->{answerName} : $self->{id}.'_'.$i);
  return $self->{answerNames}{$i};
}

package main;

#  End subclass of MultiAnswer
###################################################

$mag = Real(1);
$theta = Real(45);

$multiAns = myMultiAnswer->new(List(1,2), $mag, $theta)->with(
  format => '%2$s, %3$s',
  tex_format => '{\lVert\vec{v}\rVert = %2$s \\ \theta = %3$s^{\circ}}',
  checkTypes=>0,
  allowBlankAnswers=>1,
  checker => sub {
      my ( $correct, $student, $self, $ans_hash ) = @_;
      my ( $ggbStu, $magStu, $thetaStu ) = @{$student};
      my ( $magGGB, $thetaGGB ) = split(/,/,$ggbStu);
#      $self->{ans}[0]->{preview_text_string} = "hidden";
#      $self->{ans}[0]->{preview_latex_string} = "hidden";
      return [0,0,0] if $self->{ans}[0]->{isPreview};
      Value::Error("Very clever, but you must place \(\vec{v}\) in a quadrant to receive credit.") if ( $thetaStu == 0 || $thetaStu == 90 || $thetaStu == 180 || $thetaStu == 270 );
      Value::Error("The magnitude you entered doesn't match the magnitude from the image.") if ($magStu != $magGGB);
      Value::Error("The angle you entered doesn't match the angle from the image.") if (abs($thetaGGB-$thetaStu) > 0.25);
      return [1,1,1];
  }
);

######################################
# Create  link to applet:
###################################
# You can name your applet (anything reasonable :-) )
# Adjust the height and width as desired
# Paste the geogebra parameters in at the bottom of the page just above the
# command end command
# so that they don't get in the way
###################################


$appletName = "ggbfoo";

$applet =  GeogebraWebApplet(
   code       => "geogebra.GeoGebraApplet",
   archive    => "",
   codebase                 => "",
   appletName              => $appletName,
   appletId                    => $appletName,
   submitActionAlias     => 'getAppletValues',  # default actionAlias
   initializeActionAlias   => 'configure',  # default actionAlias
   setStateAlias             => 'setXML',
   getStateAlias             => 'getXML',
   setConfigAlias           => '',
   getConfigAlias           => '',
   returnFieldName       => '',
   width                 => 613,              # may want to modify width
   height                => 525,              # may want to modify height
  # mayscript             => "true",
   debugMode             => 0,   # set debugMode to 0 for no debug
                                              # to 1 to make xml representation visible
                                              # to 2 to add alerts detailing progression
                                              # through the applet
   onInit                => 'ggbOnInit',
   type                  => 'geogebraweb',
   submitActionScript    => qq{ getQE('answerBox').value = getAppletValues() },
   #submitActionScript => '',
   selfLoading           => 1,
   params     => GEOGEBRA_PARAMS(),# paste parameters in

);


###############################################
#
#  PNG picture of GeoGebra applet for hardcopy

 $im = init_graph(-4,-4,4,4,axes=>[0,0],grid=>[8,8],size=>[400,400]);
 $im->moveTo(0,0);
 $im->arrowTo(2,2,'blue',2);


##################################
# Setup GeogebraWebApplet --
###################################

###############################################
#
#  For persistence of GeoGebra applet state 
#  when answers are submitted or page is refreshed.
#  Uses JavaScript interface to GeoGebra.

HEADER_TEXT(
qq! 
<script language="javascript">
   function ggbOnInit(param) {
       if (param == "$appletName") {
           applet_loaded(param,1);  // report that applet is ready. 
           ww_applet_list[param].safe_applet_initialize(2);
       }
   } 
</script>
<noscript>For this question to work properly, 
you must enable JavaScript, or use another browser 
(or computer) on which JavaScript works.</noscript>
!);


#######################################################
#
#  For setting initial values in GeoGebra from WeBWorK
#  Uses JavaScript interface to GeoGebra.

TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );
<script language="javascript">
function setAppletCoefficients() {
   var applet=getApplet("$appletName");   // alert("Updating coefficients");
   if (typeof(applet)!="undefined") {
     if (typeof(applet.setValue)!="undefined") {
     
     ////////////////////////////////////////////////////////////////
     //  List the values of the parameters to be set in the applet here
     
          // Initialize anything?
          applet.setFixed('text1',true,false);
          applet.setRounding(3);
     } else {
          setTimeout("setAppletCoefficients()", 1000);
     }
   } else {
      setTimeout("setAppletCoefficients()",1000);
   }
}
ww_applet_list["$appletName"].setConfig = function() {setAppletCoefficients()};
</script>
END_SCRIPT

HEADER_TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );

<script language="javascript">
function getAppletValues() {
    console.log("getAppletValues() entered");
	// get a handle for the applet object
	var applet = $appletName;
	var mag = applet.getValue("mag");
	var theta = applet.getValue("theta");

    return( mag + ',' + theta );	
}
 </script>
END_SCRIPT

#@completed = split(/:/,$main::inputs_ref->{_data});
$completed[1] = 0;
$completed[2] = 0;
$completed[3] = 0;
$completed[4] = 0;

###################################
#insert applet into body  -- this does not need to be changed
# except to adjust the insertion of the reinitialize_button or
# a hidden AnswerBox
###################################
 TEXT($PAR,  MODES(TeX=>$im,
                      HTML=>$applet->insertAll(
                                debug =>0,
                                reinitialize_button => 1,
                                includeAnswerBox=>1,
 )));

##############################################################
#
#  Text
#
#

BEGIN_PGML

[$PAR]*

Drag [`\vec{v}`] into one of the four quadrants, then give the magnitude and direction angle for [`\vec{v}`].

    [`\lVert \vec{v} \rVert = `] [@$multiAns->ans_rule(20)@]*

    [` \theta = `] [@$multiAns->ans_rule(20)@]*

* State the angle in degrees, from 0 to 360 degree.

You must complete this problem once for each quadrant:

    [`[@ ($completed[1]==1)?'\boxtimes':'\square'@]*`] Quadrant I

    [`[@ ($completed[2]==1)?'\boxtimes':'\square'@]*`] Quadrant II

    [`[@ ($completed[3]==1)?'\boxtimes':'\square'@]*`] Quadrant III

    [`[@ ($completed[4]==1)?'\boxtimes':'\square'@]*`] Quadrant IV

END_PGML

##############################################################
#
#  Answers
#
#

NAMED_ANS($multiAns->cmp);

##############################################################
#
#  Answers
#
#


##############################################################
#
#  Hints
#
#


##############################################################
#
#  Solutions
#
#




$needed = 4;

install_problem_grader(sub {
  my ($answers, $state) = (@_);
  my $result = {
    score => 0,
    errors => '',
    type => 'persistent grader',
    msg => '',
  };

  my $name = $main::PG->{flags}{ANSWER_ENTRY_ORDER}[0];
  my $ans = $answers->{$name};
  my $score = $answers->{$name}{score};
  $result->{score} = $score;
  my $data = $main::inputs_ref->{_data} // "0:0:0:0:0";
  my ($correct,@quadrants) = split(/:/,$data);
  my ($MAG,$THETA) = split(/,/,$ans->{original_student_ans});
  if ($main::inputs_ref->{submitAnswers}) {
    if ($score == 1) {
      # Check whether credit has been given for this quadrant already
      foreach my $Q (0..3) {
        if ($THETA > $Q*90 && $THETA < ($Q+1)*90 && $quadrants[$Q]==1) {
           $ans->{score} = $result->{score} = $score = 0;
           $ans->{ans_message} = "You have already received credit for a vector in this quadrant. <BR> Reposition the vector to another quadrant and try again.";
           last;
        }
      }
    }
    if ($score == 1) {
      # if we're still correct, record credit for the quadrant
      foreach my $Q (0..3) {
        if ($THETA > $Q*90 && $THETA < ($Q+1)*90) {
           $quadrants[$Q] = 1;
           last;
        }
      }
      $correct++ if $correct < $needed;
      $state->{num_of_correct_ans}++;
      my $need = $needed - $correct;
      $ans->{ans_message} = "Good job! <BR>".
        "You need to get the problem correct $need more time".($need == 1 ? "" : "s")." for full credit."
            if $need && $ans->{ans_message} eq '';
    } else {
      $state->{num_of_incorrect_ans}++;
    }
    $state->{recorded_score} = $state->{sub_recorded_score} = $correct/$needed;
    $data = join(':',$correct,@quadrants);
  }

  INSERT_RESPONSE($name,'_data',$data);
  $main::inputs_ref->{_data} = $data;
  push @{ $main::PG->{flags}{KEPT_EXTRA_ANSWERS} }, '_data';
  $result->{msg} = '<input type="hidden" name="_data" value="'.$data.'" />' .
    ($correct ?
      ($correct == $needed ?
        'You have completed the problem.' :
        'You have '.$correct.' out of '.$needed.' correct so far.') :
      'You must get this problem correct '.$needed.' times for full credit.');

  return ($result, $state);
});

sub GEOGEBRA_PARAMS {
    my $result = {
showResetIcon=>"true" ,
enableRightClick=>"false" ,
enableLabelDrags=>"false",
showMenuBar=>"false" ,
showToolBar=>"false",
showAlgebraInput=>"false",
enableShiftDragZoom=>"false",
useBrowserForJS=>"true",
ggbbase64=>"UEsDBBQACAgIAOa6/0wAAAAAAAAAAAAAAAAWAAAAZ2VvZ2VicmFfdGh1bWJuYWlsLnBuZ+17eTyU/7v3PWM04htK0WL92opQsstSlCRLG0m2b5Zs2ca+jKWSFkQqu0LIlmxZhyS7sSa7hMFgNGMbw8xzz4R0Xue8nnNez/md5/nj6Z9q7rk/y/W5rvf7fV2fa0L1dDT2MB5mBABgj+YF9SsAAAX/CSlm2AX+tZtOmQUA5L9qqp+55tk/l7Abce16WUZ+nc30vsXQh1XnXUfo7UOYpJTPFJ6hl5i7WmPS3+QA583J4H1U0B4U9Cgoo+EJL99fj8yFVKuBfdCzDLvvHQt9ZXURejKC9cRuvt0UwshyZi6j5sy1K8dE0qsqv17O6NZMWPtGoqCU4guONOSlZWRc8BULVEh/ixEUEIDNeTAKhHR3X5ODyYeo+ulbWb09vfzt9iTldUvgGgCXW9c/PIX0J05wJCG9ktsyp72OyNgz9OP2X2AeKHEMUYtwZfhsHR0drYlpS5h4HcItSelFhfwNn6HjHA0AIlZHWa8E3m1PPiPpuWBGrOUhm5Gz4ZCkz3A4DIDw/qypqQkUFxL60oM1gLw/T6efe0uzWvhnDQSiClzwdXKyTY47YR5kYWkZo5IF3OlwA9YUdrzLB4MmCAIw3+2pAgDWhe9AiN7+gIqK06c4X8XFYbv2Zrjaepv4+vq2+6eruly4dOkedcXj1kX2ffuYu+Fy1TWcT4B2wn06fmBxdRSAvwVYgX08nKwAVJk6Ki8goCzHC04iV725KFWAYdUpgA66Y8bN9UAgvKO1EGYGAKo6FkwHVa0OoQPHgQVQFxwQwAIHx2EAQlnAL0B4mMFxqM92mXS5Lh6z6nqrHfXt241jx4+rbw6vOqYGnH0dyMKlcG2yJYafbtdfEfND5fybK+H9yQe87q4O51YSaY2VVgMyWHa8w7F//wUEovBeSEhztmHBr52Y6QL79u1rafOKPeEn/07RGpxw11Tp6cjIyDaiNKz2WXw85/Pe613emJz5626f6+rox12hxITF+ByjksMvNkgrL8txVUZV63itvg+WWtGYQafaPWHjqn4BdJy+KulnkrUfnDWHJd9VFlVRMT0se/d4X1/fE0CPZxQIGY2FJuuxBpSRjm/4gJO23gQ3TfUqXmA/rPr7ngqHgeLOdF2e6W/5KU8OnnzT70ozP2hYPVYgRGohMnJfY0NDxIcPknkY2lmAm0/mBfgOTrECv85dj2rGLevSQ7aOh1UA6rQweuaDAFSVToNmH7Wju6tBE06sE9Avnz/fLym4/ST4Kiy5ouKWidfEAa+16fSs0+JlU3+Xjj0StWiMFLTQAk8pjLYZeDpDNZNEBa+Sz3L2YrtqT/IZOpvhitLjNwsf2ocC7dV7aUsahUEW9HPLZkSmh8pLSutjgn4/sYYcFRGJHBy8RbAyDZtp6xgxHHjDIrP7rctM97FDsnfjn0WWjPgPLC9Fv3jR1J15tfVO/erPH1l+l7JeqIUwN47WBLdSsPiyoEBr7ZmEkjKDntdxavRj1IOyddHJ0s8ZpR1Ii3PbeUBqb1u8fBfngaixxekuQSEhoTuR0SZJ0oLmXBymvvnZpqiqfgIjE1OaXW9uZ5mbXWNTUyS3kneDtuyQgEYYpzy2suxU1+1wJaE9R2SkUKYcZ772uJ2EVjMASsrUjTjpwHilpXtOo9axeagXdrml53Or1mwNrnRbxxg1XipR2PMA367G0sjlMciXhyKXSdl0fnbzc+gvzFw+jS33XezQmMRpvRSXRiIXauBKw7JeeItDp6ye3EJuLOWOOGgAtt7G1CCqTleuhjBj1AEpouNAcetf4aZONfRCNoaiG0P0vKyPKYs+T7gU9SfbEoym0R5z/YI8yPV8LK6urs5lrqg3Fts1vDbdlS6qhPscyuYyEW0Qu9L66mRtxVtSIXYErYg7ramlpYiNr/Ij9bMDir0HuLnvPigGg4ULSrNiWykwtbA45DHiQOY8cMBy34uL1Y1RwsY44thjiRZszx5Lu08+lI0Vy6n219yYfrEy67XkM8EtDg/WDFM0wnMX/coSz+BuOEkoJOVh1skrqDz8St6t0s7oFkxrnLVJV+rdSv2+XAeJysVIB8h7V+OXt1xCiSSllUH+5hdiUaes26jOzEn1zdrSI5D3+3rpWE6b7Xs+TDD9iRGhP8AW7xP68KGToNZczYCYose8FNOkeMXC2d54YgFaWZS0PFewoZjoi9d0zRUafPeIGWvUW2JWAyXE6vjdyLNxRMlAzTBlCLG6VtmR2bWEw+tuZ96IV6sBW5gXH3UugJ6eflITGfvqTUqKkn+x+9zNUyhjY8NK7yLdq/T5+fmdOUayvfIK1uT5clwTWoUy2+VmpDZXFnsTpXb9nMG4gTAgokLeHzCTq8VxbeYfdayHrEP/tQcsPPNLYTAqWNJwMmCEXjWqWHo9zw4ismwxJ7A0Pm6heeFCxrJSgVWrZNZCzrjHj6eyTc+PXb30KVuTOBrM0lJ4xtJ1df6c40SWDpDoObYbxRiQgO2sF1mZ6WmWXnt9FaZw02m0GrI9haoZzGx5tk9o9xAIWNx5WL9ZBfu5vM40LQ6DKz330l03epYGK6/vqzRRjD6Ha83TAa6xrgsCVbd6vzw9kuaKactFX7a5UeUrD90Cet4T0GQPyB1o2WKnljE/4OXl5dJr7NE42Rq3EXsHsbTxI1xlPtbl4zq+eaLHwDQaRM4BMtea3n1FXEWXgLPHLoVuU8OO62v6kT0DXEVTj3yxRQ0NF1AM4C5EREWjHFEb+v0ljri1DXNzc6u5z6WDXw9tIpoaA2PAi91sQq80yMV2vW3HYd3jjit8MkQi0eqdKpuPjk5Yjivi9UizBOr4S2P3GM0gEdjaApfjp9csi9Ptr4NPrwy64qfTk2IvwxLUKF1JZAG1B39l4r2z6T3LePWaY6XtWkaWSuHRSmL5SAo5F5YD/W6CyAkyyXL5wH7cu1yca3p1tgB9OlHCiTRXElcucY2xze3MpsewvoGaHah8quvYTEi7FC/XEd19Poxzuqs0Nnvp+2mQnlw3mI059Aw71DNZ5p2R/Ymk1yW9192U1hAqxj+SO19bdqZq9gy7TA2rJakgq+rRg+n7pb4PlIpjv6uuDQaCiA0fYoPu3QT5ex+B7tQpp7LhT+lM/bax3xApfJUJCy+M7mZKHkMRJ+OuCxkqZZczvCvjNsYmM7FcZ0Soq+JOZ9xAvH6jXIjxQJESsCsz/MEhe45YWZx7rHP7Bj9HyZB7XA62E74387in0fuHwC+254PR896QpftwdVFGpryWMtyUf7t+ZtEbcQFV3JmKYeGU01tbwlqTvIOK3h9ijYPkE5mohoiDXPaVENFfOzjRJMP2Lvz+XuA8MAUqA+ZLEL2K57UhLHHngqanp8OIOjSAowmERy/f1MqXWdVo/iYv1d01QcCW6gDpTsppi+L2/yYjWMBvZqHR4GG5X9DFCjD+prYu2CjIGxYXt98L0dsHDO0z8JjV0tTkM/qbNjaVkOXOAU3NzSBOmpw0GtR3u4XoOOdSd8CgIVJQq+EVA7CHE5JMhWPyvgDX0kMRjDuI0CxNcUDacfgWu9hN9S8IPfNTvDJ0/xA45KoDWwPeHTcwdotenB8acA6vTjvxi0Lkqr/HhwOwAsDCwgJmOSYiJhadeTU9tuPxzx/1A/3uBiOytt0ZF8Ppmdjjl8i/mW8mFups/Ntu4W9224LSaDLjuOH0DMUN+5XeGC9x/PjzXOPyi9ivOYKiolrei1N8skOIq0VGi5MTEzGZT96FdaRolA+Xm9fu8prNr081r30waV3G+kE89xhr7/1tky8czczMtDIlgeR1M7dy+Wbp2kyWUmQPSACtn1S0pWrc4dUP9/K3uPn9SWVZp62XZnquC/UR2hTRPXGJp70+PDl06nwuiuzdi54T6I+9dZU3l2n7lEcvSxFxg6Vdx6EVo0FMl846m07kqaHXYoehGbjd1zpwUu6zN2IUPe6mqag4j9275TURPdlv1fLK3GTgg+VjUfO3aWnXvROZ01kOSaov40Yk6sN5PLhkR3wLEDPd09krK3MDve5vcqaliYSyT7W1rX8f/o2K0ZYba0suuMqVWBVsk7+2kAgEZM3clUXcyEiFieG+d9lTpTgUpeovLoXuJ9yrNz8UFFiCEj73MnzltePKB5eJiwbQeFnHRC9tttdLscufRUVHfAlPn0dHe1VYJcjadKYWDFM6tST62ANiTaJU8rCRWuIn84Zc0Uof7Psk5V0n/9qG09HZf5Zm3HJvstVH6SpP/3RNaP7aO+KF4TJCSfz996foV4Vwx5Kpo/eGLZ2+3+eIdZWoWv3Ha7nP2nKkyi+6yjTkUw46n71hGvEAXEXE3IuSYe/mlnRszuGivQDS5Kp2nt3a4rRVd/VajNAh44iMkjU+YEt7n3403J+uth57YhD3s5e7IvqCfs5NS6xHKdpywMtE1d4+G4H92uqmEyWkPQ5qgOgkJLm0nzDsT5qzjZqjcrFRQG75vEIphUwS3SiVgI7m2eHr3eo+f35V9/CrT6yMgwgY63n432THdxRQwDi2ErBi9vr5kxS3m6pskcN73jsMyOShJOTlDUFx/s4PDcr1WylM3oOeP55KWjbtD58OmXk0TNY3/uA3Fipsv8F9gNUPCcZMDn21ae84Gtt5P67Se8nlafbU/sDhRM8dxPdaHPAmJUnfmjUx8SdahhMpvSY+2kKXXvGSjxqXIZwFT2hfVOg11dMzI607//hyPc5e+tDAMV2lpecxd9vpfYsOAEjjKFSe3Y8vTztLXa17SW5+dfiJ5m0QLgOm/CZiDy6yq2ZNBSNzSycPE8pyj/faBBk5VvakC8zlTDES63mQg9IbTsSJl81ZeVW2+p98RMvd51wEpSrvX36c619ikvUOM2uCePQD9JRoDW6v1KaV1cg7G87a5DJxY73VuZKRi6CQMcYxMTGdwkqKZElsRf5ZPaDQF98o/FY4l/CljFUPk4hc/Q6PWlwiUjyG3F6KZl43wBMIrcTYb64uLvZ+xeKLXR0PmsXLW6QPvGYp66vwci3CGD4eQvxMeCFt9/Wj+pd+utdl2I572TiQf/yjk8yH6F+HCHmuvMHuX1RwSxGwy3+TYPeuGXvlDIlGgWxUCtS6dq3b0OMQ76ZDMeyDyL2Mg9h0ePlr265JKa680c0sUJQF41bdYOyQvEuqboTUsRk6BMizk2tcZeLJZc+YjJOS8cnnHF9meWLN3HpTSlRnMtTcRlor4cYn+cpZiCogYxfnZYrnaXFcvonmmhquLC/t6TL6A11tjRd5dDG5bh0xCeNXDTumd21yIV4e0I9R8ik2VM3QvOYz4g0PTMB20D9FkYZwl3CVH9483u82ozZyW/CmXyruvXntk4wbVw9+zahzKBmyIsXO4MsWz6iL8PcoIR7u4K9uzQ6j1llmio9ofRhXwZIy1668R8EqeemSy3Rw5uf+G2tEhLabm7G7LPfZbSWpZ3sZ8egcT/qsQ0GSmcQOymt1WV0wLzgooaRUOYzk2AYmoM6q4z5XmUnN7v8iF9YzqNaFbHRPeIWGhkKI8aoJV7sIoO5jKxzZvaWfHjKo8qIqlLkkbqDYwCScvJ7d0nIJhNdcLxjCLdzk/HbeHRwcfNgTzKwt+gttJr/rleXt2lr2Ph4o69k96y0xp6xMH2RnH9/KZAWUYbzXDn9U5rwoW9S5p/hXZjp6HUDQcQwJnIesusBm4mMTOcA/hSPHIjiv3RtlHeS0JRn2PygoKGi9+xTU/FnL8e/fn6h3SVZedW9jCVVpqIgmZLDw3eI5CEnGMIV4310tK70awsx1KXhFn/qRkf/bWzKW9pEdWeebv4gdyWS5BKGlx931u3IJ5AQFt8j7TBz9CnAIw1xAkr+AsZC/yuVho7VE+3Ub/qvArxKFmt+6Tpmoha3tO6k77ao3C+/wiRkVS2mcP0/d8MrRuvTHlub4s1GrSwfLrbKfP3/Bu/VeBkwBxls4E8y8Isx1IfKyqEZpJIo7M+SqDcdn+Fa5A6fODA/jimh0jlLah+I2FfbVLudV3q6mpDs3SLTI1m7UpoQkRa9rZB/kKxV2/eN1KWZC1GRS9GoOv7Pp33NBq9vlloSG14UV16wcYgUdW3xq+YlYREDu7mRjVp6tuk0VnfydyPDB/cfCL3mrm4R7DwkCpLD5lJ3fYJO/g0MuT/YJfZe4YMLrKoM+ifxqA+mn/wnZquMovlfmVHEezF7ovYfyHPjn1L7ks8mNphy7OsUfsmwVcUbMl86wbRg4m3oJIMF3a7KAqa7PLq6urXF+8lA3HW304/BwFlqhbM3XUw4mG6Lqr0/N0kffr0V4nNb9+6tGQ5zryPPSuMmm0CjhO+DujTAnoKpoDuYnQPuPrHTOxb0o01xK5+K6NHFgjtglPlq9f3NqVKJLu4WKXy0ian2pRqWfkE7bHRhHyYH8AHMuy3zWhsH8TRuLAPczvpuFpjBA7qPyHXyJPskOPXJ4+c83cljmlSgHFA08W6jTUMteNMXnZObnhCMaXCIXMsn9On1QVnLe4Bn0ITe7H89CFKnMQan1rgCaEIX3wFcSUeKrKX8MncfS31XiTR5OXNAop5XNaCFcDVk6w4O/PYQx7Yv/Y5CQnpAUooax7mqWbVuSO4RaWlMF7owyBATPBC+v19eFETwbWEB/0BOGVlup2gfRwalw+HL0o8wc0Q7f+9TvhaAEweY2Tz8957e/aPh1wvc8MGPu5zTyzdp87GyS7PJr4+rAY1vrW8pW/iiM5FrOa/yivHOXM5f92knr1vbx352TVEJZIDuteG0pgofmSb/NJ5gmyG20FrlK++zXgmuxwW1E3cdJk+Ci7nA4bQ8Ac/ZzMl0okXHIGZGhrgVCGxsU76ypSzVKRGTSstIf8ymNGswSo/4Yma5qKFGF5glbVuuGG885iq0SvhI2y6e0EUcDqJ6GdstKGhrXBQO3GhbQGvANBme23/V5UsKtIB9tLaocVUOAa7UsH5Sp+cUtsIBECc6DECfvcHQsoXfc0botCfzgFXRzTD00NpgUhDowYBDxa020iHHyrs7zzOLZXDjVOAG8LKsDzHej6DWZbNflo4cwumNqwFY5lLeBjkxOIaXVjbJHGrzCEDC/TfrrOYMyNjccF5m1y0Twqam9PRnnL3W4/UIbzqdqMalrJB0t7zWTKkp0BNFs55Rr+xbc1BLvJnlEvdZ6oF8mr3s9KaUR6TI4LpK0oEGtyW6iCMy5Gt2Cfq6JFSIuyrtz47R2PJM7OlrOQnJs+qa+c7035Fu+yakpb8xfPVRiFa/hd6PILvoemuFlJ/6KUdIP5VkelGzi5xBgq8Kr+t6lfRxlOcJOsDaeknWduNSL9tbFie1JtOE9qXi7wj9Mguosv+vLmIWCQuu8x6TCVB5N2Ub/T5G7c5OwDkhvjgBZSENj808bLVNwg0EmAOjT2k6joOfrTL63wO7rW/9hRcCUvBmO/EJ3WGvsq3r4yrrluumloLwZjyn0+XPhkujlkYjiQG0n8MXWABqZwMFwYAuZdJSRXSEaZLM4tiKzmlyz5ZHa5OHJ9WniJ2psbB6S3NRhgIscPsSGvcwzmM4gwr4ijxzqfLlzo0o/kj2AW4vr5ybPqMwvneFgmG/rN7ZBI9f7f+b2kEFmBLbK5jN88MqevCHSXNs7ljxf0sW0RuTYCps2XlchYqepk6tcVv9Z4nCZw/Z6d8P91kzrhlHMJCMl5Ap1OLnfE48nh24Md6zEEsmaWsKNmPykuBDgum94DyHPM2kzOGh1eL1w+ali1Kl/rOKLpsNJ8cixKnWQyLdNr8OySn+OXtJq6WC4PHLnA2bm1eutYOBtnRLMj07FZfzhcZOKg5CFfOWPaKJBCdXjNydSa5deiDll3ZQSrizy9LD0WcA3g0VhVsW99ZttyIamr9OmkwaVBYI6gB9Qfs6zD+nUShRePxW4w54BXXRVy+vw1sPjXbn+4KRyV4DOzs4WnQuM1S/59TDdixUfP35M3tAJHE/+GxYyOqk3gL6oqXnyqTF99YC8ubl52IQZ5qHuw5eyf49ELPz8jteeleVr5tsPomFFUi2Ye5vKnYCy7lNSXqGe96LvIYgZ41YEeyY3S7R4tBKxav4XfBkAmt7hnPGjSzplHYaY7jx3WMb+6MqQR9Lnz595gQORqB5CCYjp4Mlt3/OId8OqKzE4/yyeJQg1yaIdO2dTTSWGGh5rquW/9BWY+SKYhxPndH8VK2hhH+JMJZqleyo7LnPkHgYnGG9sXW1QX+NlFtZNQMzNYOb6yNS7m026dgLgWRiV094fUzkvvlZuF6cNAPJ8Qhc9r6hsc/ofg+ih5e+gkd2PLisLIqmTLUjsLMrTlOPYjx/By4r0kcF3GSCgjVkb0LY976Jvy3Eqe984X//kRQnKz0gm1SGZOtXoD/mt+6OlsGCSLUXMp0SPB8R7qmGpcK3DEfBohXiV5zcFzDtBjiob299d6uohf41XWQr8TXPM7ADXMsiWoJjKZPlNRx5lwB08RWrS1cvjZ7l1zjKtML4J/lzUzFLFc7wZfweNunhN/a+d44GT3xYmpFLG810wG7oNwiUb4eN3Kc9+HQWNihdcxYCBzvCmIGlKNhvzgH9LJGkkkbdEifDsif/4McOflF0N+Vw+aYWptzpRKmf8krf5arV0X8BGn3BanenkRZOaYyjno6sThU+eXnwRtiF8IW6Yp/mh5rPIIueE/bMqRy0F0y37FuwPUdpN+krH5Hau0E0Q8Fg0OWrgFpq9eKdb/mwbBWV6kfLsiK2fw8SEtUVHP0/aSMOSizVFE1+vzD+Bo7wS61NkE/RA68sbrxrlIbEVp2NWF3Q/+FyH7Bw2kBvSfm9orm5CEuHyhujqkOuAq4F6WV95c1reIg0xK3u5f5l4g4zvvBQaPmQWiSoH7siXoe79YRc+IEYe1ZBKLLdzc5Cvf2Qx8+aaFznFc1rJvgjh6V9ROLeRf3p1DBEj5RAtSsS9OmHcyc+udVknsZ4Ij2QUnHnqd8XE/QF1xOC831KFcbGIq1k0IgPXOU7cixXw/Xazi+BIGZQj5qWfb1DbSNdKksU0ti1hx82uVMjrnx40fVO+FNVyj4tC5pIXtuhMmcYblgx9IA7UdsN3KhzkE6Dloc5LXau3DekFkei1vraksMYHcyF2CAfNzFLJO75BixtOlfKFmLxasgImr5EyLcve9IY4Uu5N8c72HhfsNQbmLUdemI6Xn4xxmXnpMkMHajJqSgHKwIpi6Ak+2xL7vo0CYoE4nsQ38bAejeN5f5bb64diUq4YFnP2EAWLK9qHil7HTa0M2FLm6awaRkxCI/lSXpSzk73YgBV0pw1QeS5iCwLUAAaENGtAZNcE/us3GAGzSFefdqb82L64iOm1UvHV8uSNa4dO+sXl2IkeF4++b7cS40V+1DiyuE6JXi/C5EHfe82dDx931AAQl0eDSiQtl15xor6GODzGtzqEhFx5REWvX2w+ag+5UTQd4VUd/sJDSKiJUubxoF4EGQ2JDmpVdLU3jrt4b1hoKnGd9JfpbW+ywmxe1oY0Efuz7vNUAZjQYBOJIb4jdg/H0+fyFgvCUg/EGdGlL7dERJWmFeWyCPuLR1/2SGvMfQU9JibPz+NxIfpJjtdvvnPKBT6+pEyRn56mfFcI7x2O+3BdfHq1x7CEfN3DoVwSb1EtfMRYiaJI0pWqJSvtUi+LpJwKZo+5j4zerVo3AIiSsiuxuv0XycMzwULrLANzz3DC67PeYmuV9jHDJOHfPCVnChTPN0orbUzJoT974Agd+HZRUb+qYkssy0Us6lyjG4t2pFHiwYln6zn3kWVjqLfQURNGwkW7WzyTle4XHgzNePn8YDkV1jRKiTr7m4UxTNUcI7PPH8+/EhLWiUdUWrVSkAqkEouMWpJ7rgNahTDRx9J7MU45KrURGR49QGb/Thu/6EBAV511i0LAMeSG5iOmJt9cFo3OxxNEQtSlTx4S+UncJ/OitkVjzU9oO6kSl6aGO8edVCB+bOAJ45QCRCzPr6oRWzS9QDnFQ5bB5LEQp/GXUk+v+j4XJa74u76eHS5sFCC+g5n9wPJbEI4kaZ05/9aAmWCiYrHM+J77wKyW38O/RiQ2xC9KdoEuvdZF2tTN4JEw7IGsjmZZ2JBK6LmTUkrImcBUAlQFcyBu5ZqHQz8BOZ+CM87acKsZaenLauX3d2szpR9tns7PVTeNBr6V7mnveawr+qKF65xyb8inCkmv70ECMf2E2aE9Gu6XoiAvBXfsjAhNVlwvIicR8cRZhw++tZCCVd2kPiTBNMViLdKUbSzjqcuip3N07+QwoSQSqT4+g1yZMXCMIKuxN9fMGa2LAPNp7zDQa/o8Mpeao6v2FDdN9tWfNpGQULrtTUnBaQr+5NtMxYMJ9NUiRu5WF5vCKTqPypjN/PDmpxj1rmv2J4ghKWYgUJlXaJJZKLHr1i23xOTLT/Ms/Ehc3xADjKqUrR7Mh1snFV/nebsuXctQnYpaH3eeEgtYW5Olo2huu0JI+17AVEXDh6xbWX62X0XlawvRDo1cju+5Lj6fJQpph69O3Qz4XFpv3xNSQl5h9ionhhWY9ML/ygRmV8WpSlFO67doJa9xUrM0x5UW687JaTD9ieHO/1SUGV5yG8eSJernN9RcK+g/elS+gqAx0Psyf+mVgFjjs/UUYYceONdk0vq3PhTyBtZ4kl313+QqGfbfVKQNEolORdPg8CtpGJwDrLoXA6ze5EH0oJbxL5ERcz+924avTi4V/XCZIkpCwUx6Uw7kT9kiQTU0RyTFrmPPAojg5CfaqFewXGTuO+X4YxxRU25DAT4xlC98hxC1zPwMOxXt4KP5I+2YDdytRlMfED5wdvg8Jkh7CETbX6A2m16usDKjf22hv2dACZY82mBzyMnmQ6o6j/D24YGCxgTxzTQKVTSK+47cQzJ42UBXxNAuciQd0/S2IcXl7/TC/rUTmPt0Vt+Try/mAsbAhR06uTondd3zGplsbp+WtRGsCcx+hVcaL0Tp2Lytp7hR6D7889l8p5qdPxeykcRd3CQA5ro4ZbMH9fHIE4wBYcCNcnr96zwIMdW1MjZrJcZXv7trxulVOdYNa8lIxJHlEsWXSHGYKmZZWHk9qu/ow51a7gfmDtrUgGQQT1ixPDlSyhpQeUQnlgu6FwAJalPd8ahoYNZ9WMLdlAYo323nCSslwOw2YZo8lQPdbp4E3bzVom3uCwY8FruapJJxUw6oXi8L9O/triBVQOQ5rmR2RQFjygNTTZpJvwr8vOq6KZdoe+05kEDI+5iPDiLa2UU32LmYfP3iduU9Jfl5dHzuyRAwgicDVU+tzaZQ9RsvmAoLyTmZnUoWCfuYA791hFRI7v2izFRlfD1L28/We/hRCEkLUeIfGfM+5QtnGCYhYvWhfV+c0xOaCeIgHTTpFw3OWDJPKJFZNiz2TDB9gA9eOqO66Pt9083ULjAFMJOw4I7aQE+pEqCYsQOVbLw4vdDP9Xq/lCFtY6e9yQccl7rMkUjxScWX0Ai5lVeGcNafLpKgtAncOkl9qrVHyIMoHnyKEiVWgGB2AMAGqppWvlIH5dlWIlH95Tl/rjxyZdy0y2LeEbKKwU9EFliWPydfUd8h4j6qwKqViJT8s3gWdsjoGdXfknMlBfgoTJGpVJYsX7t+awEjBWl3A545CsQNQTR9C5m2fASWPPYeeb1kjT4YSG0l6y9LmAqn4xltW831RzYk37hGcvabqI+lbK3qtSAQ5zG5jtVHOVu3EuEpwCpBHig3ONDwCV3c+HO2S0C5n/4Xc7MqQJOVwBSFi6Qt5Uoe5sb/kKLeP/Ltl3/Co2LiPj8YSvvWu8D5B3E4ssO3kTl9VNrjr0SlGH5arsYB1ybLDiGiQmA0q927kNxlTvniYLoqRlx5N/DUL8gQQGioGtO9n8DeBmc9d8E54yiEtkY+1ik8o47mKjzfpWfOb0IFx33UY0heXXlGSu/FMNvdyCyfD8a9Hg+9hhTY3WYE4+BPFz3DSeic2wPCUbdUKMdk3nauWBw3GfNGz08XoNdK/V+QMFOYGWxB7JCwFnnRAxbgDgm5mfhDHLIHYaLqwdBVemYz67j78fsDWmbHys6+K4QXw/oQbevVxTOBd99IvVy+YEohvzk9vMjnRfogM+lZ6mxXOsEiWxZmnbcm0Yv2V7zZHKXtuOQZjkEjFblL59l70QMb1hO46XUe2dJhAmZS1gDPjbbrQPf0UpxtjhBxWMdncY7ZSuVGRhuXb0o78suW2LJ9AN1sbz+n9/LssXAVshw/AHowq4hSwqpDyESZEVno5xHJ1KPCaLf9mEmfCl3HIlcKscMv5sqVUaLhvEoJ80r6YsJxucvTPk0XyZX4quyP49lJWOtSjo3BprY0asZSlXRVX18oeNeeyLYERS19fXZabBylVygfgLBR7wOtnRtaWqRCkQujNedhhTt672q+c1Jvr0QNchc992zVaEOOhmyQVmq/97qA348YH70DuXDp0nObrrea9vbZ+9nZHyDGur4nlbpimmKl7W7VqrwWATTCOJ/VPeZQfBdox5UkbfdiHIPhB2QeFcPLl6ozOZ/XdGdedaiDCbu2nz7pPntjwAsKyRCGnkBM6YGZfdhfB0+olnnguDuh7SkaLPxq1HuPXOPyB8dDduT9LpfokLjKlVStl+L4NkXchYFih7SPHw9CFq4I6Sbal3PLRQd/uXfs6ttvj3brTCZ7/S0p2emeqAqk+mtDOJW8r4ErZe+EmkWfsm67ECWk3cLYrjeRvLtTE9r+5hzTMQkJWaT/ZIy1xVx/EUPAyzdqVe7YwC75b6G7dTARcXF2irq8wJI6wMjIKKCgcJN667tw5yij6s2bsfUjXV1d2oIl+84fdXF6BOw/cKBxuNInTTtOJrdXmbLuekjSggE6mpaamsoLoNhzVNz/CYqSv5N446PT67nTgQCs7+VIgXW4+rlzpWHG5e6aLi4fMgcVsl3RSiuvSnP1WAMeHdfPuTmxFtPmVew4YlJ8d+iwC+egT53OXmCDm9rSRMKhRpb6ZUd8b3PjqUjmftHdvTjn1agFc6wwtZFic+XZb5hSDsvetRfxyWBU1dEJQ0y1n4mMjDx75V78UaDMTDNEfCnUVP+QjP2rJ4dO7QIxrIL96W7QGAUH2OJNeBCrI3dptsyBtBPiZBxsyc8SEldF2LnWvspX6SQoxJjnmq3pXALBsADBqKrhgWBYQSPJrkUlH/kAVDR4nPcKgdXD1JHweLw9ZjcsOZEhIPfhfFd6OGL22+XzTw7ujQyVkJbu4dxzWOrE47qbxfY2YU+ePIkYMxM8VymqbtAia/9Nb74cpxwaGpq5ssd3ujOts8LL9dSnriwDCbbBvQEzzQamfpagFzwLzjtZ07h0hqCOG650wOzW2B0gLtBpru5XIKFBa7L9+fPnlwXk+k8Neb+1wufPn8NaPRgF4FzDXcjbYv1qhcxcyszSzGuUq+k6pfNeXyl2XJl5I40KGqhgss9Mu6WNjU0RZgoR2EV2sHNpq/Ij2ffSCbu8/yeoGK5Im5cyJ638rZVNZ6w6ClBLxTDBE9e/MABLEOZoDuNT+x9tEfevVsx/bRPsH922tGuHrRECtvtb9Eqgo+OtAUFLigmPQSkZAVK2/S51av0I1JY0CZhcD4fLnaD/b1/Qf7xzXuCP//y+l9y8QaDVnOIg+bT2oD54HOQyVd/QROte4K+kXxujCoSAHW08mxZI/lXgAZcFU/h9y/g/cAashTuvVX7dMm9VjTcvgv/lLdFh2yNQKzlV8a7jDXsNKzxPbnU3gQ8Yoncpu05e8lrs0Ei7nHKe5rm/Xud9ftjT2fmP797flXvEiw+gFgxff4FGvHhh5X72ClDFSWuAr6uro6c2wMMP8MnBZOGKGBvFHhzroUM2YQ8fMoBhfFJFa6u+irl3/778/IjMfiz1ejfcdfNqFp+wPNvHf/ny5Vx+efmYvYC3L2gL2h0rTf1tbhlcR8wuZes2ThATHkttwsQp6S1rUS2y1Xq6eRxyO/Tw/5Gj/geHUnMcWJPr+N/WPQ/Q2lCD5f9HveA/gwl/Bos/OwB6LBunHLUdDwV+AOr6zXq3ajgLs9N2nftfhgb/PgBsP/oVxDseqYnR77yNo4HFye3r32A62pXb/3W7/gnDWhAzMO+uCaTb7HgEseJ/9NcTNOPtNHk5xsLKqmmw1NWyNVb6ir7+g+14CgHunaXDdaSEXIqT0d7FzHkU1C3ndv64IjMwSd71fKyUzd+GhoZbCHLBlyHg6pGPd8cnTDq7utIMC6yumJhsAhLuv4odzhvHFRWN7qI23GltzSv/Cfzoq0iMiuLnUkCc2DNsZE4rb2ze+jO/grZ7UUAV0/jpPpNDQjDVgJvwAGdQ/febOLa6VlSTf/9EBLQtAP83P5n5b6ZKvgx61SShF6AgHZ5t/vMHOO2Y2lq+j/IwfMHNIlt62jrMmEFQIX4K3IGc/40o91/0p38nToOT/n+c/ivs+uej7iu/y4T2ZCiciuO0BHSzY3j7VvpfQYXbxtth8sF5eTCEJR36pUKfPo3+9Ons1szUbq/9h1fbX6tZtsXLS5pVA9PDlYo7e8Qu7lmf6kjpzDY8taNDrLpxzy1ZfOcVPT3ejo6Ofdzc4gVWrWdpKwtWglzQ0jqlyHcM2GLisbEgGhM/YuUMigtOIBQm3O0v1LRCJzoway9mGeS96O6+pqmj84LD1PfyfBXpLhHfLGG5Mj90nZxI7SPbGHjwD1VKrDsz04yop1vVlajkIyAi0uR2Z8xDgkgkdn50vt3Y1mYFC+Dm2cQYqqzKRTr2F2YaMgUAv+VHiOxCscPAxSRlv9Tb9WGS/3x+aIs+QzXsZr+c2mFPTc1QCwuLC4aGMb3bv+CiFmiCApl3iNl/SVAHmQIIjfz/J3QMTCyEwsL7EXf1EsN6JAD+0Tyno5571izofwFQSwcIN+0PUpYrAAAPOgAAUEsDBBQACAgIAOa6/0wAAAAAAAAAAAAAAAAWAAAAZ2VvZ2VicmFfamF2YXNjcmlwdC5qc0srzUsuyczPU0hPT/LP88zLLNHQVKiu5QIAUEsHCEXM3l0aAAAAGAAAAFBLAwQUAAgICADmuv9MAAAAAAAAAAAAAAAAFwAAAGdlb2dlYnJhX2RlZmF1bHRzMmQueG1s7Zrdc9soEMCfr38Fo6e7h9iSbNluJk4n7czNZSZNby6Zzr1iaS1zwaATKJby1xeBPuOPxo5TO23yELEYEPx2WWDR2Yd0TtE9xIJwNracjm0hYD4PCAvHViKnJyPrw/m7sxB4CJMYoymP51iOLS8vWdVTUqfvDvI8HEVjy6dYCOJbKKJY5lXGVmAhlApyyvg1noOIsA83/gzm+Ir7WOpWZlJGp93uYrHolO/r8DjshqHspEI1oPrKxNgqEqequValRU8Xd23b6f77+co0f0KYkJj5YCE1jgCmOKFSqCRQmAOTSGYRjK1pwvy8F9dfcWwhiidAxxZLKLVQUWdsDTzr/N1vZ2LGF4hP/gNf5ck4gaq8Frp5GfXzJ055jOKx5dp9CymaCs5E/8c0mmGV6gw9U5riDGJ0j2n+s87BieS+bkDnTjEVUJZVr/rMAzC/9IvyjMw1RSQkKAU4FhIRQKBTZoS21kamFVu1d9YtMCwBoUTIamBXWqhAOD17mYRpczMKW4Nw7McoTnLDOzAK1R5hcCMzCkjOiH/HQChj8xqV8sRfJAggnyqmTsQJkzfkoeiD18zVTensE2cz7IjTLOSsovd3KVfIh4b4ll1EXM0zIjOVHo5M17a1XsfraaV5zpL5Kj3qP6f/3nacgeMeWoWbAefsWoRNRo3YORBie417OPiUWM+zdJfVyD4l8X0TZq9vr/ARPwBG3eqB3cOSYWyyTuUuGqapJPT7NAb4o+l0d+M5HGqg+UMRdT3vQAbmrIbqcx4HAqVj6xpfWygrng/muYWLtXdAHEAETKlLtjg7O3EejDTn/DExj9eNub8vzF+ay5haI3ZahFzP7Bzy55sZt/hesn8gJK3NgtN7o7xnym0r7j97aXuNbHURQ1Hk/9U5j88jCulh/PTgzU+v99POdzhT7ieiPt8ZqWI7Osze7QXPdjhJCSU4zpbf9GLbutaB46p92HB/OsJPMN3nH63gf9ZyQET5H6LKbFaEgDCXKrY3pVyrY7dN36+uDp5Imr/5kkmIBegjoVga2h1AdKua+sJuY8xEHg9sn9W+f87cFJZzd9vtrF+N3SMOyy0dvB/DeP7W5MgMefuo3Hp6Mc42OYHd4rtHzO7ncAL3qk1eT/+vhVjHmt7W0l30tmI7j2MJgmC2WR8S0notvdVCI0p1lMpYPxifs/yiqgxjGqkaTv8oR3PIKKbag5EQmJmTAqHULq4NM9uMED2UOalT5GROkfNQJHQ7agAxSdFFWe+iLH7hlolemeiXCa+BcbfNoFZ4pCy9sRI8cij93XaDr+mq5JcxhB+wPikNhLVtXBqpEU82PmRKVAcZnqsK5o2EfcT+XRjzhAVL3dmPxzmCa9X12Fgyh7jhe69LuULnGXKqG0kZPSoH9BRvu2bSrR+5oCRQmOZEWfyJmgpznOopgSeC00TCjR8DsPqDBaOXBQnkLD/Ra8OckjSHYn6a8Zg8cCYrO0K5Xi+o/rahdVe1ajK6m6+eGxp5niIwC2ltXxdGqpVgQva60OM43mrdNDHaBcVBxx31nJHXs4fO8L03GjyRqjN6BtW9fZ/wgg58KzfsrnLDOPbrcGDP3rNxTDingOvz5sdSbtzpLE3TdY7s6Tp4wb2TPwP/bsLTllFtd+j+s8qoPw46zutuPcalonsMeXYbn1R1y8+2zr8BUEsHCBABtCLpBAAAWCYAAFBLAwQUAAgICADmuv9MAAAAAAAAAAAAAAAAFwAAAGdlb2dlYnJhX2RlZmF1bHRzM2QueG1s7ZfdbtsgFMev16dA3DcGx06bKm4VdReb1FaberNbYp8kbA64QJq4r7Z32DMNbEjdj2RalG2a1FyEwwGO8e9/OCSji/WiRPegNJciw7RHMAKRy4KLWYaXZnp8ii/Oj0YzkDOYKIamUi2YyXDqZm7W2V4viQfOx6oqw3nJtOY5RlXJjFuS4QIjtNb8TMgbtgBdsRxu8zks2JXMmWmizI2pzqJotVr1wvN6Us2i2cz01toGsHsVOsPeOLPhnixa9ZvpMSE0+nJ91YY/5kIbJnLAyL5HAVO2LI22JpSwAGGQqSvIsN2ogL59RskmUGb4U9N/j5FfkeG+jYvPj96N9FyukJx8hdx6jVrCZlHTidwcO3wpS6mQyvBwiJGFGVPbTnzLymrOMkx6aTu/ZDUodM9sENJ62NLIvAnReKes1BDm2oddywLakaT15lKqQqO1C2plqH374NtV2zZTp8zJ65/Woz4qF3Br6hKQmfP8mwCt3V48Hm984EUBLk/cmlHkEb6AmUvB8w7Mj8LYTLG8rMwoX6p76HKl6X5c4zRtwNL4pAFLOlgpaT80GRJKBzQ+FGYm+KJJV6QNVI4E0hVA0VgbVjab6uYEdeO9BjjdBfjdCPgMxL0FIpW2x4f4E1qToG7wrGnQnXrPgzeaOHbLiq/ROKwbh+njOBj9YCTBSDspA3ei3bt23xnmi6rkOTe784Ctue6kwdh1nxwpSvp7SU8a4ckL2cl/KvMzuhVTtkZayVwFbW2wb/3j+y+OnTtZOVMGNGeiA/7SDTwnP3gjvx1lJct6DoWS4vFG6LgeOfb9pbCP7L/Lnqb9hn5KX+BPfL1LhwOSDJKD3Sr7qrGd7N2SFU1i+1f9HPpdpnS/C4EkW7Lz5GBA/lQVJ9ureDsUCnUdjIf4QAUejQfBOAnGaTCGOy6BTpnaKbleqqn9ufdaZfJDT9VP/pX6j4H/Sm2i+9UmAWaD4sbZXXbpWzXaXY2izl+AKPzNOP8JUEsHCJizqULoAgAACA0AAFBLAwQUAAgICADmuv9MAAAAAAAAAAAAAAAADAAAAGdlb2dlYnJhLnhtbN0ZbXPbtvlz+itwXK4Xb7INgK9KpfSU9MN657S9Ouv11vYDSEISaopUCUiWsu5u/3B/ac8DgBIlxY4Tp9vdnNAAgQfP+xvo0ZebRUXWstWqqccBu6ABkXXRlKqejYOVmZ5nwZcvPhvNZDOTeSvItGkXwoyDGCF35+DtIuIJronlchwUldBaFQFZVsLgkXFQBkSV40BEcSrSMj8v4iw+j6IiP8+GOT1nWRaleUzTMJMBIRutntfNN2Ih9VIU8rqYy4W4agphLL25Mcvnl5e3t7cXHWcXTTu7nM3yi40GUiBVrceBnzwHdAeHbkMLzilllz++vnLoz1WtjagLoI8Sr9SLz56MblVdNrfkVpVmPg4SzgMyl2o2BxUkcRKQSwRagh6WsjBqLTUc7b1amc1iGVgwUeP+Ezcj1U6cgJRqrUrZjgN6EXIWR2FEszChjGdAo2mVrI2HZZ7mZYdttFby1qHFmaUY0WEK5lFa5ZUcB1NRaZBK1dMWNAoMtSt41WZbyVy03fueHzaw/wBEvZWIDezqFDEO0iQehFE4SCkdxDF13PRIxwxUZJqmspgp+Z0wElN4CBuSAUlSWOGExSSClQxWUhLiWswiEhIEYSGJIhgjXGYJ7sVwPqaEMVgmnBLOCWeEh/AaxyROSJziQQ6wydAio/AgNLADT4hrYQiPXQsjeDjOAFHs0AATcZjYWYzQgD/myL5dDDMSDYEQLsQpIyHwAO8pJYAxRPTMChFRgv8ZiRA9TwnPCOADuREz5fcYxb/vreIXjszSGSXuG4WBMfBJ4LHWOjJKdGgSsAAF2QY4MDcgu0nitqhbo6EbuBsiN8QOJnLHIwfqpKWRg4nCx4rZCRl+iJBZT0iGQoBRkHs7hAT5ZpZ/HCL/mrhX62qUUb+a4a8hvoBOksxOHilT+FEysR5VF6V3Ez2J4o5iwtjDKT7ORXdS8g+S8pHK7YiyuEc0hvyE/+1zQjL8IDnvVO0HUEwOgvAx2fkjiKf0nRnAjcyP96nkkzE1uuzq1cgzRPQcYb17G7nQyGI4JGlIEr4rIAmmeF9FUk7SmKRJr5YMsJok8b6gYDnJDgpKnPWqCpSUBBdTW6KAHtYEV2F41BWZgS8zv5+UGagK0b4wAIOIihECVcwmjK5CABd8VyN4jGWCQ0qB8sRJginqjnIBrVOj1U63c1ktd0axalT1cmUOVFcsym5qGoAWlW2MPHzZFDcvj5QthTbdHICgp9h3Lq7HOGhsnowqkcsKWsNr9ANC1qJCV7b4p01tSOcDSWDR2R5qJFdFpUol6h/A8F3D8s1qkcuW2GmDIlokeJzsmi3G9s0WVGkHUjRNW15vNfgJ2fxdtg1m1OiC9n5CiLSt2+IJvxj2f2BLFwI9PBoe7mCntb1rz9GW62tpDMividhI3alu1qqyP/9av2yqcmeuZaNq80oszaq1XTXkyBaFmtSzSlpNWhNDD1rc5M3m2uXPxOF6s11izXD089mrpmpaAvHH4xgA/Ji70cIgYzsoamGohfB2QqS7fTbkFsKOuRstFBjZseYFZZ2UtKOiNHHvBx5lHQQb2lWtzFX3YlRxsxcU4Z35Ow0eomSfCOXo8sjzRjeyrWXlvKgGO66alXZO7Exl+Vhp+Z0w80ldfi9nEH/fCcyABlA70D3HpSzUAg66da85gVb9G7DqVks5a2UnoQtIp1cfO0QvWylKPZfS7LTrfLwPZsXp2B/polVL9EOSQwa+kXtfK5UWkL/LnkQoqwamC0wmoDeDOvsBghpc4Jy8ahbLpobrhSZfKQ3Xta2E65NYmXnT2ouJMABury+3TXtj+XwjN4aIvFkjItEqpKgh55BWGriqrOXzz//E6Bfu12+rxnyxEDM3AYowB+2sSkmaKfl5LYt/rP/ZAzVzacQOuFStdIxbvfaOgL/KCr3Ey7hYiLoktS1WTrpgnyYFHQfPOGTYM6fgZmW65bXTrUeA6aqSC9AHMTby1g5V54VrixTdjTT5r7B35KU+r66x5XaU8l+PopHRLh6JqJZzgXc+b3fUfnuQVS3ab6dTLQ3ZjINzzFEwpL3d103p3aSfITWCc7yQb/34tk9oF95mDoEEF0mIG0gkTmbqJn9VZSld+WngEq4MYGJpduKi2oVhIVojNQSbB4CEh3aznD39+Uq0M9nZ7qkPT6fpE/O9Um1RyWPz0QGhZ6AztqsWPSsW91sRQgC/RngjFadG7HrKnhXvMN+HGm9vnk+meyiJsnaOqQnZUP/dZksdN+Rtt7IBXZ3bpS3zS2/9xCJaCIjXDZl08JMOasLhoJuFHukk6nBNYj9zzPxWH7gB9gtqqopjE8sNJDqNn4s6LU8gUDbQpzzbPFufDcgWfp91arA10zUlh5b0GzsUH2XJJLOmxCF3w8ca8+5Y2ynZcoxF3R2M+qsnKf6ukPhebO+Jh8lJPEzvj4cW0HUqmv5Po2GvwPM/MFuBZqot1Muva6zjrqKcVv4bKZfYb31bv2lFrfG75WHJv9s+O8RHViqcgaYnBnrpNu4306Gzv3yQoe4qPn+Ez9OLlKaMJmnGWJbEURpbA75r+ZNHxKQt3q3rZ8wHBkeVHeu9fEChWEIl2ym1fEzFZ5T3woTT6GNqvlMqY/fV/PtDhP9fVZVPFs53NSm2BX1Pj3Il65mZH3ng+sTdoNm93+Hq1UK2vd4ED9g7sqhW3sT8IuNZxFPGozRK2JB6m+vtIm8qVfSEIg9O08zflVk0pIwljL+/Xlu1dDXbyf9T+Qv5M2EZJZfk3/86e1fBPhbQYTkSMYo7w//XhYK7DPNCuVvH065JfU4+r+AS4u8ifyHPrsQb+eNP2K38cgbvbmNwvL893P98ZvzN5l3aQfLBIS8PTTfZabp5T01+z+eDDn73KeGhBWGqNrI8Di0NgtrAnFRwTev2e+GptNXX8Tn7HUmDy0x3Hx2hUrzGFBJ13xjToMsLPimKXDfVysjrAi7bdfd3QtuWDV1PceIKl/1rsf1Y5f+Q+OI/UEsHCOEEYyvGCAAAFB0AAFBLAQIUABQACAgIAOa6/0w37Q9SlisAAA86AAAWAAAAAAAAAAAAAAAAAAAAAABnZW9nZWJyYV90aHVtYm5haWwucG5nUEsBAhQAFAAICAgA5rr/TEXM3l0aAAAAGAAAABYAAAAAAAAAAAAAAAAA2isAAGdlb2dlYnJhX2phdmFzY3JpcHQuanNQSwECFAAUAAgICADmuv9MEAG0IukEAABYJgAAFwAAAAAAAAAAAAAAAAA4LAAAZ2VvZ2VicmFfZGVmYXVsdHMyZC54bWxQSwECFAAUAAgICADmuv9MmLOpQugCAAAIDQAAFwAAAAAAAAAAAAAAAABmMQAAZ2VvZ2VicmFfZGVmYXVsdHMzZC54bWxQSwECFAAUAAgICADmuv9M4QRjK8YIAAAUHQAADAAAAAAAAAAAAAAAAACTNAAAZ2VvZ2VicmEueG1sUEsFBgAAAAAFAAUATAEAAJM9AAAAAA=="
};
$result;
}

ENDDOCUMENT();        
