## DESCRIPTION
## Linear systems: Nonhomogeneous linear systems
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Nonhomogeneous systems)
## Date(01/31/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson and Veronika Furst)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.8')
## Problem1('10')
## KEYWORDS('Nonhomogeneous linear systems')


##################################
#  Initialization


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##################################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real"
);

# Define eigenvalues
$lambda1 = 0;
$c = non_zero_random(-3,3,1);
$lambda2 = 2*$c;

# Eigenvectors:  <1,-1>, <1,1>

$a[0][0] = $c;
$a[0][1] = $c;
$a[1][0] = $c;
$a[1][1] = $c;

$A11 = $c;  $A12 = $c;
$A21 = $c;  $A22 = $c;


# Correct complementary solution
$yc1top = Formula("1")->reduce;
$yc1bot = Formula("-1")->reduce;

$yc2top = Formula("e^($lambda2*t) * 1")->reduce;
$yc2bot = Formula("e^($lambda2*t) * 1")->reduce;


$multians1 = MultiAnswer($yc1top, $yc1bot, $yc2top, $yc2bot)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='0'>
    <tr>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td>&nbsp;&nbsp;</td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  
    "\left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack, \ \ 
     \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};

        my $flag1 = 0;
        my $flag2 = 0;

        Value::Error("Your solutions are not linearly independent") if
        $c[0]->typeMatch($s[0]) &&
        $c[1]->typeMatch($s[1]) &&
        $c[2]->typeMatch($s[2]) &&
        $c[3]->typeMatch($s[3]) &&
        ($s[0] * $s[3] == $s[1] * $s[2]);

        if ( $c[0]->typeMatch($s[0]) && 
             $c[1]->typeMatch($s[1]) && 
             ($s[0] != 0 || $s[1] != 0) ) 
        {
          if ( ($A11 * $s[0] + $A12 * $s[1] == $s[0]->D('t')) &&
               ($A21 * $s[0] + $A22 * $s[1] == $s[1]->D('t')) 
          ){
                $flag1=1;
          }
        }

        if ( $c[2]->typeMatch($s[2]) && 
             $c[3]->typeMatch($s[3]) && 
             ($s[2] != 0 || $s[3] != 0) ) 
        {
          if ( ($A11 * $s[2] + $A12 * $s[3] == $s[2]->D('t')) &&
               ($A21 * $s[2] + $A22 * $s[3] == $s[3]->D('t')) 
          ){
                $flag2=1;
          }
        }

        if ($flag1==1 && $flag2 ==1){
           return 1;
        } elsif ($flag1 == 1) {
           Value::Error("The vector on the left is correct");
        } elsif ($flag2 == 1) {
           Value::Error("The vector on the right is correct");
        }else{
           Value::Error("Neither of the vectors is correct");
        }
       
    }
);
# Particular solution

$g1 = Formula("e^($lambda2*t)")->reduce;

$partb1 = Formula("1");
$partb2a = Formula("1");
$partb2b = Formula("1/$c");

$partc1 = Formula("1");
$partc2 = Formula("1");
$partc3 = Formula("0");


#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

As an illustration of the difficulties that may arise in using the method of undetermined coefficients, consider
\[
\displaystyle
\ \boldsymbol{\vec{y}^{\, \prime}}
=
\left\lbrack
\begin{array}{rr}
$a[0][0] & $a[0][1] \\
$a[1][0] & $a[1][1]
\end{array} 
\right\rbrack
\boldsymbol{\vec{y}}
+ \left\lbrack
\begin{array}{c}
$g1 \\
0
\end{array} 
\right\rbrack.
\]

\{ BeginList('OL',type=>'a') \}

$ITEMSEP
$ITEM Form the complementary solution to the homogeneous equation.
$BR
$BR
\{ 
mbox(
[
' \( \boldsymbol{\vec{y}}_C(t) = c_1 \)', 
display_matrix([[$multians1->ans_rule(15)],[$multians1->ans_rule(15)]]),
' \( + \ c_2 \) ', 
display_matrix([[$multians1->ans_rule(15)],[$multians1->ans_rule(15)]]),
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM Show that seeking a particular solution of the form \( \boldsymbol{\vec{y}}_P(t) = $g1 \boldsymbol{\vec{a}} \), where \( \boldsymbol{\vec{a}} = \left\lbrack
\begin{array}{c} a_1 \\ a_2 \end{array} \right\rbrack \) is a constant vector, does not work.  In fact, if \( \boldsymbol{\vec{y}}_P \) had this form, we would arrive at the following contradiction:
$BR
$BR
\( a_2 = \) \{ ans_rule(5) \} \( \cdot a_1 \)
$BR
$BR
and
$BR
$BR
\( a_2 = \) \{ ans_rule(5) \} \( \cdot a_1 - \) \{ ans_rule(5) \}.

$ITEMSEP
$ITEM Show that seeking a particular solution of the form \( \boldsymbol{\vec{y}}_P(t) = t $g1 \boldsymbol{\vec{a}} \), where \( \boldsymbol{\vec{a}} = \left\lbrack
\begin{array}{c} a_1 \\ a_2 \end{array} \right\rbrack \) is a constant vector, does not work either.  In fact, if \( \boldsymbol{\vec{y}}_P \) had this form, we would arrive at the following contradiction:
$BR
$BR
\( a_2 = \) \{ ans_rule(5) \} \( \cdot a_1 \)
$BR
$BR
and
$BR
$BR
\( a_1 = \) \{ ans_rule(5) \}
$BR
$BR
and
$BR
$BR
\( a_2 = \) \{ ans_rule(5) \}.


\{ EndList('OL',type=>'a') \}
END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $partb1->cmp() );
ANS( $partb2a->cmp() );
ANS( $partb2b->cmp() );
ANS( $partc1->cmp() );
ANS( $partc2->cmp() );
ANS( $partc3->cmp() );

ENDDOCUMENT();









$multianshom = MultiAnswer( $corv11, $corv21, $corv12, $corv22 )->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "C1 &lt;%s, %s&gt; + C2 &lt;%s, %s&gt;",
    tex_format =>  "\text{\(C1\)} <%s, %s> + \text{\(C2\)} <%s, %s>" ,
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($stuv11, $stuv21, $stuv12, $stuv22) = @{$student}[0..3];


        # if both vectors are nonempty and not multiples of each other (takes care of nonzero assumption)
        if ( 
             ($stuv11 && $stuv21 && $stuv12 && $stuv22) && 
             ($stuv11 * $stuv22 != $stuv21 * $stuv12)
            ) { 
              my $deriv11 = $stuv11->D('t');
              my $deriv21 = $stuv21->D('t');
              my $deriv12 = $stuv12->D('t');
              my $deriv22 = $stuv22->D('t');
              if (
                    ($a[0][0]*$stuv11 + $a[0][1]*$stuv21 == $deriv11) &&
                    ($a[1][0]*$stuv11 + $a[1][1]*$stuv21 == $deriv21) &&
                    ($a[0][0]*$stuv12 + $a[0][1]*$stuv22 == $deriv12) &&
                    ($a[1][0]*$stuv12 + $a[1][1]*$stuv22 == $deriv22)
              )
               { return 1; }
        }

        # if the left vector is nonempty and nonzero
        if ( ($stuv11 && $stuv21) && ($stuv11 != 0 || $stuv21 != 0) ) {
              my $deriv11 = $stuv11->D('t');
              my $deriv21 = $stuv21->D('t');
              if (
                    ($a[0][0]*$stuv11 + $a[0][1]*$stuv21 == $deriv11) &&
                    ($a[1][0]*$stuv11 + $a[1][1]*$stuv21 == $deriv21)
              )
           { 
              $self->setMessage(1,"The vector on the left is correct.");
              return 0.67;
           }
        }

        # if the right vector is nonempty and nonzero
        if ( ($stuv12 && $stuv22) && ($stuv12 != 0 || $stuv22 != 0) ) {
              my $deriv12 = $stuv12->D('t');
              my $deriv22 = $stuv22->D('t');
              if (
                    ($a[0][0]*$stuv12 + $a[0][1]*$stuv22 == $deriv12) &&
                    ($a[1][0]*$stuv12 + $a[1][1]*$stuv22 == $deriv22)
              )
           { 
              $self->setMessage(2,"The vector on the right is correct.");
              return 0.67;
           }
        }
       
        return 0;

    }
);


