## DESCRIPTION
## Linear systems: Nonhomogeneous linear systems
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Nonhomogeneous systems)
## Date(04/12/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-NH-IVP)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.8')
## Problem1('7')
## KEYWORDS('Nonhomogeneous linear systems')


##################################
#  Initialization


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##################################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real"
);

# Eigenvectors:  <1,-1>, <1,-2>
# Eigenvalues:    -1,      1

#  The matrix for the system
$A11 = -3;  $A12 = -2;
$A21 =  4;  $A22 =  3; 


# Redefine true eigenvalues of matrix A
$lambda1 = -1;
$lambda2 =  1;
 

#########################################
#  Homogeneous (complementary) solution

# Correct complementary solution
$yc1top = Compute(" 1 e^(-t)");
$yc1bot = Compute("-1 e^(-t)");

$yc2top = Compute(" 1 e^(t)");
$yc2bot = Compute("-2 e^(t)");


$multians1 = MultiAnswer($yc1top, $yc1bot, $yc2top, $yc2bot)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='0'>
    <tr>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
      <td>&nbsp;&nbsp;</td>
      <td><table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table></td>
    </tr>
    </table>",
    tex_format =>  
    "\left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack, \ \ 
     \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};

        my $flag1 = 0;
        my $flag2 = 0;

        Value::Error("Your solutions are not linearly independent") if
        $c[0]->typeMatch($s[0]) &&
        $c[1]->typeMatch($s[1]) &&
        $c[2]->typeMatch($s[2]) &&
        $c[3]->typeMatch($s[3]) &&
        ($s[0] * $s[3] == $s[1] * $s[2]);

        if ( $c[0]->typeMatch($s[0]) && 
             $c[1]->typeMatch($s[1]) && 
             ($s[0] != 0 || $s[1] != 0) ) 
        {
          if ( ($A11 * $s[0] + $A12 * $s[1] == $s[0]->D('t')) &&
               ($A21 * $s[0] + $A22 * $s[1] == $s[1]->D('t')) 
          ){
                $flag1=1;
          }
        }

        if ( $c[2]->typeMatch($s[2]) && 
             $c[3]->typeMatch($s[3]) && 
             ($s[2] != 0 || $s[3] != 0) ) 
        {
          if ( ($A11 * $s[2] + $A12 * $s[3] == $s[2]->D('t')) &&
               ($A21 * $s[2] + $A22 * $s[3] == $s[3]->D('t')) 
          ){
                $flag2=1;
          }
        }

        if ($flag1==1 && $flag2 ==1){
           return 1;
        } elsif ($flag1 == 1) {
           Value::Error("The vector on the left is correct");
        } elsif ($flag2 == 1) {
           Value::Error("The vector on the right is correct");
        }else{
           Value::Error("Neither of the vectors is correct");
        }
       
    }
);



#############################################
#  Particular solution

# Random entry in nonhomogeneous part
$k = non_zero_random(-4,4,2);

# Vector g(t) = < $k sin(t), 0 >
# Guess solution of the form \vec{a} sin(t) + \vec{b} cos(t)

$g1 = Formula("$k sin(t)");
$g2 = Formula("0");

$a = 1.5 * $k;
$b = -0.5 * $k;
$c = -2 * $k;

$yptop = Compute("$a sin(t) + $b cos(t)");
$ypbot = Compute("$c sin(t)");

$multians2 = MultiAnswer( $yptop, $ypbot )->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='5'>
    <tr><td> %s </td></tr>
    <tr><td> %s </td></tr>
    </table>",
    tex_format =>  "\left\lbrack\begin{array}{rr} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};
        my @score = ();

        foreach my $i (0..1) {
          if ( $c[$i]->typeMatch($s[$i]) && $c[$i] == $s[$i] ) { 
             $score[$i] = 1; 
          } else { 
             $score[$i] = 0; 
          }
        }

        if ($score[0] == 1 && $score[1] == 1) {
           return 1;
        } elsif ($score[0] == 1) {
           Value::Error('The top entry is correct');
        } elsif ($score[1] == 1) {
           Value::Error('The bottom entry is correct');
        } else {
           Value::Error('Neither entry is correct');
        }

   }
);



#########################################
#  IVP solution

$C1 = non_zero_random(-5,5,1);
do { $C2 = non_zero_random(-5,5,1); } until ($C2 != $C1);

$init1 = -1 * $C1 - ($C2 / 2) - ($k / 2);
$init2 = $C1 + $C2;


#IVP solution

$y1ivp = Compute("-$C1 e^(-t) - ($C2 / 2) e^t + $a sin(t) + $b cos(t)");
$y2ivp = Compute("$C1 e^(-t) + $C2 e^t + $c sin(t)");

$multians3 = MultiAnswer( $y1ivp, $y2ivp )->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='5'>
    <tr><td> %s </td></tr>
    <tr><td> %s </td></tr>
    </table>",
    tex_format =>  "\left\lbrack\begin{array}{rr} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};
        my @score = ();

        Value::Error('Answers should be non-constant functions') unless 
           $c[0]->typeMatch($s[0])
        && $c[1]->typeMatch($s[1])
        && $s[0]->isFormula 
        && $s[1]->isFormula;

        foreach my $i (0..1) {
          if ( $c[$i]->typeMatch($s[$i]) && $c[$i] == $s[$i] ) { 
             $score[$i] = 1; 
          } else { 
             $score[$i] = 0; 
          }
        }

        if ($score[0] == 1 && $score[1] == 1) {
           return 1;
        } elsif ($score[0] == 1) {
           Value::Error('The top entry is correct');
        } elsif ($score[1] == 1) {
           Value::Error('The bottom entry is correct');
        } else {
           Value::Error('Neither entry is correct');
        }

   }
);



#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Consider the initial value problem
\[
\displaystyle
\ \boldsymbol{\vec{y}^{\, \prime}}
=
\left\lbrack
\begin{array}{rr}
$A11 & $A12 \\
$A21 & $A22
\end{array} 
\right\rbrack
\boldsymbol{\vec{y}}
+ \left\lbrack
\begin{array}{c}
$g1 \\
$g2
\end{array} 
\right\rbrack, 
\quad \quad
\boldsymbol{\vec{y}}(0)
= \left\lbrack
\begin{array}{c}
$init1 \\
$init2
\end{array} 
\right\rbrack.
\]

\{ BeginList('OL',type=>'a') \}

$ITEMSEP
$ITEM Form the complementary solution to the homogeneous equation.
$BR
$BR
\{ 
mbox(
[
' \( \boldsymbol{\vec{y}}_C(t) = c_1 \)', 
display_matrix([[$multians1->ans_rule(15)],[$multians1->ans_rule(15)]]),
' \( + \ c_2 \) ', 
display_matrix([[$multians1->ans_rule(15)],[$multians1->ans_rule(15)]]),
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM Construct a particular solution by assuming 
the form \( \boldsymbol{\vec{y}}_P(t) = (\sin t) \boldsymbol{\vec{a}} 
+ (\cos t) \boldsymbol{\vec{b}} \) and solving for the undetermined 
constant vectors \( \boldsymbol{\vec{a}} \) and 
\( \boldsymbol{\vec{b}} \).
$BR
$BR
\{ 
mbox(
[
'\( \boldsymbol{\vec{y}}_P(t) = \)',
display_matrix([[$multians2->ans_rule(15)],[$multians2->ans_rule(15)]]),
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM Form the general solution \( \boldsymbol{\vec{y}}(t) = 
\boldsymbol{\vec{y}}_C(t) + \boldsymbol{\vec{y}}_P(t) \) and 
impose the initial condition to obtain the solution of the 
initial value problem.
$BR
$BR
\{
mbox(
[
display_matrix([['\( y_1(t) \)'],['\( y_2(t) \)']]),
'\( = \)',
display_matrix([[$multians3->ans_rule(50)],[$multians3->ans_rule(50)]]),
],
allowbreaks=>'yes'
)
\}

\{ EndList('OL',type=>'a') \}
END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );
ANS( $multians3->cmp() );

;

ENDDOCUMENT();
