## DESCRIPTION
## Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 8/6/06

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Compositions and combinations of functions)
## Date(6/3/2002)
## Institution(Rochester)
## Author(K. Lesh)
## MLT(ccf-decompose-find-inner)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Precalculus')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('2.6')
## Problem1('43')
## TitleText2('College Algebra')
## AuthorText2('Stewart, Redlin, Watson')
## EditionText2('3')
## Section2('4.7')
## Problem2('45')
## KEYWORDS('algebra','function','composition of functions', 'combining functions', 'domain','domain', 'function' 'composition')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

###################################
# Setup
$a = non_zero_random(-9,9,1);
$b = random(2,9,1);

$f=Formula("x^$b")->reduce;
$g=Formula("x+$a ")->reduce;

$fcompg = $f ->substitute(x=>$g);
#$gcompf = $g ->substitute(x=>$f);

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT
If the function
\( h(x)=$fcompg\) is expressed in the form \(f \circ g\) with
\(f(x)=$f,\) then find the function \(g(x)\).
$PAR
\(g(x)\)= \{ans_rule(15)\}
END_TEXT
Context()->normalStrings;

###################################
# Answers

ANS( $fcompg->cmp( 
   checker=>sub {
   my ( $correct, $student, $ansHash ) = @_;
      return  $correct == ($f ->substitute(x=>$student))->reduce;
   } 
) );

##ANS($g->cmp);

###################################

;
ENDDOCUMENT();
