# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Rewrite an English sentence as a mathematical equation
# The quotient of a number and $a is $b
#
# Last edited: Jordan 6/26/13
#
# ENDDESCRIPTION

## DBCCSS('6.EE.2.a')
## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Algebraic expressions)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MLT(English_to_expression)
## Level(2)
## MO(1)
## KEYWORDS('equation','english-to-math')



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserImplicitPlane.pl",
  "contextFraction.pl",
  "contextTypeset.pl",
  "PGcourse.pl"
);

##############################################

Context("Fraction");

$sol = random(1,20,1);
$a = $sol;
while($a==$sol or gcd($a,$sol)!=1)
{
    $sol = random(1,20,1);
    $a = random(2,40,1);
}
#do {$a = random(2,40,1)} until ($a != $sol and gcd($a,$sol) != 1);
$b = Fraction($sol,$a);
($bnum, $bden) = $b->value;

Context("ImplicitPlane");
Context()->flags->set(showExtraParens=>0);
Context()->flags->set(reduceConstants=>0);
$ans = Compute("x/$a=$bnum/$bden");

$bperl = $sol/$a;
$bdenom = ($b->value)[1];



##############################################

TEXT(beginproblem());


BEGIN_PGML

Translate the following sentence into a math equation. Use [`x`] to represent the unknown number.

    The quotient of a number and [@numberWord($a)@] is [@numberWord($bperl,denominator=>$bdenom)@].

    [_____________]

END_PGML

##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));


BEGIN_PGML_SOLUTION

The word 'is' implies equality. So we have 

    The quotient of a number and [@numberWord($a)@] [`=`] [@numberWord($bperl,denominator=>$bdenom)@].

The word 'quotient' implies division, and that division is done left to right. So the equation we want is

    [`[@Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX@]`]


END_PGML_SOLUTION

ENDDOCUMENT();


