# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Rewrite an English sentence as a mathematical equation
# $a times a number is $b
#
# Last edited: Jordan 6/26/13
#
# ENDDESCRIPTION

## DBCCSS('6.EE.2.a')
## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Algebraic expressions)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## MLT(English_to_expression)
## Level(2)
## MO(1)
## KEYWORDS('equation','english-to-math')



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserImplicitPlane.pl",
  "contextTypeset.pl",
  "PGcourse.pl"
);

##############################################

Context("ImplicitPlane");
Context()->flags->set(showExtraParens=>0);


$a = random(2,10,1);
$sol = random(1,8,1);
$b = $a*$sol;

$ans = Compute("$a x=$b");

##############################################

TEXT(beginproblem());

BEGIN_PGML

Translate the following sentence into a math equation. Use [`x`] to represent the unknown number.

    [@numberWord($a,capital=>1)@] times a number is [@numberWord($b)@].

    [_____________]

END_PGML

##############################################

ANS($ans->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));



BEGIN_PGML_SOLUTION

The word 'is' implies equality. So we have 

    [`[$a]`] times [`x`] [`=`] [`[$b]`]

The word 'times' implies multiplication. So the equation we want is

    [`[@Parser::Formula(Context("Typeset") , $ans->{correct_ans})->TeX@]`]

END_PGML_SOLUTION

ENDDOCUMENT();


