# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Rewrite an English sentence as a mathematical equation
# The product of $a and a number, increased by $b, is $c
#
#
# Last edited: Jordan 6/26/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('equation','english-to-math')
## DBCCSS('6.EE.2.a')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserImplicitPlane.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
$var=RandomVariableName();
Context()->variables->are($var=>"Real");

$a = random(2,8,1);
$b = random(2,8,1);

$sol = random(5,50,1);
$c = $a*$sol+$b;

Context("ImplicitPlane");
Context()->variables->are($var=>"Real");
Context()->noreduce('(-x)-y','(-x)+y');
Context()->flags->set(showExtraParens=>0);
Context()->flags->set(reduceConstants=>0);

$ans1string = "$a*$var+$b=$c";
$ans1 = Compute($ans1string);

$eqWord = list_random('is','gives','yields');
$addPhrase = list_random('increased by', 'added to');

$left = Compute("$a*$var+$b")->reduce;
$right = "$c";


##############################################

TEXT(beginproblem());

BEGIN_PGML

Translate the following phrase or sentence into a math expression or equation (whichever is appropriate).
[@KeyboardInstructions('Use [`[$var]`] to represent the unknown number.')@]**

    The product of [@numberWord($a)@] and a number, [$addPhrase]  [@numberWord($b)@], [$eqWord] [@numberWord($c)@].

    [_______________]

END_PGML

##############################################

ANS($ans1->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans1string)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans1string)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  Context()->variables->are($var=>"Real");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));


BEGIN_PGML_SOLUTION

The word '[$eqWord]' implies equality. So we have 

    The product of [@numberWord($a)@] and a number, [$addPhrase]  [@numberWord($b)@] [`=`] [@numberWord($c)@].

The comma tells us that we should completely translate the content that comes before it, before we do anything else. The word 'product' implies multiplication. The two items to be multiplied are '[@numberWord($a)@]' and 'a number', so we have:

    [`[$a][$var]`], [$addPhrase]  [@numberWord($b)@] [`=`] [@numberWord($c)@].

The phrase '[$addPhrase]' implies addition, so we add [`[$a][$var]`] and [`[$b]`] and have:

    [``[$left]=[$right]``]


END_PGML_SOLUTION

ENDDOCUMENT();


