# WeBWorK problem written by Alex Jordan
# Portland Community College
#
# Rewrite an English sentence as a mathematical equation
# The quotient of a number and $a is $b
#
# Last edited: Jordan 6/26/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Algebraic Expressions')
## KEYWORDS('equation','english-to-math')
## DBCCSS('6.EE.2.a')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserImplicitPlane.pl",
  "contextFraction.pl",
  "contextTypeset.pl",
  "PGcourse.pl",
);

##############################################

Context("Fraction");
$var=RandomVariableName();
Context()->variables->are($var=>"Real");

$sol = random(1,20,1);
$a = $sol;
while($a==$sol or gcd($a,$sol)!=1)
{
    $sol = random(1,20,1);
    $a = random(2,40,1);
}
#do {$a = random(2,40,1)} until ($a != $sol and gcd($a,$sol) != 1);
$b = Fraction($sol,$a);
($bnum, $bden) = $b->value;

Context("ImplicitPlane");
Context()->variables->are($var=>"Real");
Context()->flags->set(showExtraParens=>0);
Context()->flags->set(reduceConstants=>0);

$ans1string = "$var/$a=$bnum/$bden";
$ans1 = Compute($ans1string);

$bperl = $sol/$a;
$bdenom = ($b->value)[1];

$left = Compute("$var/$a");
$right = "\frac{$bnum}{$bden}";

##############################################

TEXT(beginproblem());


BEGIN_PGML

Translate the following phrase or sentence into a math expression or equation (whichever is appropriate).
[@KeyboardInstructions('Use [`[$var]`] to represent the unknown number.')@]**

    The quotient of a number and [@numberWord($a)@] is [@numberWord($bperl,denominator=>$bdenom)@].

    [_____________]

END_PGML

##############################################

ANS($ans1->cmp(
  cmp_class => "a linear equation",
  correct_ans=>Parser::Formula(Context("Typeset") , $ans1string)->string,
  correct_ans_latex_string=>Parser::Formula(Context("Typeset") , $ans1string)->TeX, 
  checker => sub {
  my ($correct,$student,$ansHash) = @_;
  $result = ($correct == $student);
  my $context = Context();
  Context("Typeset");
  Context()->variables->are($var=>"Real");
  $student = Parser::Formula(Context("Typeset"), $ansHash->{original_student_ans});
  $ansHash->{preview_text_string} = $student->string;
  $ansHash->{preview_latex_string} = $student->TeX;
  $ansHash->{student_ans} = $student->string;  
  Context($context);
  return $result
}));


BEGIN_PGML_SOLUTION

The word 'is' implies equality. So we have 

    The quotient of a number and [@numberWord($a)@] [`=`] [@numberWord($bperl,denominator=>$bdenom)@].

The word 'quotient' implies division, and that division is done left to right. So the equation we want is

    [``[$left]=[$right]``]


END_PGML_SOLUTION

ENDDOCUMENT();


