# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>
# ENDDESCRIPTION

## DBsubject('Precalculus')
## DBchapter('Trigonometric Functions')
## DBsection('Inverse Trigonometric Functions')
## KEYWORDS('arctan','arcsin','arccos','reference angle')
## TitleText1('Functions Modeling Change')
## EditionText1('3')
## AuthorText1('Connally')
## Section1('6.Review')
## Problem1('40')
## TitleText2('Functions Modeling Change');
## EditionText2('4')
## AuthorText2('Connally')
## Section2('8.Review')
## Problem2('51')
## Author('Adam Spiegler')
## Institution('Loyola University Chicago')

DOCUMENT();

loadMacros("PG.pl",
           "PGbasicmacros.pl",
           "PGanswermacros.pl",
           "PGstandard.pl",
           "extraAnswerEvaluators.pl",
           "PGauxiliaryFunctions.pl",
           "PGgraphmacros.pl",
           "MathObjects.pl",
           "PGchoicemacros.pl"
          # "parserNumberWithUnits.pl",
          # "parserFormulaWithUnits.pl",
          # "parserFormulaUpToConstant.pl",
          # "PGcourse.pl",
          );

TEXT(beginproblem());

Context("Numeric");


$showPartialCorrectAnswers = 1;

$h = random(1,5,1);
$value = random(0.11,0.39,.01);

$eqn = "\tan{ (\theta - $h )} = $value";
$ans = (arctan($value) + $h);

Context()->texStrings;
BEGIN_TEXT

Using your calculator, find, (if possible), a solution for \( \theta \)
in radians.  If no solution exists, enter ${BBOLD}NONE${EBOLD}.
$PAR
\( \displaystyle $eqn \)
$BR
\( \theta = \) \{ ans_rule(15) \}
$BR
(enter your answer accurate to at least four decimal places)

END_TEXT
Context()->normalStrings;

# ANS(num_cmp( $ans,tol=>0.001) );
ANS( Compute( "arctan($value) + $h" )->cmp( checker=>sub{
    my ( $c, $s, $a ) = @_;
    return $c == $s || Compute($value) == tan($s - $h);
} ) );

$rounded = round(atan($value)*10000)/10000;
$ans_rounded = round($ans*10000)/10000;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
Using the inverse tangent function on a calculator we get
\( \theta - $h = \arctan{($value)} \approx $rounded \).
Solving for \( \theta \), we get \( \theta = $ans_rounded \).

END_SOLUTION
Context()->normalStrings;



COMMENT('MathObject version');
ENDDOCUMENT();
