## DESCRIPTION
## Find a value of x for which a set of 3 vectors in R^3, depending on x, is linearly dependent.
##   The vectors have small integer entries.
##   There will usually (always?) only be one possible solution. If multiple solutions work, any of them will be accepted.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Euclidean spaces')
## DBsection('Linear independence')
## Date('February 28, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('2.3: Subspaces, Spans, and Linear Independence')
## Problem1('')
## KEYWORDS('Euclidean space','linear dependence')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "contextFraction.pl",
    "PGchoicemacros.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Matrix");

do {
    $c1 = random(-2,3);
    $c2 = random(-2,3);
    $c3 = random(-2,3);
    $c4 = random(-2,3);
    $c5 = random(-2,3);
    $c6 = random(-2,3);
    $c7 = random(-2,3);
    $c8 = random(-2,3);
    $c9 = random(-2,3);
    $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
        
    $r1 = random(-2,4);
    $r2 = random(-2,4);
    $R = Matrix([[1,0,$r1],[0,1,$r2],[0,0,0]]);
        
    $A = $C * $R;
    
    $r = random(1,3);
    $c = random(1,3);
    
    $ans = $A->element($r,$c);
    
    foreach $i  (0..2) {
        foreach $j (0..2) {
            $em[$i][$j] = 0;
        }
    }
    $em[$r-1][$c-1] = 1;

    $Em = Matrix(@em);

    $D = $A + $Em;
    $F = $A - $Em;
} until (abs($C->det) > 0.1 && abs($D->det) > 0.1 && abs($F->det) > 0.1);

foreach $i  (0..2) {
    foreach $j (0..2) {
        if($i == $r-1 && $j == $c-1){
            $ad[$i][$j] = "x";
        }else{
            $ad[$i][$j] = $A->element($i+1,$j+1);
        }
    }
}

   
##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a real number \( x \) for which the set of vectors \( \lbrace ($ad[0][0],$ad[0][1],$ad[0][2]), ($ad[1][0],$ad[1][1],$ad[1][2]), ($ad[2][0],$ad[2][1],$ad[2][2])\rbrace \) is linearly dependent.
$BR
$BR
\( x \): \{ ans_rule(3) \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $ans->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    
    foreach $i  (0..2) {
        foreach $j (0..2) {
            if($i == $r-1 && $j == $c-1){
                $ansm[$i][$j] = $student;
            }else{
                $ansm[$i][$j] = $A->element($i+1,$j+1);
            }
        }
    }
    $AnsMat = Matrix(@ansm);
    
    if( abs($AnsMat->det) < 0.1 )
        { return 1; } else { return 0; }
}));

ENDDOCUMENT();
