# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus','derivative')
## Tagged by glr 12/19/08

## DBsubject('Calculus')
## DBchapter('Limits and Derivatives')
## DBsection('Derivatives')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('2.4')
## Section2('2.4')
## Problem1('3')
## Problem2('1')

## Textbook tags
## HHChapter1('Key Concept: The Derivative')
## HHChapter2('Key Concept: The Derivative')
## HHSection1('Interpretations of the Derivative')
## HHSection2('Interpretations of the Derivative')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->strings->add( gallons=>{}, dollars=>{}, "gallons/dollar"=>{},
			 "dollars/gallon"=>{}, "gallons per dollar"=>{},
		         "dollars per gallon"=>{} );
$showPartialCorrectAnswers = 1;

# include javascript for pop-up explaining units
HEADER_TEXT(<<EOF);
<script language="javascript" type="text/javascript">
<!-- //
function unitspopup() {
    var url = "/webwork2_files/units.html";
    var opt = "height=625,width=600,location=no,menubar=no,status=no" +
              "resizable=yes,scrollbars=yes,toolbar=no,";
    window.open(url,'examdata_info',opt).focus();
}
// -->
</script>
EOF

$a = random(100,300,25);
$b = random(200,400,25);
$c = random(1.2,1.8,0.2);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

The cost, \( C \) (in dollars) to
produce \( g \) gallons of ice cream can
be expressed as \( C = f(g) \).
$PAR
${BBOLD}(a)$EBOLD
In the expression \( f($a) = $b \),
$BR
what are the units of $a? \{ ans_rule(10) \}
$BR
what are the units of $b? \{ ans_rule(10) \}
$PAR

${BBOLD}(b)$EBOLD
In the expression \( f'($a) = $c \),
$BR
what are the units of $a? \{ ans_rule(10) \}
$BR
what are the units of $c? \{ ans_rule(10) \}
$PAR

${BITALIC}(Be sure that you can carefully put into words the
meanings of each of these statement in terms of ice cream.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( String("gallons")->cmp() );
ANS( String("dollars")->cmp() );
ANS( String("gallons")->cmp() );
ANS( String("dollars/gallon")->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    return $student == $correct || $student == String("dollars per gallon");
} ) );

$ap1 = $a + 1;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The statement \(f($a) = $b\) means that it costs 350 dollars to
produce 200 gallons of ice cream.
$PAR

${BBOLD}(b)$EBOLD
The statement \(f'($a) = $c\) means that when the number of
gallons produced is $a, costs are increasing by about $c dollars/gallon
In other words, it costs about $c dollars to produce the
${ap1}st gallon of ice cream.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
