# DESCRIPTION
# Find the side lengths of an isosceles right triangle using a quadratic equation.
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Applications and models)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(3)
## MO(1)
## KEYWORDS('equations', 'quadratic')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPowers.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

package my::Function::numeric;
our @ISA = ('Parser::Function::numeric');
sub sqrt {
  my $self = shift;
  my $value = $self->context->flag("setSqrt");
  return $value+1 if $value && $_[0] == 1; # force sqrt(1) to be incorrect
  return $value if $value;
  return $self->SUPER::sqrt(@_);
}
package main;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
Context()->flags->set(reduceConstantFunctions=>0);
LimitedPowers::OnlyPositiveIntegers();
$k = random(1,9);
$b = 2*$k;
$c = -($k)**2;
$disc = ($b)**2 - 4*2*$c;
$t = 2*$k;
$legs = Formula("$k + $k*sqrt(2)");
$hyp = Formula("$t + $k*sqrt(2)");
$da = "legs are \($k + $k\sqrt{2}\) cm long and the hypotenuse is \($t + $k\sqrt{2}\) cm long";



###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
The hypotenuse of an isosceles right triangle is $k cm longer than either of its legs. Note that an isosceles right triangle is a right triangle whose legs are the same length. Find the exact length of its legs and its hypotenuse.
$PAR
Legs: \{ans_rule(20)\} cm
$PAR
Hypotenuse: \{ans_rule(20)\} cm
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS( $legs-> cmp( checker => sub {
    my ($correct,$student,$ans) = @_;
    return 0 if $ans->{isPreview} || $correct != $student;
    $student = $ans->{student_formula};
    $correct = $correct->{original_formula} if defined $correct->{original_formula};
    Context()->flags->set(setSqrt => 1);
    delete $correct->{test_values};
    delete $student->{test_values};
    my $OK = ($correct == $student);
    Context()->flags->set(setSqrt => 0);
    Value::Error("Check to see if your answer is simplified.") unless $OK;
    return $OK;
    }, 
    formatStudentAnswer=>"reduced",
  )
);
ANS( $hyp-> cmp( checker => sub {
    my ($correct,$student,$ans) = @_;
    return 0 if $ans->{isPreview} || $correct != $student;
    $student = $ans->{student_formula};
    $correct = $correct->{original_formula} if defined $correct->{original_formula};
    Context()->flags->set(setSqrt => 1);
    delete $correct->{test_values};
    delete $student->{test_values};
    my $OK = ($correct == $student);
    Context()->flags->set(setSqrt => 0);
    Value::Error("Check to see if your answer is simplified.") unless $OK;
    return $OK;
    }, 
    formatStudentAnswer=>"reduced",
  )
);




COMMENT('Desc: Find the side lengths of an isosceles right triangle using a quadratic equation.');
ENDDOCUMENT();
