## DESCRIPTION
## Algebraic equations
## ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## BookChapter('Rules for Expressions and the Reasons for Them')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring polynomials: general)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Algebra: Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('2.3')
## Problem1('34')
## KEYWORDS('Equations')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(z=>"Real");

$k = random(2,4,2);
$a = 1;
$b = list_random(3,5,7);

$p2 = $k;
$p1 = $k * ($b - $a);
$p0 = $k * $a * $b;

$expression = "\( $p2 z^2 + $p1 z - $p0 \)";

$fac0 = Formula("$k");
$fac1 = Formula("(z - 1)");
$fac2 = Formula("(z + $b)");

$multians = MultiAnswer($fac0, $fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f0stu, $f1stu, $f2stu ) = @{$student};
      my ( $f0, $f1, $f2 ) = @{$correct};
      if ( ($f0 == $f0stu && $f1 == $f1stu && $f2 == $f2stu) ||
           ($f0 == $f0stu && $f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1,1];
      } else {
          if      ($f0 == $f0stu && ($f1 == $f1stu || $f2 == $f1stu)) {
              return [1,1,0];
          } elsif ($f0 == $f0stu && ($f1 == $f2stu || $f2 == $f2stu)) {
              return [1,0,1];
          } elsif ($f0 == $f0stu && ($f1 != $f1stu && $f2 != $f2stu)) {
              return [1,0,0];
          } elsif ( ($f0 != $f0stu && ($f1 == $f1stu && $f2 == $f2stu)) ||
                    ($f0 != $f0stu && ($f1 == $f2stu && $f2 == $f1stu)) ) {
              return [0,1,1];
          } elsif ($f0 != $f0stu && ($f1 == $f1stu || $f2 == $f1stu)) {
              return [0,1,0];
          } elsif ($f1 != $f0stu && ($f1 == $f2stu || $f2 == $f2stu)) {
              return [0,0,1];
          } else {
              return [0,0,0];
          }
      }
  }
);



################################################
#  Main text

BEGIN_TEXT

Factor the expression $expression.  Simplify your 
answer as much as possible, and put the greatest
common factor in the first answer box.
$BR
$BR
\{ $multians->ans_rule(5) \}
\( \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \)
\{ AnswerFormatHelp("formula") \}

END_TEXT


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());


;
ENDDOCUMENT();
