## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Institution(TCNJ)
## MLT(diag3x3)
## Level(4)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('6.4')
## Problem1('')
## KEYWORDS('diagonalization' 'matrices')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context("Matrix");

$a11 = random(-1,1,1);
$a21 = random(-1,1,2);
$a31 = random(-1,1,2);
$b1 = random(-1,1,2);
$a12 = $b1 * $a11;
$m = random(-1,1,2);
$a22 = $b1 * $a21 + $m;
$a32 = $b1 * $a31;
$c = random(-1,1,1);
$d = random(-1,1,2);
$n = random(-1,1,2);
$a13 = $c * $a11 + $d * $a12 + $n;
$a23 = $c * $a21 + $d * $a22;
$a33 = $c * $a31 + $d * $a32;

$a = Matrix([[$a11,$a12,$a13],[$a21,$a22,$a23],[$a31,$a32,$a33]]);
$b = $a->inverse();

# b has integer entries since det($a) = +/-1

$e11=Real(random(-5,-2,1));
$e22=Real(random(-1,1,1));
$e33=Real(random(2,5,1));

$e = Matrix([[$e11,0,0],[0,$e22,0],[0,0,$e33]]);

$matrix = $a * $e *$b;

## to check that we have a correct answer, just verify that
##    the appropriate matrix multiplication works

$mm = MultiAnswer($a,$e11,$e22,$e33)->with(
   singleResult => 1,
   checker => sub {
     my ($correct, $student, $self ) = @_;
     my ( $corr_a, $corr_e11, $corr_e22, $corr_e33) = @{$correct};
     my ( $stu_a, $stu_e11, $stu_e22, $stu_e33) = @{$student};
     my $stu_b = $stu_a->inverse();
     my $stu_e = Matrix([[$stu_e11,0,0],[0,$stu_e22,0],[0,0,$stu_e33]]);
     $matrix == $stu_a*$stu_e*$stu_b ? return(1) : return(0); 
   }
);


Context()->texStrings;
BEGIN_TEXT
Let \(A = $matrix\).
$PAR
Find \(S\) and \(D\) such that \(A=SDS^{-1}\).
\{ mbox( '\(S=\)', $mm->ans_array, ', \( D=\)',
   display_matrix(
	[[ $mm->ans_rule(4), 0, 0 ],
         [ 0, $mm->ans_rule(4), 0 ],
         [ 0, 0, $mm->ans_rule(4) ]]), '.')
\}
END_TEXT
Context()->normalStrings;

ANS($mm->cmp() );

ENDDOCUMENT();       # This should be the last executable line in the problem.

