## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Systems of linear equations)
## DBsection(Systems with 3 variables)
## Date(March 2016)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem());

Context('Matrix');

$a = non_zero_random(-9,9,1);
$b = non_zero_random(-8,8,2);
$c = non_zero_random(-9,9,2);


$A = Matrix([[$b],[- $a],[0]]);
$B = Matrix([[$c],[0],[- $a]]);

$multians = MultiAnswer($A, $B)->with(
  singleResult => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my @s = @{$student};
    my @c = @{$correct};
    my $s0 = Matrix($s[0]);
    my $s1 = Matrix($s[1]);
    my $grad = Matrix([[$a,$b,$c]]);

    # if $s0 and $s1 are dependent, return 0.
    my $M = Matrix([ 
    [$s0->element(1,1), $s1->element(1,1)], 
    [$s0->element(2,1), $s1->element(2,1)], 
    [$s0->element(3,1), $s1->element(3,1)]]);
    if ( (($M->transpose) * $M)->det == 0 ) { return 0; }

    # return 1 if both $s0 and $s1 are solutions, 0 otherwise
    return ( ($grad * $s0 == Matrix([[0]]) ) && ($grad * $s0 == Matrix([[0]]) ) );
  }
);

Context()->texStrings;
BEGIN_TEXT
Solve the equation \[ $a x + $b y + $c z = 0. \]
$PAR
\(
\left[
\begin{array}{c}
x \\
y \\
z
\end{array}
\right]
= s \)
\{ $multians->ans_array(3) \}
+ \( t \)
\{ $multians->ans_array(3) \}
END_TEXT
Context()->normalStrings;

ANS( $multians->cmp() );


ENDDOCUMENT();
