##DESCRIPTION
##  College Algebra Review: 
##  Factoring quadratics with leading term greater than 1
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## DBsubject(Algebra)
## Date(Spring 2012)
## Institution(University of Missouri)
## MO(1)
## Status(R)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.4')
## Problem1('')
## KEYWORDS('algebra', 'factoring', 'quadratic')

########################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());


#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(x=>"Real");

do {
  $a = random(1,4);
  do {$b = non_zero_random(-4,4)} until (gcd($a,$b) == 1);
  $c = random(1,4);
  do {$d = non_zero_random(-4,4)} until (gcd($c,$d) == 1);
  $e1 = $a*$c;
  $e2 = $a*$d+$b*$c;
  $e3 = $b*$d;
} until ($e1 > 1 && $e2 > 0 && $e3 < 0 && ($a != $c && $b != $d));


$fac1 = Formula(" $a x + $b ");
$fac2 = Formula(" $c x + $d  ")->reduce;


$expression = "\( $e1 x^2 + $e2 x + $e3 \)";
$texanswer = "\( \left($a x + $b \right)\left($c x + $d \right)\)";
if ($a == 1) {
$texanswer = "\( \left( x + $b \right)\left($c x + $d \right)\)";
} elsif ($c == 1) {
$texanswer = "\( \left($a x + $b \right)\left( x + $d \right)\)";
}
if ($e2 == 1) {
  $expression = "\( $e1 x^2 + x + $e3 \)";
} elsif ($e2 == -1) {
  $expression = "\( $e1 x^2 -  x + $e3 \)";
}

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your 
answer as much as possible:
$PAR
${BCENTER}
$expression \( = \) 
\( \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \)
${ECENTER}
END_TEXT



################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

ENDDOCUMENT();