##DESCRIPTION
##  Algebra problem: solve a quadratic equation
##ENDDESCRIPTION

## pphi tagged and PAID on 12-12-2003

## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Solving equations)
## Date(6/3/2002)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.6')
## Problem1('47')
## KEYWORDS('algebra', 'quadratic equations')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

Context("Numeric");
$a = Real(random(2,7,1));
$b = $a**3;

$ans1 = Compute("-$a*sqrt($a)");
$ans2 = Compute("$a*sqrt($a)");

$ans = MultiAnswer($ans1,$ans2)->with(
    singleResult => 1,
    checker => sub {
        my ($correct,$student,$multi_ans,$ans_hash) = @_; 
        my ($corr1,$corr2) = @{$correct};
        my ($x1,$x2) = @{$student};  
        return (($corr1 == $x1 && $corr2 == $x2) or ($corr1 == $x2 && $corr2 == $x1)) ;
    },
);


BEGIN_PGML;
Find all real solutions of the equation [`x^2-[$b]=0`].

    [` x = `] [_]{$ans}{15} [`\ `] or [`\ x = `] [_]{$ans}{15}  
  
END_PGML

BEGIN_PGML_SOLUTION
    [`x=\pm [$a]\sqrt{[$a]}`]
END_PGML_SOLUTION



ENDDOCUMENT();        # This should be the last executable line in the problem.
