##DESCRIPTION
##  Algebra problem: quadratic equation
##ENDDESCRIPTION

## Tagged by sawblade

## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Solve by factoring)
## Date(6/3/2002)
## MLT(quadraticFactorMonic)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('3.3')
## Problem1('1')
## KEYWORDS('algebra', 'quadratic equations')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

$showPartialCorrectAnswers = 1;

Context("Numeric");

$x1 =  Real(random(1,4,1));
$x2 =  Real($x1+random(2,4,1));

# (x+$x1)(x-$x2)=0

$b = $x1-$x2;
$c = -$x1*$x2;

$f=Compute("x^2+$b*x+$c")->reduce();

$ans = MultiAnswer(-$x1,$x2)->with(
    singleResult => 1,
    checker => sub {
        my ($correct,$student,$multi_ans,$ans_hash) = @_; 
        my ($corr1,$corr2) = @{$correct};
        my ($stu1,$stu2) = @{$student};  
        return (($corr1 == $stu1 && $corr2 == $stu2) or ($corr1 == $stu2 && $corr2 == $stu1));
    }
);

BEGIN_PGML;
Solve the equation [`[$f]=0`] by factoring.

    [` x = `] [_]{$ans}{15} or [` x = `] [_]{$ans}{15} 
  
END_PGML

BEGIN_PGML_SOLUTION
    [``[$f]=(x+[$x1])(x-[$x2])``] 
    
so 

    [`0=[$f]`] if and only if [`x=-[$x1]`] or [`x=[$x2]`].
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.
