##DESCRIPTION
##  Algebra problem: quadratic equation
##ENDDESCRIPTION

## tsch tagged and PAID on 12-12-2003

## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Solving equations)
## Date(6/3/2002)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('3.3')
## Problem1('31')
## KEYWORDS('algebra', 'quadratic equations')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl','parserPopUp.pl', 'PGcourse.pl');

Context("Numeric");
Context()->variables->are(z=>"Real");

$b =  Real(random(4,8,1));
$c =  Real(random(1,3,1));
$D = $b**2-4*$c;


$f = Compute("$c+$b*z+z^2");

$z1 = Compute("(-$b+sqrt($D))/2");
$z2 = Compute("(-$b-sqrt($D))/2");

$YN=DropDown(["Yes","No"],"Yes");

$ans = MultiAnswer($z1,$z2)->with(
    singleResult => 1,
    checker => sub {
        my ($correct,$student,$multi_ans,$ans_hash) = @_; 
        my ($corr1,$corr2) = @{$correct};
        my ($stu1,$stu2) = @{$student};  
        return (($corr1 == $stu1 && $corr2 == $stu2) or ($corr1 == $stu2 && $corr2 == $stu1));
    }
);

BEGIN_PGML;
Does the equation [`[$f]=0`] have real solutions? [_]{$YN}

If your answer is Yes, enter the solutions:
  
    [` z = `] [_]{$ans}{25} or [` z = `] [_]{$ans}{25} 
  
END_PGML

BEGIN_PGML_SOLUTION
The quadratic formula says the solutions are 

    [``z=\frac{-[$b]\pm \sqrt{[$b]^2-(4)(1)([$c])}}{(2)(1)}``]

The solutions are real because [`[$b]^2-(4)(1)([$c])=[$D]\geq 0`] so its square root is a real number.
END_PGML_SOLUTION
ENDDOCUMENT();        # This should be the last executable line in the problem.