##DESCRIPTION
##  Algebra problem: quadratic equation
##ENDDESCRIPTION

## Tagged by sawblade

## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Solving equations)
## Date(6/3/2002)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('3.3')
## Problem1('25')
## KEYWORDS('algebra', 'quadratic equations')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserPopUp.pl','parserMultiAnswer.pl', 'PGcourse.pl');

Context("Numeric");

$a =  Real(random(2,4,1));
$b =  Real(random(1,6,1));
$c =  Real(random(1,6,1));

$f=Compute("$a*x^2+$b*x-$c");

$YN = DropDown(["Yes","No"],"Yes");

$D = $b**2+4*$a*$c;

$x1 = Compute("(-$b-sqrt($D))/(2*$a)");
$x2 = Compute("(-$b+sqrt($D))/(2*$a)");

$ans = MultiAnswer($x1,$x2)->with(
    singleResult => 1,
    checker => sub {
        my ($correct,$student,$multi_ans,$ans_hash) = @_; 
        my ($corr1,$corr2) = @{$correct};
        my ($stu1,$stu2) = @{$student};  
        return (($corr1 == $stu1 && $corr2 == $stu2) or ($corr1 == $stu2 && $corr2 == $stu1));
    }
);

BEGIN_PGML;
  
Does the equation [`[$f] = 0`] have real solutions? [_]{$YN}

If your answer is Yes, enter the solutions:
  
    [` x = `] [_]{$ans}{15} or [` x = `] [_]{$ans}{15} 
    
*Note:* WeBWorK can compute square roots and many other mathematical functions.  For example
to input a number like [``\frac{5+\sqrt{19}}{4}``] you may simply enter *(5+sqrt(19))/4* instead of typing out a decimal approximation. For more information click the link [@ helpLink("syntax")@]*, and look for  Mathematical Functions Available in WeBWorK. 
END_PGML

BEGIN_PGML_SOLUTION
If [`a,b,c`] are real numbers and [`a\neq 0`] then the quadratic formula says that the quadratic equation [`ax^2+bx+c=0`] has real solutions [``x=\frac{-b\pm\sqrt{b^2-4ac}}{2a}``], and these solutions are real if and only if [`b^2-4ac\geq 0`].  

In our case [`a=[$a]`],[`b=[$b]`] and [`c=-[$c]`] so [`b^2-4ac = [$D]\geq 0`], and so the equation has real solutions

    [``x=\frac{-[$b]\pm \sqrt{[$D]}}{2([$a])}``]
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.
