## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Graphs)
## Level(2)
## KEYWORDS('domain', 'range', 'intercepts', 'rational function', 'rational functions')
## Author(Kyle Schlitt)
## Institution(TRU)
## Language(en)

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "contextFraction.pl"
);

Context("Fraction");

# choose the x-intercepts of f
$xintercept1num = random(1,4,1);
$xintercept1den = 5;
$xintercept1 = Fraction($xintercept1num, $xintercept1den)->reduce;

$xintercept2 = random(2,5);

# build the numerator of f so it has the desired xintercepts
$f_num_a = $xintercept1den;
$f_num_b = -($xintercept1den*$xintercept2 + $xintercept1num);
$f_num_c = $xintercept1num*$xintercept2;
$f_num = Compute("$f_num_a x^2 + $f_num_b x + $f_num_c")->reduce;

# choose the points at which f will not be defined
$singularity1 = random(5,9,1);

do {
  $singularity2 = random(1,7,1);
} while ($singularity2 == $singularity1);

# build the denominator of f so it has the desired singularities
$f_den_a = 1;
$f_den_b = -($singularity1 + $singularity2);
$f_den_c = $singularity1*$singularity2;

$f_den = Compute("$f_den_a x^2 + $f_den_b x + $f_den_c")->reduce;

$f = $f_num/$f_den;

$popup = DropDown(['a polynomial', 'a rational function', 'an exponential function', 'a quadratic function'], 'a rational function', placeholder => 'Select One');

BEGIN_PGML

Consider the function [`f(x) = \displaystyle{[$f]}`].

What type of function is [`f`]? 

**Answer:** [`f`] is [_]{$popup}.
----

END_PGML

Context("Interval");

$a = 0;$b = 0;

if ($singularity1 < $singularity2) {
  $a = $singularity1; $b = $singularity2;
} else {
  $a = $singularity2; $b = $singularity1;
}

$domain = Union("(-infinity, $a) U ($a, $b) U ($b, infinity)");

BEGIN_PGML

What is the domain of [`f`]?  Express your answer using interval notation.

**Answer:** [`\text{dom}(f)=`] [_]{$domain}

----

END_PGML

$intercepts = MultiAnswer($xintercept1, $xintercept2)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($x1stu, $x2stu) = @{$student};
    my ($x1, $x2) = @{$correct};
    
    if (($x1 == $x1stu && $x2 == $x2stu) || ($x1 == $x2stu && $x2 == $x1stu)) {
      return [1,1];
    } elsif ($x1 == $x1stu || $x2 == $x1stu) {
      return [1,0];
    } elsif ($x1 == $x2stu || $x2 == $x2stu) {
      return [0,1];
    } else {
      return [0,0];
    }
  }
);

BEGIN_PGML

What are the [`x`]-intercepts of the graph of [`f`]?

**Answer:** [`x =`] [_]{$intercepts} and [`x = `] [_]{$intercepts}

----
END_PGML

Context("Fraction");

$yintercept = Fraction("$f_num_c/$f_den_c")->reduce;

BEGIN_PGML

What is the [`y`]-intercept of the graph of [`f`]?

**Answer:** [`y =`] [_]{$yintercept}

END_PGML

ENDDOCUMENT();