# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 07/07/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector calculus)
## DBsection(Line integrals)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(my_new_Line_Integrals)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('18.2')
## Problem1('33')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('6')
## Section2('18.2')
## Problem2('41')
## KEYWORDS('line integral', 'integrals', 'parametric curve', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Line Integrals')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real', y=>'Real' );
$showPartialCorrectAnswers = 1;

## a useful display routine
sub dfrac {
    ($n, $d) = reduce( $_[0], $_[1] );
    return ( $d == 1 ) ? $n : "\frac{$n}{$d}";
}

$m = random(2,8,1);
$x0 = random(2,5,1);
$x0 = 3;
$y0 = $m*$x0;

$sgn = list_random(-1,1);
$sgn = list_random(-1,1);
$a = $sgn*random(1,5,1);
$b = $sgn*random(1,5,1);

$intsign = PopUp( [ '?', 'greater than zero', 'less than zero',
		    'equal to zero' ], ( $sgn > 0 ) ? 'greater than zero' :
		  'less than zero' );

$fx = Compute("$a*x*y")->reduce();
$fy = Compute("$b*x")->reduce();

$int = Compute( "$a*$m*$x0^3/3 + $b*$m*$x0^2/2" );

$relate = PopUp( [ '?', 'is greater than', 'is less than', 'is equal to' ],
		 'is equal to' );

$t0 = Compute("0");
$t1 = Compute("$x0");

if ( $x0 == 2 ) {
    $xt = "\frac{1}{3}\left(4\,t^3 - 12\,t^2 + 11\,t\right)";
    $xpt = "\frac{1}{3}\left(12\,t^2 - 24\,t + 11\right)";
    ($n,$d) = reduce($m, 6);
    $yc = ( $d == 1 ) ? $n : "\frac{$n}{$d}";
    $yc = '' if ( $yc eq '1' );
    $yt = "$yc\left(4\,t^3 - 12\,t^2 + 11\,t\right)";
    $ypt = "$yc\left(12\,t^2 - 24\,t + 11\right)";

    ## some equations for the solution
    $ic1 = dfrac( $m, 27 );
    $thb = 3*$b;
    $igr = "$ic1 t (4t^2 - 12t + 11)(12t^2 - 24t + 11)" .
        "($thb + $a t(4t^2 - 12t + 11))";
} elsif ( $x0 == 4 ) {
    $xt = "\frac{1}{3}\left(t^3 - 6\,t^2 + 11\,t\right)";
    $xpt = "\frac{1}{3}\left(3\,t^2 - 12\,t + 11\right)";
    ($n,$d) = reduce($m, 12);
    $yc = ( $d == 1 ) ? $n : "\frac{$n}{$d}";
    $yc = '' if ( $yc eq '1' );
    $yt = "$yc\left(t^3 - 6\,t^2 + 11\,t\right)";
    $ypt = "$yc\left(3\,t^2 - 12\,t + 11\right)";

    ## for the solution
    $ic1 = dfrac($m, 27);
    $thb = 3*$b;
    $igr = "$ic1(t^2 - 6t^2 + 11)(3t^2 - 12t + 11)($thb + $a t(t^2 - 6t + 11))";
} else {
    $d = 3*$x0*$x0;
    $k1 = 24*$x0;
    $k2 = 11*$x0*$x0;
    $twok1 = 2*$k1;
    $xt = "\frac{1}{$d}\left(16\,t^3 - $k1\,t^2 + $k2\,t\right)";
    $xpt = "\frac{1}{$d}\left(48\,t^2 - $twok1\,t + $k2\right)";
    ($n,$d) = reduce($m, $d);
    $yc = ( $d == 1 ) ? $n : "\frac{$n}{$d}";
    $yc = '' if ( $yc eq '1' );
    $yt = "$yc\left(16\,t^3 - $k1\,t^2 + $k2\,t\right)";
    $ypt = "$yc\left(48\,t^2 - $twok1\,t + $k2\right)";

    ## for the solution
    $ic1 = dfrac( $a*$m, 27*$x0**6 );
    $ic2 = dfrac( abs($b)*$m, 9*$x0**4 );
    $s = ( $sgn > 0 ) ? '+' : '-';
    $j1 = 48*$x0;
    $j2 = 11*$x0*$x0;
    $j3 = 24*$x0;
    $igr = "$ic1(48t^2 - $j1 t + $j2)(16t^3 - $j3 t^2 + $j2 t)^2 " .
        "$s $ic2(48t^2 - $j1 t + $j2)(16t^3 - $j3 t^2 + $j2 t)";
}

$myContext = Context()->copy();

Context("ImplicitEquation");
Context()->variables->set( x=>{limits=>[0,$x0]}, y=>{limits=>[0,$y0]} );
$eqn = ImplicitEquation( "y = $m*x" );

Context($myContext);
$tc1 = Compute("$x0/2 - $x0*sqrt(3)/12");
$tc2 = Compute("$x0/2 + $x0*sqrt(3)/12");
$tchg = MultiAnswer( $tc1, $tc2 )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ans ) = @_;
	if ( ($c->[0]==$s->[0] && $c->[1]==$s->[1]) ||
	     ($c->[0]==$s->[1] && $c->[1]==$s->[0]) ) {
	    return [1,1];
	} elsif ( $c->[0]==$s->[0] || $c->[1]==$s->[0] ) {
	    return [1,0];
	} elsif ( $c->[0]==$s->[1] || $c->[1]==$s->[1] ) {
	    return [0,1];
	} else {
	    return [0,0];
	}
    }
);

$integr = Compute("(1/(27*$x0^6))*$m*t*" .
		  "(768*t^4 - 1920*t^3*$x0 + 1856*t^2*$x0^2 - " .
		  "792*t*$x0^3 + 121*$x0^4)*(3*$b*$x0^2 + " .
		  "$a*t*(16*t^2 - 24*t*$x0 + 11*$x0^2))");
$int2 = Compute( "$a*$m*$x0^3/3 + $b*$m*$x0^2/2" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Suppose \( C \) is the line segment from the point \( (0,0) \) to the
point \( ($x0,$y0) \) and \( \vec{F}=$fx\,\vec i + $fy\,\vec j \).

$PAR
${BBOLD}(a)$EBOLD
Is \( \int_C \vec{F} \cdot d\vec{r} \) greater than, less than, or
equal to zero?  \{ $intsign->menu() \}
$BR
${BITALIC}(Be sure you are able to give a geometric explanation for
your answer.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
A parameterization of \( C \) is \( (x(t),y(t)) = (t,$m t) \) for
\(  0 \leq t \leq $x0 \).  Use this to compute
\( \int_C \vec{F} \cdot d\vec{r} \).
$BR
\( \int_C \vec{F} \cdot d\vec{r} = \) \{ ans_rule(45) \}

$PAR
${BBOLD}(c)$EBOLD
Suppose a particle leaves the point \( (0,0) \), moves along
the line toward the point \( ($x0,$y0) \), stops before reaching it and
backs up, stops  again and reverses direction, then completes its
journey to the endpoint.  All travel takes place along the line segment joining
the point \( (0,0) \) to the point \( ($x0,$y0) \).
Call this path \( C' \).  How is \(\int_{C'} \vec{F} \cdot d\vec{r}\)
related to \(\int_{C} \vec{F} \cdot d\vec{r} \)?
$BR
\(\int_{C'} \vec{F} \cdot d\vec{r}\)
\{ $relate->menu() \}
\(\int_{C} \vec{F} \cdot d\vec{r}\)
$BR
${BITALIC}(Be sure you can explain why this is.)$EITALIC

$PAR
${BBOLD}(d)$EBOLD
A parameterization for a path like \( C' \) is given by
\[
(x(t),y(t)) = \left($xt, $yt)\right).
\]
For what value of \(t\) is this parameterization
$BR
at \((0,0)\)?  \( t = \) \{ ans_rule(15) \}
$BR
at \(($x0,$y0)\)?  \( t = \) \{ ans_rule(15) \}
$PAR
What equation must \((x(t),y(t))\) satisfy?
$BR
\{ ans_rule(45) \}
$BR
${BITALIC}(Enter your equation in terms of \(x\) and \(y\): thus,
if \(y(t) = 3 x(t)^2 - 6\), you would enter \(y = 3x^2 - 6\).)$EITALIC.
$BR
By hand, on a separate sheet of paper, verify that this equation
is in fact satisfied by this parameterization.
$PAR
At what values of \(t\) does the particle change direction?
$BR
At \( t = \) \{ $tchg->ans_rule(25) \} and
\{ $tchg->ans_rule(25) \}

$PAR
${BBOLD}(e)$EBOLD
Set up and find \(\int_{C'} \vec{F} \cdot d\vec{r}\) using the
parameterization in part ${BBOLD}(d)$EBOLD.
$BR
With \(a = \) \{ ans_rule(15) \} and
\( b = \) \{ ans_rule(15) \},
$BR
\(\int_{C'} \vec{F} \cdot d\vec{r} = \int_a^b\)
\{ ans_rule(35) \} \(dt =\) \{ ans_rule(25) \}
$BR
${BITALIC}(Note whether you get the same answer as you did in part
${BBOLD}(b)$EBOLD.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS($intsign->cmp() );
ANS($int->cmp() );
ANS($relate->cmp() );
ANS($t0->cmp() );
ANS($t1->cmp() );
ANS($eqn->cmp() );
ANS($tchg->cmp() );
ANS($t0->cmp() );
ANS($t1->cmp() );
ANS($integr->cmp() );
ANS($int2->cmp() );

$which = ( $sgn > 0 ) ? 'same' : 'opposite';
$angle = ( $sgn > 0 ) ? 'less than' : 'greater than';
$am = $a*$m;
$s2 = ( $sgn > 0 ) ? '+' : '-';
$absb = abs($b);
$absbm = $absb*$m;
$alpha1 = dfrac( $am, 3 );
$alpha2 = dfrac($absbm, 2 );
$val = dfrac( (2*$a*$m*$x0**3 + 3*$b*$m*$x0**2), 6 );

$x0o2 = dfrac( $x0, 2 );
$x0o24 = dfrac( $x0, 24 );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The line integral \(\int_C ($a xy\,\vec i $s2 $absb x\,\vec j) \cdot d\vec{r}\)
is \{ $intsign->correct_ans() \}.
This follows from the fact that all of the vectors of 
\($a xy\,\vec i $s2 $absb x\,\vec j\)
at points along \(C\) point approximately in the $which direction
as \(C\) (meaning the angles between the vectors and the direction
of \(C\) are $angle \(\pi/2\)).

$PAR
${BBOLD}(b)$EBOLD
Using the parameterization \(x(t)=t\), \(y(t)=$m t\), with \(x'(t)=1\),
\(y'(t)=$m\), we have
\[
\int_C \vec{F} \cdot d\vec{r}  = \int_0^{$x0} \vec{F}(t,$m t) \cdot
(\vec i + 3\vec j)\,dt
= \int_0^{$x0} ($am t^2\vec i $s2 $absb t\vec j) \cdot (\vec i + $m\vec j)\,dt.
\]
Expanding the dot product and integrating, we have
\[
\int_C \vec{F} \cdot d\vec{r} = \int_0^{$x0} ($am t^2 $s2 $absbm t) \,dt
= \left($alpha1 t^3 $s2 $alpha2 t^2\right) \bigg|_0^{$x0} = $val.
\]

$PAR
${BBOLD}(c)$EBOLD
Call the point on the oriented path \(C'\) where the particle stops and
reverses direction \(P\) and the point where it changes direction again
\(Q\).  Then we can view  \(C'\) as a sum of three paths: \(C_1\), on
which the particle moves from \((0,0)\) along the line connecting \((0,0)\)
and \(($x0,$y0)\) until it gets to the point \(P\); \(C_2\), the path
connecting the point \(P\) and the point \(Q\); and \(C_3\), the path
along the line that extends from \(Q\) to the point \(($x0,$y0)\).
Thus, \(C'=C_1+C_2+C_3\).  Along the parts of \(C_1\) and
\(C_2\) that overlap, the line integrals cancel, so we are left with the
line integral over the part of \(C_1\) that does not overlap with \(C_2\),
followed by the line integral over \(C_3\). Thus, the line integral over \(C'\)
is the same as the line integral over the direct
route from the point \((0,0)\) to the point \(($x0,$y0)\).

$PAR
${BBOLD}(d)$EBOLD
The parameterization
\[ (x(t),y(t)) = \left($xt, $yt\right) \]
has \((x(0),y(0)) = (0,0)\) and \((x($x0),y($x0)) = ($x0,$y0)\).
To lie along the line, we must have that
\[ y(t) = y = $m x = $m x(t).\]
We can see, with a little bit of algebraic rearrangement, that the
expression \(y(t)\) is in fact equal to \($m\,x(t)\), therefore all
points on the parameterized curve lie on the line \(y = $m x\) which
is the curve \(C\).
$PAR
We can find the points at which the particle changes direction by
looking for where \(x'(t)\) (and \(y'(t)\)) change sign.  We have
\[ x(t) = $xt, \]
so that
\[ x'(t) = $xpt. \]
Setting this equal to zero and solving for \(t\) using the quadratic
formula, we have
\[ t = $x0o2 \pm $x0o24\sqrt{3}. \]

$PAR
${BBOLD}(e)$EBOLD
Using the parameterization given in part (d), we have
\[ \vec r'(t)=x'(t)\vec i+y'(t)\vec j=
$xpt\,\vec i + $ypt\,\vec j. \]
Thus,
\[
\int_{C'} \vec{F} \cdot d\vec{r} = \int_0^{$x0}
  \vec{F}($xt,$yt) \cdot
  ($xpt\,\vec i+ $ypt\,\vec j )\,dt,
\]
which is, after expanding the dot product and simplifying,
\[
\int_0^{$x0} $igr\,dt.
\]
Numerical integration yields an answer of \($val\)
which agrees with the answer found in part b.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
