# WeBWorK problem written by Alex Jordan 
# Portland Community College
#
# Template:
# Convert the expression to radical notation.
#
#       x^(m/n)
#
#
# Last updated: Spoddeck 11/1/15
# ENDDESCRIPTION


##############################################


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "niceTables.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "parserRoot.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);


##############################################


TEXT(beginproblem());

Context("Numeric");
parser::Root->Enable;
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
'/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
'//' => {class => 'bizarro::BOP::divide', isCommand => 1},
);
$var = list_random('x','y','z','t','r','m','n','a','b','c');
Context()->variables->are($var=>'Real');
Context()->reductions->set('(-x)+y'=>0,'(-x)-y'=>0);
Context()->flags->set(reduceConstants=>0);

$m = random(2,5,1);
$n = $m+1;

$display = Formula("$var^($m/$n)");


Context()->flags->set(limits=>[0.5,5]);

$ans = Formula("root($n,$var^($m))");

$evaluator = $ans->cmp(
     checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is not in the expected form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0);
      return 1;
  }
);



##############################################


BEGIN_PGML

Convert the expression to radical notation.  

[@KeyboardInstructions('To enter something like [`\sqrt[3]{x+1}`], type [|root(3,x+1)|]*.')@]**

[``[$display]``] = [__]{$evaluator}{7}

END_PGML


##############################################


BEGIN_PGML_SOLUTION

When we have fractions as exponents, the numerator indicates the power while the denominator indicates the radical: [` x^\frac{m}{n} = \sqrt[n]{x^m} `]. So we can rewrite our expression as follows:

    [``
\begin{aligned}
   [$var]^\frac{[$m]}{[$n]} &= \sqrt[[$n]]{[$var]^[$m]} \\
\end{aligned}
    ``]

END_PGML_SOLUTION


##############################################


ENDDOCUMENT();
