## DESCRIPTION
## Find a unit LU decomposition of a 3x3 matrix.
##    The matrix and the solution will have small integer entries.
##    If the given matrix is not invertible, it might have multiple unit LU decompositions. All of them are accepted.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Matrix factorizations')
## DBsection('LU factorization')
## Date('March 12, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('2.D: The LU Decomposition')
## Problem1('')
## KEYWORDS('linear algebra','matrix','LU decomposition')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

do {
    $l1 = random(-3,4);
    $l2 = random(-3,4);
    $l3 = random(-3,4);
    $L = Matrix([[1,0,0],[$l1,1,0],[$l2,$l3,1]]);
    
    $u1 = list_random(-2,-1,1,2,3);
    $u2 = list_random(-2,-1,1,2,3);
    $u3 = list_random(-2,-1,1,2,3);
    $u4 = list_random(-2,-1,1,2,3);
    $u5 = random(-2,3);
    $u6 = random(-2,3);
    $U = Matrix([[$u1,$u2,$u3],[0,$u4,$u5],[0,0,$u6]]);
    
    $A = $L * $U;
} until ((abs($l1) + abs($l2) > 0) || (abs($l1) + abs($l3) > 0) || (abs($l2) + abs($l3) > 0));

$multians = MultiAnswer($L,$U)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if($s0->element(1,1) == 1 && $s0->element(2,2) == 1 && $s0->element(3,3) == 1 && $s0->element(1,2) == 0 && $s0->element(1,3) == 0 && $s0->element(2,3) == 0 && $s1->element(2,1) == 0 && $s1->element(3,1) == 0 && $s1->element(3,2) == 0 && $A == $s0*$s1) {
            return 1;
        } else {
            return 0;
        }
    }
);


##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a unit LU Decomposition of the matrix \( \displaystyle A = $A \). That is, find a unit lower triangular matrix \( L \) and an upper triangular matrix \( U \) such that \( A = LU \).
$BR
$BR
\( L \) = \{ $multians->ans_array \}
$BR
\( U \) = \{ $multians->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
