## DESCRIPTION
## Find a unit LU decomposition of a 2x2 matrix.
##    The matrix and the solution will have small integer entries.
##    If the given matrix is not invertible, it might have multiple unit LU decompositions. All of them are accepted.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Matrix factorizations')
## DBsection('LU factorization')
## Date('March 11, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('2.D: The LU Decomposition')
## Problem1('')
## KEYWORDS('linear algebra','matrix','LU decomposition')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

$l1 = list_random(-3,-2,-1,1,2,3,4,5);
$L = Matrix([[1,0],[$l1,1]]);
    
$u1 = list_random(-2,-1,1,2,3,4);
$u2 = list_random(-2,-1,1,2,3,4);
$u3 = random(-2,4);
$U = Matrix([[$u1,$u2],[0,$u3]]);
    
$A = $L * $U;

$multians = MultiAnswer($L,$U)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if($s0->element(1,1) == 1 && $s0->element(2,2) == 1 && $s0->element(1,2) == 0 && $s1->element(2,1) == 0 && $A == $s0*$s1) {
            return 1;
        } else {
            return 0;
        }
    }
);


##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a unit LU Decomposition of the matrix \( \displaystyle A = $A \). That is, find a unit lower triangular matrix \( L \) and an upper triangular matrix \( U \) such that \( A = LU \).
$BR
$BR
\( L \) = \{ $multians->ans_array \}
$BR
\( U \) = \{ $multians->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
