##DESCRIPTION
##   Algebra: Lines
##ENDDESCRIPTION

##Tagged by up2p

## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Equations of lines: standard form)
## Date(6/3/2002)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('2.4')
## Problem1('19')
## KEYWORDS('Algebra', 'Lines')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "parserImplicitPlane.pl",
  "PGcourse.pl"
);

#Context('');
Context()->variables->are(x=>'Real',y=>'Real');

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = random(-10,10,1);
$b = random(-10,10,1);
do {
	$c = random(-10,10,1);
	$d = random(-10,10,1);
} until ($a != $c or $b != $d);

$A = $b-$d;
$B = -($a-$c);
$C = $d*($a-$c) - $c*($b-$d);
# pretty up the correct answer
$GCD = gcd(gcd($A, $B), $C);
$A = $A/$GCD;
$B = $B/$GCD;
$C = $C/$GCD;
if($A<0) {
	$A = -$A;
	$B = -$B;
	$C = -$C;
}
$mp = MultiAnswer($A, $B, $C)->with(
	singleResult => 1,
	checker => sub {
		my ($correct,$student,$self) = @_;  # get the parameters
		my ($sa, $sb, $sc) = @{$student};
		my ($ca, $cb, $cc) = @{$correct};
		Value::Error("If all coefficients are 0, the equation does not define a line") if($sa == 0 and $sb == 0 and $sc == 0);
		return ImplicitPlane("$ca x + $cb y + $cc = 0")== ImplicitPlane("$sa x + $sb y + $sc = 0");
	},
);

BEGIN_TEXT
The equation of the line that goes through the points
\( ( $a ,$b ) \) and \( ( $c ,$d ) \) can be written in general form
\( A x + B y + C = 0 \) where
$BR 
$BR 
\(A=\) \{$mp->ans_rule(5) \}
$BR
\(B=\) \{$mp->ans_rule(5) \}
$BR
\(C=\) \{$mp->ans_rule(5) \}
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();        # This should be the last executable line in the problem.

