# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Douglas Quinney, <d.a.quinney@maths.keele.ac.uk>
# revised by Gavin LaRose, <glarose@umich.edu> 11/2008
# ENDDESCRIPTION

## KEYWORDS('calculus','derivative','rate of change','velocity','average velocity','instantaneous velocity')
## Tagged by glr 11/19/08

## DBsubject('Calculus')
## DBchapter('Limits and Derivatives')
## DBsection('Derivatives and Rates of Change')
## Date('')
## Author('Gavin LaRose');
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('4')
## AuthorText1('Hughes-Hallett')
## Section1('2.1')
## Problem1('17')

## Textbook tags
## HHChapter('Key Concept: The Derivative')
## HHSection('How Do We Measure Speed')

DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
"PG.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGauxiliaryFunctions.pl",
"MathObjects.pl",
"parserNumberWithUnits.pl",
);

Context("Numeric");

## Do NOT show partial correct answers
$showPartialCorrectAnswers = 0;
TEXT(beginproblem());

# include javascript for pop-up explaining units
HEADER_TEXT(<<EOF);
<script language="javascript" type="text/javascript">
<!-- //
function unitspopup() {
    var url = "/webwork2_files/units.html";
    var opt = "height=625,width=600,location=no,menubar=no,status=no" +
              "resizable=yes,scrollbars=yes,toolbar=no,";
    window.open(url,'examdata_info',opt).focus();
}
// -->
</script>
EOF

$a1 = random(0.3,0.5,0.05);
$a2 = random(1.3,1.9,0.1);
$b  = ($a1-0)/0.2;
$c  = ($a2-$a1)/0.2;
$v  = ($b+$c)/2;

Context()->texStrings;
BEGIN_TEXT

Consider a car whose position, \( s\), is given by the table
$PAR
$BCENTER
\{begintable(7) \}
\{row("\(t\) (s)" ,0,0.2,0.4,0.6,0.8,1.0) \}
\{row("\(s\) (ft)",0,$a1,$a2,3.8,6.5,9.6) \}
\{endtable()\}
$ECENTER

$PAR
Find the average velocity over the interval \(0 \le t \le 0.2\).
$BR
average velocity = \{ ans_rule(8) \}
(include \{ htmlLink("javascript:unitspopup()","units") \})
$PAR
Estimate the velocity at \(t=0.2\).
$BR
velocity = \{ ans_rule(8) \}
(include \{ htmlLink("javascript:unitspopup()","units") \})

END_TEXT
Context()->normalStrings;

ANS( NumberWithUnits( $b, "ft/s" )->cmp() );

$bval = NumberWithUnits( $b, "ft/s" );
$cval = NumberWithUnits( $c, "ft/s" );
$vval = NumberWithUnits( $v, "ft/s" );
ANS( NumberWithUnits( $v, "ft/s" )->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( $student == $correct || $student == $cval || $student == $bval ) {
	return 1;
    } else {
	return 0;
    } } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
The average velocity over the interval \(0 \le t \le 0.2\) is given by

$PAR
\(\frac{s(0.2)-s(0)}{0.2}=\frac{$a1}{0.2}=$b \) ft/s.

$PAR
The average velocity over the interval \(0.2 \le t \le 0.4\) is given by

$PAR
\(\frac{s(0.4)-s(0.2)}{0.2}=\frac{$a2-$a1}{0.2}=$c \) ft/s.


$PAR

A reasonable estimate for the velocity at \(t=0.2\) is the average
\(\frac{1}{2}\cdot ($b+$c) = $v\) ft/s.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();

