##DESCRIPTION
##  College Algebra Review: Factoring quadratics with leading term 1
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring trinomials)
## Date(Spring 2012)
## Institution(University of Missouri)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.4')
## Problem1('')
## KEYWORDS('algebra', 'factoring', 'quadratic')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(x=>"Real");

do {
$a = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6,7,8,9,10);
do {$b = list_random(-8,-6,-5,-4,-3,-2,2,3,4,5,6,8)} until ($b != $a && $b != -$a);
$e2 = $a + $b;
$e3 = $a *$b;
} until ($e2 > 0 && $e3 > 0);

$fac1 = Formula(" x + $a ");
$fac2 = Formula(" x + $b  ")->reduce;
$texanswer = "\(\left(x + $a\right) \left(x + $b\right)\)";
$expression = "x^2 + $e2 x + $e3";
if ($e2 == 1) {
  $expression = "x^2 + x + $e3";
} elsif ($e2 == - 1) {
  $expression = "x^2 - x + $e3";
}

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your 
answer as much as possible:
$PAR
${BCENTER}
\($expression  = \) 
\( \big( \)
\{ $multians->ans_rule(25) \}
\( \big) \big( \)
\{ $multians->ans_rule(25) \}
\( \big) \)
${ECENTER}
END_TEXT

Context()->normalStrings;


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

COMMENT('MathObject version');
ENDDOCUMENT();
