## DBsubject('Calculus')
## DBchapter('Applications of Differentiation')
## DBsection('Antiderivatives')
## KEYWORDS('antiderivatives')
## TitleText1('Calculus')
## EditionText1('5e')
## AuthorText1('Stewart')
## Section1('4.10')
## Problem1('1')
## Author('')
## Institution('UCSB')


DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
);

TEXT(&beginproblem);
$showPartialCorrectAnswers = 1;
$a=random(1,10,1)*random(-1,1,2);
$b=random(1,10,1)*random(-1,1,2);
$c=random(1,10,1)*random(-1,1,2);

Context("Numeric");
Context()->variables->add(C=>'Real');

$funct = Formula("$a x^2 + $b x + $c");
$ans = Formula("($a)*(x^3)/3+$b*x^2/2+$c*x+C");


# Subtract off student answer from correct one, take the derivative w.r.t. x, if zero, then correct; if not, then incorrect.

BEGIN_TEXT

$PAR
Find the most general antiderivative of \(f(x) = $a x^2 + $b x + $c.\) 
$PAR
Note: Any arbitrary constants used must be an upper-case "C".

$PAR
\(F(x) =\) \{ans_rule(50)\}

END_TEXT

ANS( $ans->cmp( checker=>sub {
my ( $correct, $student, $ansHash ) = @_;

$checkStuDeriv = $student->D('x');

$checkforC = $student->D('C');

if (($checkStuDeriv==$funct)&&($checkforC != 0)) {
   return 1; }
elsif (($checkStuDeriv==$funct)&&($checkforC == 0)) {
   Value->Error("Don't forget + C");
   return 0; }
else {return 0;}
} ) );


ENDDOCUMENT();
