## DESCRIPTION
## Laplace transforms: partial fractions
## ENDDESCRIPTION

## KEYWORDS('Laplace transform')

## DBsubject('Differential Equations')
## DBchapter('Laplace Transforms')
## DBsection('Partial Fractions')
## Date('')
## Author('')
## Institution('Fort Lewis College')
## TitleText1('Elementary Differential Equations')
## EditionText1('2006')
## AuthorText1('Kohler and Johnson')
## Section1('5.3')
## Problem1('20')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"PGunion.pl",
"answerHints.pl",
"parserAssignment.pl",
"contextInequalities.pl",
"parserFunction.pl",
"parserMultiAnswer.pl"
);

TEXT(beginproblem());


#############################
#  Setup1

Context("Numeric");
Context()->variables->are(
t=>"Real",
s=>"Real"
);

# Redefine the Heaviside function so the answer checker will take either 1 or h(t).
parserFunction("h(t)" => "1");

# Make Y(s) a function for answer checking.
parserFunction("Y(s)" => "1.5 * sin(e*s) + 5*pi/3 + arctan(s)");

$a = random(2,4,1);
$D = random(1,5,1);
$C = $a*$D;
$B = -$D;
$r = $C*$a;

$y0 = random(2,9,1);

$LHS = Formula("s*Y(s) - $y0 + $a*Y(s)")->reduce;
$RHS = Formula("$r/s^2")->reduce;

$multians = MultiAnswer($LHS, $RHS)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);

$Ys = Formula("$r/(s^2*(s+$a)) + $y0/(s+$a)")->reduce;
$yt = Formula("$B + $C*t + ($D+$y0)*e^(-$a*t)")->reduce;


#############################
#  Main text1

Context()->texStrings;
BEGIN_TEXT

Consider the initial value problem
\[ y' + $a y = $r t, \quad \quad y(0) = $y0. \]

\{ BeginList('OL', type=>'a') \}

$ITEM Take the Laplace transform of both sides of the given differential equation to create the corresponding algebraic equation.  Denote the Laplace transform of \( y(t) \) by \( Y(s) \).  Do not move any terms from one side of the equation to the other (until you get to part (b) below).
$BR
$BR
\{ $multians->ans_rule(40) \}
\( = \)
\{ $multians->ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}

$ITEMSEP
$ITEM Solve your equation for \( Y(s) \).
$BR
$BR
\( \displaystyle Y(s) = {\mathcal L}$LBRACE y(t) $RBRACE = \)
\{ ans_rule(60) \}

$ITEMSEP
$ITEM Take the inverse Laplace transform of both sides of the previous equation to solve for \( y(t) \).
$BR
$BR
\( y(t) = \)
\{ ans_rule(60) \}

\{ EndList('OL', type=>'a') \}
END_TEXT
Context()->normalStrings;


##############################
#  Answer evaluation1

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );
ANS( $Ys->cmp() );
ANS( $yt->cmp() );

ENDDOCUMENT();