## DESCRIPTION
## Laplace transforms: partial fractions
## ENDDESCRIPTION

## KEYWORDS('Laplace transform')

## DBsubject('Differential Equations')
## DBchapter('Laplace Transforms')
## DBsection('Partial Fractions')
## Date('')
## Author('')
## Institution('Fort Lewis College')
## TitleText1('Elementary Differential Equations')
## EditionText1('2006')
## AuthorText1('Kohler and Johnson')
## Section1('5.3')
## Problem1('12')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"PGunion.pl",
"answerHints.pl",
"parserAssignment.pl",
"parserMultiAnswer.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
s=>"Real",
t=>"Real"
);

$a = random(1,4,1);
do { $b = random(1,4,1); } until ($b != $a);
$aplusb = $a+$b;
$C = random(1,5,1);
$D = random(1,5,1);

$Fs = Compute("(($C+$D)*s + (-$C*$b - $D*$a))/(s^2 - $aplusb*s + $a*$b)")->reduce; 
$fac1 = Compute("$C/(s-$a)");
$fac2 = Compute("$D/(s-$b)");
$ft = Compute("$C*e^($a*t) + $D*e^($b*t)");

$mp = MultiAnswer($fac1, $fac2)->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	my ( $f1stu, $f2stu ) = @{$student};
	my ( $f1, $f2 ) = @{$correct};
	if ( ( $f1 == $f1stu && $f2 == $f2stu ) ||
	     ( $f1 == $f2stu && $f2 == $f1stu ) ) {
	    return [1,1];
	} else {
	    if ( $f1 == $f1stu || $f2 == $f1stu ) {
 		$self->setMessage(1,"correct");
 		$self->setMessage(2,"incorrect");
		return [1,0];
	    } elsif ( $f1 == $f2stu || $f2 == $f2stu ) {
 		$self->setMessage(1,"incorrect");
 		$self->setMessage(2,"correct");
		return [0,1];
	    } else {
		return [0,0];
	    }
	}
    }
);


#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Consider the function \( \displaystyle F(s) = $Fs \).

\{ BeginList('OL', type=>'a') \}

$ITEM Find the partial fraction decomposition of \( F(s) \):
$BR
$BR
\( \displaystyle $Fs = \) \{ $mp->ans_rule(20) \} + \{ $mp->ans_rule(20) \}

$ITEMSEP
$ITEM  Find the inverse Laplace transform of \( \displaystyle F(s) \).
$BR
$BR
\( \displaystyle f(t) = {\mathcal L}^{-1} \left$LBRACE F(s) \right$RBRACE = \)
\{ ans_rule(50) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL', type=>'a') \}
END_TEXT
Context()->normalStrings;


##############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $mp->cmp() );
ANS( $ft->cmp() );

ENDDOCUMENT();