## DBsubject('Statistics')
## DBchapter('Hypothesis Testing')
## DBsection('Independent samples - t')
## Level(4)
## Date('2014/01/27')
## Author('Jonathan Baik')
## Institution('University of British Columbia')
## TitleText1('No Text')
## EditionText1('?')
## AuthorText1('?')
## Section1('?')
## Problem1('?')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "RserveClient.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context->("Numeric");

# Problem text
$q_text = 
"
How does evidence of economic recession influence an individual's
attitudes toward redistribution of wealth? Is it the case that during a
period of economic crisis, attitudes toward redistribution are influenced by
how attractive one is to potential mates? White$BITALIC et al.$EITALIC (2013)
conducted a series of studies to investigate these and related questions. In
their study 2a, subjects were split at random into two groups. One, the
control group, read an article about a person organizing their personal
office space. The other group read an article of the same length about the
economic crisis and how it had affected the life of a man who had
unexpectedly lost his job. Afterwards, participants responded to
questionnaires about their political attitudes, level of attractiveness to
potential mates (via the Landolt Mate-Value Scale), and socioeconomic
status. The study on political attitudes had a particular focus on views
relating to$BITALIC wealth redistribution,$EITALIC and responses were scored on a
scale on which a higher score is consistent with views favoring more
progressive wealth redistribution and social policies. Via the Mate-Value
Scale, responders of each sex could be categorised as of 'low mate-value' 
(so relatively unattractive to potential mates) or 'high mate-value' (relatively
attractive to potential mates). We focus here only on results for men
classified as being of low mate-value. 
$BR $BR
White, A. E., Kenrick, D. T., Neel, R. and Neuberg, S.L. (2013):
'From the Bedroom to the Budget Deficit: Mate Competition
Changes Men's Attitudes Toward Economic Redistribution.' 
$BITALIC Journal of Personality and Social Psychology$EITALIC,$BBOLD 105$EBOLD, No. 6,
924-940.";

# Part a
$qu1 = "Responses of the two experimental groups for the men of low
mate-value are displayed in the article via bar graphs. Which of the
following would have been a preferable method to display the two data sets?";
$ans1 = "Side-by-side boxplots";
$fake1a = "A histogram";
$fake2a = "A stem-and-leaf plot";
$fake3a = "A pie chart";
$fake4a = "Two pie charts";

$mc1 = new_multiple_choice();
$mc1->qa(
  $qu1,
  $ans1
);
$mc1->extra(
$fake1a, $fake2a, $fake3a, $fake4a
);

# Part b
$qu2 = "For the subset of the data under consideration here, which best
describes a sensible null hypothesis of interest?";
$ans2 = 
"For men of relatively low mate-value, awareness of economic crisis makes no difference to their mean attitudes towards wealth redistribution as
measured by the questionnaire.";
$fake1b = "For men of relatively low mate-value, awareness of economic crisis
increases their mean attitudes towards wealth redistribution as measured by
the questionnaire.";
$fake2b = "For men of relatively low mate-value, awareness of economic crisis
decreases their mean attitudes towards wealth redistribution as measured by
the questionnaire.";
$fake3b = "An increased awareness of economic crisis does not change the mean
mate-value rating of men as measured by the Landolt Mate-Value Scale.";
$fake4b = "An increased disposition toward wealth redistribution does not change
the mean mate-value rating of a man as measured by the Landolt Mate-Value
Scale.";

$mc2 = new_multiple_choice();
$mc2->qa(
  $qu2,
  $ans2
);
$mc2->extra(
$fake1b, $fake2b, $fake3b, $fake4b
);

# Part c
$qu3 = "Read the first column of p.927 of the article. Based on what the
authors write, do you think they should adopt a one--sided or two--sided
alternative hypothesis?";
$ans3 = "One-sided";
$fake1c = "Two-sided";
$fake2c = "It does not matter";

$mc3 = new_multiple_choice();
$mc3->qa(
  $qu3,
  $ans3
);
$mc3->extra(
$fake1c, $fake2c
);

# Part d

# Let us set some stuff up for this question in R

########################################################
# START R
########################################################
rserve_start();

## Randomly generate data
rserve_eval("x <- rnorm(11, mean=4.3, sd=1.3)");
rserve_eval("y <- rnorm(12, mean=5.4, sd=1.4)");

# Calculate values for the table
rserve_eval("xbar <- round(mean(x), 2)");
rserve_eval("ybar <- round(mean(y), 2)");
rserve_eval("xsd <- round(sd(x), 2)");
rserve_eval("ysd <- round(sd(y), 2)");

# Perform t test
# Do it like the student would to avoid rounding errors
rserve_eval("vpool <- ( (11-1)*xsd^2 + (12-1)*ysd^2 )/( (11-1) + (12-1) )");
rserve_eval("test_stat <- (xbar - ybar)/(sqrt((vpool/11) + (vpool/12)))");

# Get the P-value
rserve_eval("p_value <- pt(test_stat, df=21)");
# According to how t.test is calculating the p value, using (n1-1) + (n2-1) df

# Now assign our values so that webwork can access them
$xbar = join ", ", rserve_eval("xbar");
$ybar = join ", ", rserve_eval("ybar");
$xsd = join ", ", rserve_eval("xsd");
$ysd = join ", ", rserve_eval("ysd");
$poolsd = join ", ", rserve_eval("sqrt(vpool)");
#$ans_d = join ", ", rserve_eval("test_stat");
$ans_d = join ", ", rserve_eval("test_stat<- round(test_stat,2)");
$ans_e = join ", ", rserve_eval("p_value");

# Round these last numbers to 3 decimal places, to look nice
$ans_d_round = sprintf("%.3f", $ans_d);
$ans_e_round = sprintf("%.4f", $ans_e);

rserve_finish();
########################################################
# END R
########################################################

$part_d_text = "
A test that could reasonably be applied to the data here is the$BITALIC
two-sample t-test.$EITALIC This tests the null hypothesis that the means of two
Normal distributions, \(\mu _{1}\) and \(\mu _{2}\) say, are equal. The test
takes samples of sizes \(n_{1}\) and \(n_{2}\) respectively from the two
distributions, and computes their sample means, \(\bar{x}_{1}\) and \(\bar{x}_{2}\), 
and standard deviations, \(s_{1}\) and \(s_{2}\), say. There are variants
of the test, but often the test statistic is taken as
$BR $BR
$BCENTER
\[
t = \frac{\bar{x}_{1}-\bar{x}_{2}} {\sqrt{\frac{s_{p}^{2}}{n_{1}} + \frac{s_{p}^{2}}{n_{2}}}}
\]
$ECENTER
$BR $BR
in which 
$BR $BR
$BCENTER
\[
s_{p}^{2} = \frac{ \left( n_{1}-1 \right) s_{1}^{2} + \left( n_{2}-1 \right) s_{2}^{2}} {\left( n_{1}-1 \right) + \left( n_{2}-1 \right) }
\]
$ECENTER
$BR $BR
Under the null hypothesis, the test statistic \(t\) will follow a \(t\)
distribution on \(n_{1}+n_{2}-2\) degrees of freedom. 
$BR $BR
Suppose the researchers collect data with the following summary statistics:
$BR $BR
$BCENTER
\[
\begin{array}{c|ccc}
& \text{Sample size} & \text{Mean} & \text{SD} \\ \hline
\text{Control group} & 11 & $xbar & $xsd \\ 
\text{'Bad economy' group} & 12 & $ybar & $ysd
\end{array}
\]
$ECENTER
$BR $BR
Compute the test statistic for the two-sample t-test here, giving your
answer to two decimal places.
";

# Part e
$part_e_text = "
The P-value (or p-value) of a test statistic is the
probability of observing a value of the test statistic$BITALIC at least as
inconsistent$EITALIC with the null hypothesis as the value observed, assuming the
null hypothesis is true. Find the P-value of the test statistic you found
above, giving your answer to three decimal places.
";

##############################################################
#
#  Text
#
#

Context()->texStrings;

BEGIN_TEXT

$q_text

$BR
$BR

$BBOLD Part a) $EBOLD
$BR
\{ $mc1->print_q() \}
$BR
\{ $mc1->print_a() \}

$BR
$BR

$BBOLD Part b) $EBOLD
$BR
\{ $mc2->print_q() \}
$BR
\{ $mc2->print_a() \}

$BR
$BR

$BBOLD Part c) $EBOLD
$BR
\{ $mc3->print_q() \}
$BR
\{ $mc3->print_a() \}

$BR
$BR

$BBOLD Part d) $EBOLD
$BR
$part_d_text
$BR
\{ ans_rule(12) \} 
$BR
$BR

$BBOLD Part e) $EBOLD
$BR
$part_e_text
$BR
\{ ans_rule(12) \} 
$BR
$BR

END_TEXT  

#######

$showPartialCorrectAnswers = 0;


ANS( radio_cmp($mc1->correct_ans()) );
ANS( radio_cmp($mc2->correct_ans()) );
ANS( radio_cmp($mc3->correct_ans()) );
ANS( Real($ans_d)->cmp( tol=> 0.01, tolType=>"absolute", checker=>sub {
        my ($correct, $student, $ans) = @_;
       undef($ans->{checker});
        return abs($correct)->cmp_compare(abs($student), $ans)
}) );

ANS( num_cmp($ans_e, tol=> 0.001, tolType=>"absolute") );

##############################################################
#
#  Solution
#

Context()->texStrings;
Context()->{format}{number} = "%.2f";

BEGIN_SOLUTION

$BBOLD Part a) $EBOLD
$BR $BR
$ans1
$BR $BR

$BBOLD Part b) $EBOLD
$BR $BR
$ans2
$BR $BR

$BBOLD Part c) $EBOLD
$BR $BR
One-sided seems appropriate for the researchers, as they suggest "We predicted that the economic threat manipulation would
lead low mate-value men to adopt more pro-social orientation...", the view expressed apparently based on previous research.
$BR $BR

$BBOLD Part d) $EBOLD
$BR $BR
The test statistic can be calculated with the equations described in the question. First, we calculate the pooled standard deviation,
$BR $BR
\[
\begin{align*}
s_{p}^{2} &= \frac{ \left( n_{1}-1 \right) s_{1}^{2} + \left( n_{2}-1 \right) s_{2}^{2}} {\left( n_{1}-1 \right) + \left( n_{2}-1 \right) } \\
&= \frac{ \left( 11-1 \right) $xsd^{2} + \left( 12-1 \right) $ysd^{2}} {\left( 11-1 \right) + \left( 12-1 \right) }
\end{align*}
\]
$BR $BR

and then the test statistic:

$BR $BR
\[
\begin{align*}
t &= \frac{\bar{x}_{1}-\bar{x}_{2}} {\sqrt{\frac{s_{p}^{2}}{n_{1}} + \frac{s_{p}^{2}}{n_{2}}}} \\
&= \frac{$xbar-$ybar} {\sqrt{\frac{s_p^2}{11} + \frac{s_p^2}{12}}} \\
&= $ans_d_round
\end{align*}
\]
$BR $BR

$BBOLD Part e) $EBOLD
$BR $BR
The P-value can be found using R's $BBOLD pt $EBOLD function with $(n_1-1) + (n_2-1)$ degrees of freedom (this is degrees of freedom used by the $BBOLD t.test $EBOLD function in R).
$BR $BR

$BBOLD
pt($ans_d_round, df=(11-1 + 12-1))
$EBOLD

$BR $BR
The P-value is approximately $ans_e_round .

END_SOLUTION

ENDDOCUMENT();
