## DESCRIPTION
##   Estimating partial derivatives
## ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Partial derivatives)
## Date(10/15/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(2varapp)
## Level(5)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('14.1')
## Problem1('13 15')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('14.1')
## Problem2('13 15')
## KEYWORDS('Derivative', 'Partial')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "niceTables.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(t=>'Real',w=>'Real');

# t = temperature
# w = wind speed
# http://www.nws.noaa.gov/om/windchill/
# f(t,w) = wind chill

$f = Formula("35.74 + 0.6215 t - 35.75 w^(0.16) + 0.4275 t w^(0.16)");

@windspeed = (5,10,15,20,25);

$starttemp = random(15,35,5);
for $i (0..7) {
@temp[$i] = $starttemp - 5*$i;
}  

foreach $j (0..scalar(@temp)-1) {
foreach $i (0..scalar(@windspeed)-1) {
$windchill[$i][$j] =  round($f->eval(w=>$windspeed[$i],t=>$temp[$j])->reduce);
}
}

$rA = random(1,3,1); # row index for part A
$cA = random(1,6,1); # column index for part A
$windspeedrA = Real("$windspeed[$rA]");  
$tempcA = Real("$temp[$cA]"); 

# slopes of two secant lines approximating f_w(a,b)
$fw1 = Real("($windchill[$rA+1][$cA] - $windchill[$rA][$cA])/5");
$fw2 = Real("($windchill[$rA][$cA] - $windchill[$rA-1][$cA])/5");

$fw1neg = Real("-($fw1)");
$fw2neg = Real("-($fw2)");


$rB = random(1,3,1); # row index for part B
while ($rA == $rB) { $rB = random(1,3,1); }
$cB = random(1,6,1); # column index for part B
$windspeedrB = Real("$windspeed[$rB]");  
$tempcB = Real("$temp[$cB]"); 

# slopes of two secant lines approximating f_T(a,b)
$fT1 = Real("($windchill[$rB][$cB-1] - $windchill[$rB][$cB])/5");
$fT2 = Real("($windchill[$rB][$cB] - $windchill[$rB][$cB+1])/5");



##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

The temperature adjusted for wind chill is a temperature that tells you how cold it feels, as a result of the combination of wind speed, w, in miles per hour, and actual temperature, T, as measured by a thermometer in degrees Fahrenheit.  The table below shows the temperature adjusted for wind chill, f(w,T), as a function of w and T.  The left column of the table gives the wind speed in miles per hour, the top row gives the actual temperature in degrees Fahrenheit, and the remaining entries give the temperature adjusted for wind chill in degrees Fahrenheit as a function of wind speed and actual temperature.
$PAR
\{
DataTable(
 [ ["w\T",$temp[0],$temp[1],$temp[2],$temp[3],$temp[4],$temp[5],$temp[6],$temp[7]],
   [$windspeed[0],$windchill[0][0],$windchill[0][1],$windchill[0][2],$windchill[0][3],$windchill[0][4],$windchill[0][5],$windchill[0][6],$windchill[0][7]],
   [$windspeed[1],$windchill[1][0],$windchill[1][1],$windchill[1][2],$windchill[1][3],$windchill[1][4],$windchill[1][5],$windchill[1][6],$windchill[1][7]],
   [$windspeed[2],$windchill[2][0],$windchill[2][1],$windchill[2][2],$windchill[2][3],$windchill[2][4],$windchill[2][5],$windchill[2][6],$windchill[2][7]],
   [$windspeed[3],$windchill[3][0],$windchill[3][1],$windchill[3][2],$windchill[3][3],$windchill[3][4],$windchill[3][5],$windchill[3][6],$windchill[3][7]],
   [$windspeed[4],$windchill[4][0],$windchill[4][1],$windchill[4][2],$windchill[4][3],$windchill[4][4],$windchill[4][5],$windchill[4][6],$windchill[4][7]],
 ],
 midrules => 1,
 caption => "Table of windchills.",
 align => 'c|c|c|c|c|c|c|c|c'
);
\}
$PAR
(a)  Estimate \( f_w \, ($windspeed[$rA], $temp[$cA]) \) and give a practical interpretation.
\( f_w \, ($windspeed[$rA], $temp[$cA]) \approx \) \{ans_rule(30)\}
$BR$BR
If the wind is blowing \{ans_rule(10)\} mph and the actual temperature is \{ans_rule(10)\} degrees Fahrenheit,
it will feel \{ans_rule(10)\} degrees Fahrenheit $BITALIC colder $EITALIC for each 1 mph increase in wind speed. 
$BR
$BR
(b)  Estimate \( f_T \, ($windspeed[$rB], $temp[$cB]) \) and give a practical interpretation.
\( f_T \, ($windspeed[$rB], $temp[$cB]) \approx \) \{ans_rule(30)\}
$BR$BR
If the wind is blowing \{ans_rule(10)\} mph and the actual temperature is \{ans_rule(10)\} degrees Fahrenheit,
it will feel \{ans_rule(10)\} degrees Fahrenheit $BITALIC warmer $EITALIC for each 1 degree Fahrenheit increase in actual temperature. 

END_TEXT
Context()->normalStrings;

##################################################
#  Answer Checkers

$showPartialCorrectAnswers=1;

ANS($fw1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
#return $student == $correct || $student == $fw2;
if (
 (($correct <= $student) && ($student <= $fw2)) ||
 (($fw2 <= $student) && ($student <= $correct))
) 
{ return 1; } else { return 0; }
}));
ANS($windspeedrA->cmp());
ANS($tempcA->cmp());
ANS($fw1neg->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
#return $student == $correct || $student == $fw2neg;
if (
 (($correct <= $student) && ($student <= $fw2neg)) ||
 (($fw2neg <= $student) && ($student <= $correct))
) 
{ return 1; } else { return 0; }
}));


ANS($fT1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
#return $student == $correct || $student == $fT2;
if (
 (($correct <= $student) && ($student <= $fT2)) ||
 (($fT2 <= $student) && ($student <= $correct))
) 
{ return 1; } else { return 0; }
}));
ANS($windspeedrB->cmp());
ANS($tempcB->cmp());
ANS($fT1->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
#return $student == $correct || $student == $fT2;
if (
 (($correct <= $student) && ($student <= $fT2)) ||
 (($fT2 <= $student) && ($student <= $correct))
) 
{ return 1; } else { return 0; }
}));



;
ENDDOCUMENT();
