## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(4)

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
install_problem_grader(~~&std_problem_grader);

Context('Matrix');

$P = GL2Z();
$Pinv = $P->inverse;

$e1 = random(2,3,1);
$e2 = random(-2,-1,1);

$D = Matrix([
[$e1,0],
[0,$e2],
]);

$n = random(5,8,1);

$Dn = Matrix([
[($e1)**($n),0],
[0,($e2)**($n)],
]);

$Dnalt = Matrix([
[($e2)**($n),0],
[0,($e1)**($n)],
]);

$A = $P * $D * ($Pinv);

$An = $P * $Dn * $Pinv;

$multians = MultiAnswer($P, $Dn, $Pinv)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = (0,0,0); # (0) x 4; # initialize with zeros
      
      $score[1] = 1 if (($s[1] == $Dn) or ($s[1] == $Dnalt));
      if (($s[0]->det != 0) and ($s[2]->det != 0)) {
        $score[0] = 1 if (($An == $s[0] * $s[1] * $s[2]) and ($score[1] == 1));
        $score[2] = 1 if (($An == $s[0] * $s[1] * $s[2]) and ($score[1] == 1));
      }

      $self->setMessage(1,"P and P^(-1) are marked correct or incorrect simultaneously") if ($score[0] == 0);
      $self->setMessage(2,"Since D was incorrect, P and P^(-1) were not checked for correctness") if ($score[1] == 0);
      $self->setMessage(3,"P and P^(-1) are marked correct or incorrect simultaneously") if ($score[2] == 0);

      return @score;
  }
);

Context()->texStrings;
BEGIN_TEXT
Suppose \[ A = $A. \]
Find an invertible matrix \( P \) and a diagonal matrix \( D \) so that \( A = P D P^{-1} \).
Use your answer to find an expression for \( A^{$n} \) in terms of \( P \), a power of \( D \), and \( P^{-1} \) in that order.
$BR
$BR
\( A^{$n} = \) 
\{ $multians->ans_array(5) \}
\{ $multians->ans_array(5) \}
\{ $multians->ans_array(5) \}.
END_TEXT
Context()->normalStrings;

ANS( $multians->cmp );

COMMENT('Always diagonalizable.');

ENDDOCUMENT();
