## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(4)

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
install_problem_grader(~~&std_problem_grader);

Context('Matrix');

$P = GL2Z();
$Pinv = $P->inverse;

$D = Matrix([
[random(2,3,1),0],
[0,random(-2,-1,1)],
]);

$e1 = $D->element(1,1);
$e2 = $D->element(2,2);

$Dalt = Matrix([[$e2,0],[0,$e1]]);

$v1 = $P->column(1);
$v2 = $P->column(2);

$A = $P * $D * ($Pinv);

$multians = MultiAnswer($A, $P, $D, $Pinv)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = (0,0,0,0); # (0) x 4; # initialize with zeros
      
      $score[0] = 1 if ($s[0] == $c[0]);
      $score[2] = 1 if (($s[2] == $c[2]) or ($s[2] == $Dalt));
      if (($s[1]->det != 0) and ($s[3]->det != 0)) {
        $score[1] = 1 if (($c[0] == $s[1] * $s[2] * $s[3]) and ($score[2] == 1));
        $score[3] = 1 if (($c[0] == $s[1] * $s[2] * $s[3]) and ($score[2] == 1));
      }

      $self->setMessage(2,"P and P^(-1) are marked correct or incorrect simultaneously") if ($score[1] == 0);
      $self->setMessage(3,"Since D was incorrect, P and P^(-1) were not checked for correctness") if ($score[2] == 0);
      $self->setMessage(4,"P and P^(-1) are marked correct or incorrect simultaneously") if ($score[3] == 0);

      return @score;
  }
);

Context()->texStrings;
BEGIN_TEXT
Suppose that a \( 2 \times 2 \) matrix \( A \) has
an eigenvalue \( $e1 \) with corresponding eigenvector \( $v1 \) and 
an eigenvalue \( $e2 \) with corresponding eigenvector \( $v2 \).
$BR
$BR
Find an invertible matrix \( P \) and a diagonal matrix \( D \) so that \( A = P D P^{-1} \).
Enter your answer as an equation of the form \( A = P D P^{-1} \).
You must enter a number in every answer blank for the answer evaluator to work properly.
$BR
$BR
\{ $multians->ans_array(5) \} \( = \) 
\{ $multians->ans_array(5) \}
\{ $multians->ans_array(5) \}
\{ $multians->ans_array(5) \}
END_TEXT
Context()->normalStrings;

ANS( $multians->cmp );

COMMENT('Always diagonalizable.');

ENDDOCUMENT();
