##DESCRIPTION
##  Rational Expressions and Functions: Add and Subtract
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Simplifying expressions)
## Institution(The College of Idaho)
## Author(RA Cruz)
## MLT(Least_common_denominator)
## Level(2)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('6.2')
## Problem1('')
## KEYWORDS('rational expressions','rational functions')
## Date: 2007/11 updated 2017/10 --rac

DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

######################################
#  Setup

Context("Numeric");

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '+' => {class => 'bizarro::BOP::add', isCommand => 1},
   '+ ' => {class => 'bizarro::BOP::add', isCommand => 1},
   ' +' => {class => 'bizarro::BOP::add', isCommand => 1},);

# --- Form: a/bx^2 + c/x -------------------------------------------
#  Answer: (a + bcx)/bx^2

$a = random(2,5,1);
$b = random(2,7,1);
if (gcd($a,$b)!=1) {$b = $a + 1;}
$c = random(1,5,1);

$den1 = Formula("$b x^2")->reduce;
$den2 = Formula("x")->reduce;
$f1 = Formula("($a) / $den1")->reduce->TeX;
$f2 = Formula("($c) / $den2")->reduce->TeX;
$mult2 = Formula("$b x")->reduce;
$f2_lcd =  Formula("($c*$b) x")->reduce->TeX;

######################################
#  Main text

BEGIN_TEXT
Perform the indicated operation.  Note that the denominators
are different.  Simplify the result, if possible.
\[ $f1 + $f2 \]
$PAR
Answer:  \{ ans_rule(25) \} 
$BR $BR
\{ 
knowlLink("help(entering your answer)", 
value=>'Enter your answer  '.
'in the form: ${BITALIC}numerator/denominator$EITALIC. '.
'Use parentheses as needed.'.
'$BR Your answer should be in reduced form with integer ' .
'coefficients. For example:'.
'$BR $SPACE $SPACE If the correct, reduced answer is: 3/(2x+4)'.
'$BR $SPACE $SPACE The following will be scored as incorrect: 9/(6x+12) or 1.5/(x+2) or 1/(2/3x+4/3) or (3/2)(1/(x+2))'
) \}
END_TEXT

######################################
#  Answer

$showPartialCorrectAnswers = 1;

$errText = "Your answer is equivalent to the correct answer, but ".
                 "needs to be simplified further or is not in the expected format. $BR".
                 "See ${BBOLD}help(entering your answer)$EBOLD ".
                 "given below for the expected format.";

$answer = Formula("($a + $b*$c x)/($b x^2)")->reduce;

ANS($answer -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1,bizarroAdd=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      $student = Formula("$student"); $correct = Formula("$correct"); 
      $correctAlt = Formula("$answerAlt");
      my $OK = ($correct == $student || $correctAlt == $student);
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0,bizarroAdd=>0);
      Value::Error($errText) unless $OK;
      return $OK;
})); 


######################################
#  Solution

Context()->texStrings;
BEGIN_SOLUTION
$BR
Consider the two denominators: \($den1\) and \($den2\)
$BR
The LCD is \($b x^2\).
\[
\begin{aligned}
$f1 + $f2  & = $f1 + $f2 \cdot \frac{$mult2}{$mult2}  \\
                & =$f1 + \frac{ $f2_lcd }{$den1} \\
                & = $answer
\end{aligned}
\]
END_SOLUTION
Context()->normalStrings;ENDDOCUMENT();

ENDDOCUMENT();


























