##DESCRIPTION
##  Rational Expressions and Functions: Add and Subtract
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Institution(The College of Idaho)
## Author(RA Cruz)
## MLT(Add_fractions)
## Level(2)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('6.2')
## Problem1('')
## KEYWORDS('rational expressions','rational functions')
## Date: 2007/11 updated 2017/10 --rac

DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

######################################
#  Setup

Context("Numeric");

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '+' => {class => 'bizarro::BOP::add', isCommand => 1},
   '+ ' => {class => 'bizarro::BOP::add', isCommand => 1},
   ' +' => {class => 'bizarro::BOP::add', isCommand => 1},);

# --- Form: d + a/(bx+c) ---------------------------------------
#  Note: Don't let d factor out to give bx+c
#  Answer: (dbx+cd+a)/(bx+c)

$b = random(2,5,1);
do {$c = non_zero_random(-5,5,1);} until (gcd($b,$c)==1);
$d = random(1,4,1);
do {$a = random(1,4,1);} while (($c*$d+$a)/$d==$c);

@den = (Formula("1"), Formula("$b x + $c") );
@f = (Formula("$d")->reduce->TeX,
               Formula("($a) / $den[1]")->reduce->TeX);

$p1 = random(0,1,1);
$p2 = 1-$p1;

# --- For the alternative answer ------------------
$numFactor = gcd($d*$b, $c*$d+$a);  
$numCoef1 = ($d*$b)/$numFactor;
$numCoef2 = ($c*$d+$a)/$numFactor;

# --- For the Solution----------------------------------
$lcd = $den[1];
$mult = $den[1]->TeX;
$num[0] = Formula("($d*$b)x +($c*$d)")->reduce;
$num[1] = Formula("$a")->reduce;

######################################
#  Main text

BEGIN_TEXT
Perform the indicated operation.  Note that the denominators
are different.  Simplify the result, if possible.
\[ $f[$p1] + $f[$p2] \]
$PAR
Answer:  \{ ans_rule(25) \} 
$BR $BR
\{ 
knowlLink("help(entering your answer)", 
value=>'Enter your answer  '.
'in the form: ${BITALIC}numerator/denominator$EITALIC. '.
'Use parentheses as needed.'.
'$BR Your answer should be in reduced form with integer ' .
'coefficients. For example:'.
'$BR $SPACE $SPACE If the correct, reduced answer is: 3/(2x+4)'.
'$BR $SPACE $SPACE The following will be scored as incorrect: 9/(6x+12) or 1.5/(x+2) or 1/(2/3x+4/3) or (3/2)(1/(x+2))'
) \}
END_TEXT

######################################
#  Answer

$showPartialCorrectAnswers = 1;

$errText = "Your answer is equivalent to the correct answer, but ".
                 "needs to be simplified further or is not in the expected format. $BR".
                 "See ${BBOLD}help(entering your answer)$EBOLD ".
                 "given below for the expected format.";

$answer = Formula("(($b*$d) x + ($c*$d+$a) ) / ($b x + $c)")->reduce;

#--The alternative answer has possible factoring in numerator ------
$answerAlt = Formula("( ($d*$b)x+($d*$c-$a) ) / ($b x+$c)")->reduce;
if ($numFactor>1) {
   $answerAlt = Formula("( $numFactor($numCoef1 x+$numCoef2) ) / ($b x+$c)")->reduce;
}

ANS($answer -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1,bizarroAdd=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      $student = Formula("$student"); $correct = Formula("$correct"); 
      $correctAlt = Formula("$answerAlt");
      my $OK = ($correct == $student || $correctAlt == $student);
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0,bizarroAdd=>0);
      Value::Error($errText) unless $OK;
      return $OK;
})); 

######################################
#  Solution

$displayStep = ($p1==1) ? "\frac{$f[0]}{1} \cdot \frac{$mult}{$mult} + $f[1]" :
                                      " $f[1] + \frac{$f[0]}{1} \cdot \frac{$mult}{$mult}";

Context()->texStrings;
BEGIN_SOLUTION
$BR
The LCD is \($lcd\).
\[
\begin{aligned} 
$f[$p0] + $f[$p1] & = $displayStep \\
          & = \frac{$num[$p0]}{$lcd} + \frac{$num[$p1]}{$lcd} \\
          & = $answer               
\end{aligned}
\]
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();


























