##DESCRIPTION
##  Rational Expressions and Functions: Add and Subtract
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Simplifying)
## Institution(The College of Idaho)
## Author(RA Cruz)
## MLT(Subtract_fractions)
## Level(2)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('6.2')
## Problem1('')
## KEYWORDS('rational expressions','rational functions')
## Date: 2007/11

DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "CofIdaho_macros.pl",
  "contextRationalFunction.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

######################################
#  Setup
#  Form: ax/(x+b)(x+c) - d/(x+c) 
#  Note: b neq + or -c, 
#        (a-d) does not factor out of -bd, 
#         to give c.
#  Answer: ((a-d)x-bd)/(x+b)(x+c)

$b = random(1,5,1);
$c = non_zero_random(-4,4,1);
if ($c*$c==$b*$b) {$c = $b + random(1,3,1);}
$d = random(1,5,1);
do {$a = $d + random(1,4,1);} while ((-1*$b*$d)/($a-$d)==$c);

Context("RationalFunction");

$f1 = Formula("($a x) / (x^2+($b+$c)x+($b*$c))")->reduce->TeX;
$f2 = Formula("($d) / (x+$c)")->reduce->TeX;

######################################
#  Main text

BEGIN_TEXT
Perform the indicated operation.  Note that the denominators
are different.  Simplify the result, if possible.
\[ $f1 - $f2 \]
$PAR
Answer:  \{ ans_rule(25) \} 
END_TEXT

######################################
#  Answer

#$answer=(($a-$d)*x - $b*$d ) / ((x+$b)*(x+$c))";
#ANS(RationalExpEvaluator($answer,"x"));  

$answer=Formula("(($a-$d)*x - $b*$d ) / ((x+$b)*(x+$c))")->reduce;
ANS($answer->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
}));

$showPartialCorrectAnswers = 1;

######################################
# Solution

Context()->texStrings;
BEGIN_SOLUTION
$PAR Solution: $BR $BR
Answer: \(\displaystyle $answer\)
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();


























