##DESCRIPTION
##  Rational Expressions and Functions: Add and Subtract
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Simplifying expressions)
## Institution(The College of Idaho)
## Author(RA Cruz)
## MLT(Same_denominator)
## Level(2)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('6.2')
## Problem1('')
## KEYWORDS('rational expressions','rational functions')
## Date: 2007/11 updated 2017/09 --rac

DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

######################################
#  Setup

Context("Numeric");

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '+' => {class => 'bizarro::BOP::add', isCommand => 1},
   '+ ' => {class => 'bizarro::BOP::add', isCommand => 1},
   ' +' => {class => 'bizarro::BOP::add', isCommand => 1},);

# --- Form: (ex^2+(ce+d+f)xy+gy^2)/(ax+by)(cx+dy) - ((e-c)x^2+fxy+(g-de)y^2)/(ax+by)(cx+dy)  
#  Note:  b=1,-1, e>c, c and d rel prime, a and b rel prime
#  Answer: (x+ey)/(ax+by), 

Context()->variables->are(x=>'Real',y=>'Real');

$a = random(2,4,1);
$b = random(-1,1,2);
$c = random(2,3,1);
do {$d = non_zero_random(-5,5,1);} until (gcd($c,$d)==1);
$e = $c + random(1,5,1);
$f =  non_zero_random(-5,5,1);
$g =  non_zero_random(-5,5,1);

$den = Formula("($a*$c)x^2+($b*$c+$a*$d)xy+($b*$d)y^2")->reduce;
$f1 = Formula("($e x^2+($c*$e+$d+$f)xy+$g y^2) / $den")->reduce->TeX;
$f2 = Formula("(($e-$c)x^2+$f xy+($g-$d*$e)y^2) / $den")->reduce->TeX;
$unreduced = Formula("$c x^2+($c*$e+$d)xy+($d*$e)y^2")->reduce;
$factorNum = Formula("(x+$e y)($c x+$d y)")->reduce;
$factorDen = Formula("($a x+$b y)($c x+$d y)")->reduce;

######################################
#  Main text

BEGIN_TEXT
Perform the indicated operation.  Note that the denominators
are the same.  Simplify the result, if possible.
\[ $f1 - $f2 \]
$PAR
Answer:  \{ ans_rule(25) \} 
$BR $BR
\{ 
knowlLink("help(entering your answer)", 
value=>'Enter your answer  '.
'in the form: ${BITALIC}numerator/denominator$EITALIC. '.
'Use parentheses as needed.'.
'$BR Your answer should be in reduced form with integer ' .
'coefficients. For example:'.
'$BR $SPACE $SPACE If the correct, reduced answer is: 3/(2x+4)'.
'$BR $SPACE $SPACE The following will be scored as incorrect: 9/(6x+12) or 1.5/(x+2) or 1/(2/3x+4/3) or (3/2)(1/(x+2))'
) \}
END_TEXT

######################################
#  Answer

$showPartialCorrectAnswers = 1;

$errText = "Your answer is equivalent to the correct answer, but ".
                 "needs to be simplified further or is not in the expected format. $BR".
                 "See ${BBOLD}help(entering your answer)$EBOLD ".
                 "given below for the expected format.";

$answer = Formula("(x+$e*y)/($a*x+$b*y)")->reduce;

ANS($answer -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1,bizarroAdd=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      $student = Formula("$student"); $correct = Formula("$correct");
      my $OK = ($correct == $student);
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0,bizarroAdd=>0);
      Value::Error($errText) unless $OK;
      return $OK;
})); 

######################################
#  Solution

Context()->texStrings;
BEGIN_SOLUTION
$BR
The denominators are equal.  Add the numerators.
Factor and reduce
\[
\begin{aligned} 
    $f1 - $f2       & = \frac{$unreduced}{$den} \\
                        & =  \frac{$factorNum}{$factorDen}\\
                        & = $answer
\end{aligned}
\]
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
