# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('multivariable', 'functions', 'level surface')
## Tagged by glr 05/14/09


## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('12.5')
## Problem1('19')

## Textbook tags
## HHChapter1('Functions of Several Variables')
## HHSection1('Functions of Three Variables')

Context("Numeric");
Context()->variables->add( y=>'Real', z=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(1,8,1);
$b = random(1,8,1);
$c = random(1,8,1);
$g = Compute( "$a*x + $b*y" )->reduce();
$f = Compute("$a*x + $b*y - z + $c")->reduce();
$fn = Compute("z - $a*x - $b*y + $c")->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a function \( f(x,y,z) \) whose level surface \( f=$c \)
is the graph of the function \( g(x,y)=  $g\).

$PAR
\( f(x,y,z) = \) \{ ans_rule(45) \}

END_TEXT
Context()->normalStrings;

ANS( $f->cmp( checker=>sub { my ( $c, $s, $self ) = @_;
			     return $s == $c || $s == $fn; } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR


The graph of \(g(x,y)=$g\) is the set of all points \((x,y,z)\)
satisfying \(z=$g\), or \($g - z = 0\).  This is a level surface,
but we want the surface equal to the constant value $c, not 0, so
we can add $c to both sides to get \($g - z + $c = $c\).  Thus,
\(f(x,y,z) = $f\) has level surface \(f=$c\) identical to the graph
of \(g(x,y)=  $g\).
$PAR
Of course, we could also say that the graph of \(g(x,y)\) is
equivalent to \(z - $g = 0\), so that \(f = $fn\).


END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
