# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('multivariable', 'functions', 'level surface')
## Tagged by glr 05/14/09


## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('12.5')
## Problem1('23')

## Textbook tags
## HHChapter1('Functions of Several Variables')
## HHSection1('Functions of Three Variables')

Context("Numeric");
Context()->variables->add( y=>'Real', z=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(1,5,1);
$asq = $a*$a;
$xy = list_random('x','y');
$yx = ( $xy eq 'x' ) ? 'y' : 'x';

$f = Compute( "sqrt($asq - $xy^2)" );
$g1 = Compute( "$f - z" );
$g2 = Compute( "z - ($f)" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

The surface \( S \) is the graph of \( f(x,y) = $f \).
$PAR
Be sure you can explain why \( S \) is the upper half of a circular
cylinder of radius $a, centered along the  \( $yx \)-axis.  Write three
or four sentences on a sheet of paper that clearly demonstrate this.
$PAR
Find a level surface \( g(x,y,z)=c \) representing \( S \).
$PAR
\( g(x,y,z) = \) \{ ans_rule(45) \},
$BR
with \(c = \) \{ ans_rule(10) \}


END_TEXT
Context()->normalStrings;

ANS( $g1->cmp( checker=>sub { my ( $cor, $stu, $self ) = @_;
			      return $stu == $cor || $stu == $g2; } ) );
ANS( Compute(0)->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR


The graph of \(f(x,y)\) is obtained by plotting points \((x,y,z)\),
where \(z=f(x,y)\). Since the square root function is never negative, we
have \(z \ge 0\).  Setting \(z = \sqrt{$asq - $xy^2}\) and squaring
both sides leads to \($xy^2 + z^2 = $asq\), which is the equation for
a circular cylinder of radius $a lying along the \($yx\)-axis (since
\($yx\) is missing from the equation). The graph of the function includes
only those points where \(z \ge 0\), that is, the upper half of the
cylinder.
$PAR
If we take \(g(x,y,z) = f(x,y) - z = $f - z\), then the level surface
\(g(x,y,z) = 0\) is the surface \(S\).
$PAR
Clearly, we could also take \(g(x,y,z) = z - f(x,y) = z - $f\) as well.


END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
