# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserImplicitEquation.pl",
"parserMultiAnswer.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('multivariable', 'functions', 'level surface')
## Tagged by glr 05/14/09


## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('12.5')
## Problem1('31')

## Textbook tags
## HHChapter1('Functions of Several Variables')
## HHSection1('Functions of Three Variables')

Context("ImplicitEquation");
Context()->variables->add( z=>'Real' );
Context()->flags->set( ignoreEndpointTypes=>1 );

$showPartialCorrectAnswers = 1;

$a = random(-3,3,1);
$b = random(-3,3,1);
$a = 0;
$e = Compute( "e^(-((x-$a)^2 + (y - $b)^2 + z^2))" )->reduce();

$asq = $a*$a;
$bsq = $b*$b;

Context()->variables->set(
    x=>{limits=>[$a-2,$a+2]},
    y=>{limits=>[$b-2,$b+2]},
    z=>{limits=>[-2,2]}
);

$cxz = Compute( "1/(e^($bsq+1))" );
$xz = ImplicitEquation( "(x-$a)^2 + z^2 = 1" )->reduce();
$cyz = Compute( "1/(e^($asq+1))" );
$yz = ImplicitEquation( "(y-$b)^2 + z^2 = 1" )->reduce();
$cxy = Compute( "1/(e^(1))" );
$xy = ImplicitEquation( "(x-$a)^2 + (y-$b)^2 = 1" )->reduce();

$xzAns = MultiAnswer( $cxz, $xz )->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checkTypes => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $cstu, $eqstu ) = @{$student};
        my ( $cans, $eqans ) = @{$correct};
	my @ret = (0,0);
	$cstu = Real($cstu) if ( $cstu ne "" );
	if ( ref($cstu) eq ref($cans) && $cstu > Compute(0) &&
	     $cstu < Compute("1/e^($bsq)") ) {
	    $ret[0] = 1;
	}
	if ( ref($cstu) eq ref($cans) && $eqstu ) {
	    my $eq = ImplicitEquation( "(x-$a)^2 + z^2 = -ln($cstu) - $bsq" );
	    $eqstu = ImplicitEquation( $eqstu );
	    $ret[1] = ( $eq == $eqstu ) ? 1 : 0;
	}

	return [ @ret ];
    }
);
$yzAns = MultiAnswer( $cyz, $yz )->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checkTypes => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $cstu, $eqstu ) = @{$student};
        my ( $cans, $eqans ) = @{$correct};
	my @ret = (0,0);
	$cstu = Real($cstu) if ( $cstu ne "" );
	if ( ref($cstu) eq ref($cans) && $cstu > Compute("0") &&
	     $cstu < Compute("1/e^($asq)") ) {
	    $ret[0] = 1;
	}
	if ( ref($cstu) eq ref($cans) && $eqstu ) {
	    my $eq = ImplicitEquation( "(y-$b)^2 + z^2 = -ln($cstu) - $asq" );
	    $eqstu = ImplicitEquation( $eqstu );
	    $ret[1] = ( $eq == $eqstu ) ? 1 : 0;
	}
	return [ @ret ];
    }
);
$xyAns = MultiAnswer( $cxy, $xy )->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checkTypes => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $cstu, $eqstu ) = @{$student};
        my ( $cans, $eqans ) = @{$correct};
	my @ret = (0,0);
	$cstu = Real($cstu) if ( $cstu ne "" );
	if ( ref($cstu) eq ref($cans) && $cstu > Compute(0) &&
	     $cstu < Compute("1") ) {
	    $ret[0] = 1;
	}
	if ( ref($cstu) eq ref($cans) && $eqstu ) {
	    my $eq = ImplicitEquation( "(x-$a)^2 + (y-$b)^2 = -ln($cstu)" );
	    $eqstu = ImplicitEquation( $eqstu );
	    $ret[1] = ( $eq == $eqstu ) ? 1 : 0;
	}
	return [ @ret ];
    }
);


Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Describe in words the level surfaces of
\( g(x,y,z) = $e \).
Think what the intersections of these
surfaces with the three coordinate planes look like.
$PAR

${BBOLD}(a)$EBOLD
In symbolic form, these level surfaces are given by
\( $e = c \)
where \(c\) is a constant.  For what values of \(c\)
are level surfaces defined?
$BR
\(c \in\) \{ ans_rule(35) \}
$BR
${BITALIC}(Give your answer as an interval or list of intervals; for
example, if c is less than 1 or greater than or equal to 2, enter
${BBOLD}(-inf,1),[2,inf)$EBOLD.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
Pick a level surface that intersects the \(xz\)-plane.  What is your
value of \(c\)? $BR
\( c = \) \{ $xzAns->ans_rule(15) \}
$BR
Write the equation for the intersection:
$BR
\{ $xzAns->ans_rule(35) \}

$PAR
${BBOLD}(c)$EBOLD
Pick a level surface that intersects the \(yz\)-plane.  What is your
value of \(c\) (which might be different from that you used in
${BBOLD}(a)$EBOLD? $BR
\( c = \) \{ $yzAns->ans_rule(15) \}
$BR
Write the equation for the intersection:
$BR
\{ $yzAns->ans_rule(35) \}

$PAR
${BBOLD}(d)$EBOLD
Pick a level surface that intersects the \(xy\)-plane.  What is your
value of \(c\)? $BR
\( c = \) \{ $xyAns->ans_rule(15) \}
$BR
Write the equation for the intersection:
$BR
\{ $xyAns->ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( Interval( "(0,1]" )->cmp() );
ANS( $xzAns->cmp() );
ANS( $yzAns->cmp() );
ANS( $xyAns->cmp() );

$xzLHS = Compute( "(x - $a)^2 + z^2" )->reduce();
$yzLHS = Compute( "(y - $b)^2 + z^2" )->reduce();
$xyLHS = Compute( "(x - $a)^2 + (y - $b)^2" )->reduce();

$asqp1 = $a*$a + 1;
$bsqp1 = $b*$b + 1;

$c1 = ( $b == 0 ) ? 1 : "e^{-$bsqp1}";
$c2 = ( $a == 0 ) ? 1 : "e^{-$asqp1}";

$masq = ( $a == 0 ) ? '' : "- $asq";
$mbsq = ( $b == 0 ) ? '' : "- $bsq";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The level surfaces are the graphs of
\(g(x,y,z) = $e = c\)
for constant values of \(c\) such that
\(0 < c < \infty\).
$PAR
Then \( $ex = -\ln c\), which is
the graph of a sphere since \(-\ln c \ge 0\) (or a point if \(c=1\)).

$PAR
${BBOLD}(b)$EBOLD
In the \(xz\)-plane, \(y=0\), so the intersections of the level
surfaces with the plane are circles,
\( $xzLHS = -\ln(c) $mbsq \).  Thus we must have \(c < e^{-$bsq}\),
for example, \(c = $c1\), for which the intersection of the level
curve with the \(xz\)-plane is
\[ $xzLHS = 1. \]

$PAR
${BBOLD}(c)$EBOLD
Similarly, in the \(yz\)-plane, \(x=0\), so the intersections of the level
surfaces with the plane are circles,
\( $yzLHS = -\ln(c) $masq \).  Thus we must have \(c < $c2\),
for example, \(c = e^{-$asqp1}\), for which the intersection of the level
curve with the \(xz\)-plane is
\[ $yzLHS = 1. \]

$PAR
${BBOLD}(d)$EBOLD
Finally, in the \(xz\)-plane, \(z=0\), so the intersections of the level
surfaces with the plane are circles,
\( $xyLHS = -\ln(c) \).  Thus we must have \(c < 1\),
for example, \(c = e^{-1}\), for which the intersection of the level
curve with the \(xz\)-plane is
\[ $xyLHS = 1. \]


END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
