##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Trigonometry')
## DBchapter('Polar coordinates & vectors')
## DBsection('Polar and rectangular coordinates')
## Date('8/6/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "answerHints.pl",
   "PGgraphmacros.pl",
   "AppletObjects.pl",
   "niceTables.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Complex");
Context()->flags->set(
    reduceConstants=>0, 
    reduceConstantFunctions=>0,
    tolType=>'absolute',
    tolerance=>0.0005);

# Ensure that at least one of our complex numbers is not in Q1
do {
    $alpha1 = random(-1,1,2)*random(2,5,1);
    $beta1 = random(-1,1,2)*random(2,4,1);
} until ($alpha1<0 || $beta1<0);

$complex1 = Complex("$alpha1+$beta1 i");
$radicand1 = $alpha1**2 + $beta1**2;
$modulus1 = Formula("sqrt($radicand1)");

## for static image / PDF
$x1 = $alpha1/2;
$y1 = $beta1/2;
$lr1 = ($alpha1<0)?"right":"left";
$tb1 = ($beta1<0)?"bottom":"top";
$lrVec1 = ($alpha1*$beta1<0)?"left":"right";

# Repeat for second complex number/vector, ensure different complex number
do {
    $alpha2 = random(-1,1,2)*random(2,5,1);
    $beta2 = random(-1,1,2)*random(2,4,1);
} until ($alpha1!=$alpha2 || $beta1!=$beta2);

$complex2 = Complex("$alpha2+$beta2 i");
$radicand2 = $alpha2**2 + $beta2**2;
$modulus2 = Formula("sqrt($radicand2)");

## for static image / PDF
$x2 = $alpha2/2;
$y2 = $beta2/2;
$lr2 = ($alpha2<0)?"right":"left";
$tb2 = ($beta2<0)?"bottom":"top";
$lrVec2 = ($alpha=2*$beta2<0)?"left":"right";

# for ggbChecker
$complexList = List($complex1,$complex2);

# for display purposes
$complex1tex = $complex1->TeX;
$complex2tex = $complex2->TeX;

# determine product 
$productAlpha = $alpha1*$alpha2 - $beta1*$beta2;
$productBeta = $alpha1*$beta2 + $alpha2*$beta1;
$product = Complex("$productAlpha + $productBeta i");
$productRadicand = $radicand1*$radicand2;
$productModulus = Formula("sqrt($productRadicand)");

# for display 
$producttex = $product->TeX;

# for the solution


$quad1 = ($alpha1>0)? (($beta1>0)? "\text{I}" : "\text{IV}" ): (($beta1>0)? "\text{II}" : "\text{III}" );
$adjust1 = ($alpha1>0)? (($beta1>0)? "\arctan($beta1/$alpha1)" : "\arctan(($beta1)/$alpha1)" ): (($beta1>0)? "\arctan($beta1/($alpha1))+180^{\circ}" : "\arctan(($beta1)/($alpha1))+180^{\circ}" );
$quad2 = ($alpha2>0)? (($beta2>0)? "\text{I}" : "\text{IV}" ): (($beta2>0)? "\text{II}" : "\text{III}" );
$adjust2 = ($alpha2>0)? (($beta2>0)? "\arctan($beta2/$alpha2)" : "\arctan(($beta2)/$alpha2)" ): (($beta2>0)? "\arctan($beta2/($alpha2))+180^{\circ}" : "\arctan(($beta2)/($alpha2))+180^{\circ}" );
$quadproduct = ($productAlpha>0)? (($productBeta>0)? "\text{I}" : "\text{IV}" ): (($productBeta>0)? "\text{II}" : "\text{III}" );
$adjustproduct = ($productAlpha>0)? (($productBeta>0)? "\arctan($productBeta/$productAlpha)" : "\arctan(($productBeta)/$productAlpha)" ): (($productBeta>0)? "\arctan($productBeta/($productAlpha))+180^{\circ}" : "\arctan(($productBeta)/($productAlpha))+180^{\circ}" );

##############################################
#  Begin trig functions in degrees

Context()->functions->remove("cos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cos {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub cos {Parser::Function->call('cos',@_)} 
#  Add the new functions to the Context
Context()->functions->add( cos => {class => 'NewFunc', TeX => '\cos'}, );


Context()->functions->remove("sin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sin {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub sin {Parser::Function->call('sin',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sin => {class => 'NewFunc', TeX => '\sin'}, );


Context()->functions->remove("tan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub tan {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180)/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub tan {Parser::Function->call('tan',@_)}
#  Add the new functions to the Context
Context()->functions->add( tan => {class => 'NewFunc', TeX => '\tan'}, );


Context()->functions->remove("sec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sec {
  shift; my $x = shift;
  return 1.0/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub sec {Parser::Function->call('sec',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sec => {class => 'NewFunc', TeX => '\sec'}, );


Context()->functions->remove("csc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub csc {
  shift; my $x = shift;
  return 1.0/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub csc {Parser::Function->call('csc',@_)} 
#  Add the new functions to the Context
Context()->functions->add( csc => {class => 'NewFunc', TeX => '\csc'}, );


Context()->functions->remove("cot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cot {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180)/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub cot {Parser::Function->call('cot',@_)}
#  Add the new functions to the Context
Context()->functions->add( cot => {class => 'NewFunc', TeX => '\cot'}, );



#sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))}
#sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])}
#sub atan {CORE::atan2($_[1],1)}
#sub acot {CORE::atan2(1,$_[1])}
#sub asec {acos($_[0],1.0/$_[1])}
#sub acsc {asin($_[0],1.0/$_[1])}


Context()->functions->remove("acos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acos {Parser::Function->call('acos',@_)}
#  Add the new functions to the Context
Context()->functions->add( acos => {class => 'NewFunc', TeX => '\arccos'}, );


Context()->functions->remove("asin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub asin {Parser::Function->call('asin',@_)}
#  Add the new functions to the Context
Context()->functions->add( asin => {class => 'NewFunc', TeX => '\arcsin'}, );


Context()->functions->remove("atan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub atan {CORE::atan2($_[1],1)*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub atan {Parser::Function->call('atan',@_)}
#  Add the new functions to the Context
Context()->functions->add( atan => {class => 'NewFunc', TeX => '\arctan'}, );


Context()->functions->remove("asec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asec {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
#sub asec {acos($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub asec {Parser::Function->call('asec',@_)}
#  Add the new functions to the Context
Context()->functions->add( asec => {class => 'NewFunc', TeX => '\arcsec'}, );


Context()->functions->remove("acsc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acsc {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
#sub acsc {asin($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub acsc {Parser::Function->call('acsc',@_)}
#  Add the new functions to the Context
Context()->functions->add( acsc => {class => 'NewFunc', TeX => '\arccsc'}, );


Context()->functions->remove("acot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acot {CORE::atan2(1,$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acot {Parser::Function->call('acot',@_)}
#  Add the new functions to the Context
Context()->functions->add( acot => {class => 'NewFunc', TeX => '\arccot'}, );


#  End trig functions in degrees
###################################################

$ggbCheck = sub {
    my ($correct,$student,$ansHash,$value) = @_;
    my $n = scalar(@$student);  # number of student answers
    my $score = 0;              # number of correct student answers
    my @errors = ();            # stores error messages
    my ($correct1,$correct2) = @{$correct};
    my $purple = Complex($student->[0]);
    my $orange = Complex($student->[1]);
    $ansHash->{preview_latex_string} = "{\zeta_1 : $purple \\ \zeta_2 : $orange}";
    $ansHash->{preview_text_string} = "{purple: $purple; orange: $orange";
    return 0 if $ansHash->{isPreview};
    push(@errors,"It looks like you graphed the given complex number twice.") if ($purple == $orange);
    if ($orange!=$correct1 && $orange!=$correct2) {
        push(@errors,"The product of \(\zeta_1\) and \(\zeta_2\) is incorrect.<BR>Check the product box and manipulate \(\zeta_2\) until the product matches \($producttex\).");
    } else {
        $score=2;
    }
    return ($score,@errors);
};

# Determine argument in degrees, adjusting for Q2 & Q3

$adjustQuadrant1 = ($alpha1<0)? 1 : 0; # adjust angle for Q2 & Q3? 
$argument1 = atan($beta1/$alpha1) + 180*$adjustQuadrant1;
$argument1 = Real($argument1); #->with(period=>360);

$adjustQuadrant2 = ($alpha2<0)? 1 : 0; # adjust angle for Q2 & Q3?
$argument2 = atan($beta2/$alpha2) + 180*$adjustQuadrant2;
$argument2 = Real($argument2); #->with(period=>360);

$productArgument = Real($argument1+$argument2); #->with(period=>360);
#$productArgument = ($productArgument>360)?Real("$productArgument-360")->reduce:Real($productArgument)->reduce;
#$productArgument = $productArgument->with(period=>360);

######################################
# Create  link to applet:
###################################
# You can name your applet (anything reasonable :-) )
# Adjust the height and width as desired
# Paste the geogebra parameters in at the bottom of the page just above the
# command end command
# so that they don't get in the way
###################################


$appletName = "ggbfoo";

$applet =  GeogebraWebApplet(
   code       => "geogebra.GeoGebraApplet",
   archive    => "",
   codebase                 => "",
   appletName              => $appletName,
   appletId                    => $appletName,
   submitActionAlias     => 'getAppletValues',  # default actionAlias
   initializeActionAlias   => 'configure',  # default actionAlias
   setStateAlias             => 'setXML',
   getStateAlias             => 'getXML',
   setConfigAlias           => '',
   getConfigAlias           => '',
   returnFieldName       => '',
   width                 => 538,              # may want to modify width
   height                => 469,              # may want to modify height
  # mayscript             => "true",
   debugMode             => 0,   # set debugMode to 0 for no debug
                                              # to 1 to make xml representation visible
                                              # to 2 to add alerts detailing progression
                                              # through the applet
   onInit                => 'ggbOnInit',
   type                  => 'geogebraweb',
   submitActionScript    => qq{ getQE('answerBox').value = getAppletValues() },
   #submitActionScript => '',
   selfLoading           => 1,
   params     => GEOGEBRA_PARAMS(),# paste parameters in

);


###############################################
#
#  PNG picture of GeoGebra applet for hardcopy

$im = init_graph(-6,-6,6,6,axes=>[0,0], grid=>[12,12], size=>[400,400]);
$im->moveTo(0,0);
$im->arrowTo($alpha1,$beta1,'blue',3);
$im->lineTo($alpha1,0,'orange',2,'dashed');
$im->lineTo(0,0,'orange',2,'dashed');
$im->lb( new Label($alpha1,$y1," $beta1 ",'orange',$lr1,'middle'));
$im->lb( new Label($x1,0,"$alpha1",'orange','center',$tb1));
$im->lb( new Label($x1,$y1," $complex1 ",'blue',$lrVec1,'bottom'));

$im->moveTo(0,0);
$im->arrowTo($alpha2,$beta2,'blue',3);
$im->lineTo($alpha2,0,'orange',2,'dashed');
$im->lineTo(0,0,'orange',2,'dashed');
$im->lb( new Label($alpha2,$y2," $beta2 ",'orange',$lr2,'middle'));
$im->lb( new Label($x2,0,"$alpha2",'orange','center',$tb2));
$im->lb( new Label($x2,$y2," $complex2 ",'blue',$lrVec2,'bottom'));


##################################
# Setup GeogebraWebApplet --
###################################

###############################################
#
#  For persistence of GeoGebra applet state 
#  when answers are submitted or page is refreshed.
#  Uses JavaScript interface to GeoGebra.

HEADER_TEXT(
qq! 
<script language="javascript">
   function ggbOnInit(param) {
       if (param == "$appletName") {
           applet_loaded(param,1);  // report that applet is ready. 
           ww_applet_list[param].safe_applet_initialize(2);
       }
   } 
</script>
<noscript>For this question to work properly, 
you must enable JavaScript, or use another browser 
(or computer) on which JavaScript works.</noscript>
!);


#######################################################
#
#  For setting initial values in GeoGebra from WeBWorK
#  Uses JavaScript interface to GeoGebra.

TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );
<script language="javascript">
function setAppletCoefficients() {
   var applet=getApplet("$appletName");   // alert("Updating coefficients");
   if (typeof(applet)!="undefined") {
     if (typeof(applet.setValue)!="undefined") {
     
     ////////////////////////////////////////////////////////////////
     //  List the values of the parameters to be set in the applet here
     
          // hide product at first
          applet.setValue('d',0);
          // re-label vectors with complex variable names
          applet.setCaption('v','\$\$ \\zeta_1 \$\$');
          applet.setCaption('w','\$\$ \\zeta_2 \$\$');
          applet.setLabelStyle('v',3);
          applet.setLabelStyle('w',3);
          // lock to grid
          applet.setPointCapture(1,2);
          // set and lock one complex vector
          applet.setCoords('v',$alpha1,$beta1);
          applet.setFixed('v',true,false);

     } else {
          setTimeout("setAppletCoefficients()", 1000);
     }
   } else {
      setTimeout("setAppletCoefficients()",1000);
   }
}
ww_applet_list["$appletName"].setConfig = function() {setAppletCoefficients()};
</script>
END_SCRIPT

HEADER_TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );

<script language="javascript">
function getAppletValues() {
    console.log("getAppletValues() entered");
	// get a handle for the applet object
	var applet = $appletName;
	var vX = applet.getXcoord("v");
	var vY = applet.getYcoord("v");
	var wX = applet.getXcoord("w");
	var wY = applet.getYcoord("w");
    // return string = "purple complex, orange complex"
    return( vX + '+' + vY + '*i,' + wX + '+' + wY + '*i' );	
}
 </script>
END_SCRIPT

 TEXT($PAR,  MODES(TeX=>$im,
                      HTML=>$applet->insertAll(
                                debug =>0,
                                reinitialize_button => 0,
                                includeAnswerBox=>1,
 )));

NAMED_ANS('answerBox'=>$complexList->cmp(list_checker=>$ggbCheck));


$multiAns = MultiAnswer($productModulus,$productArgument,$modulus1,$argument1,$complex2,$modulus2,$argument2)->with(
    singleResult=>1,
    allowBlankAnswers=>1,
    checkTypes=>0,
    tex_format=>'{|\zeta_1\cdot\zeta_2|=%1$s \\ \arg(\zeta_1\cdot\zeta_2)=%2$s \\ |\zeta_1|=%3$s \\ \arg(\zeta_1)=%4$s \\ \zeta_2 = %5$s \\ |\zeta_2|=%6$s \\ \arg(\zeta_2)=%7$s}',
    checker => sub {
        my ( $correct, $student, $self, $ansHash ) = @_;
        my $score = 0;
        my $n = scalar(@$student);
        my $i;
        my @ansObj = ('\(|\zeta_1\cdot\zeta_2|\)','\(\arg(\zeta_1\cdot\zeta_2)\)','\(|\zeta_1|\)','\(\arg(\zeta_1)\)','\(\zeta_2\)','\(|\zeta_2|\)','\(\arg(\zeta_2)\)');

        for ($i = 0; $i < $n; $i++) {
            my $cor = $correct->[$i];
            my $stu = $student->[$i];
            next if $stu eq ""; # don't choke if students leave a blank answer
            # periodic answers are failing _this_ logic 
            if ( $stu != $cor ) {
                # but satisfying _this_ logic...
                if ( abs($cor-$stu) == 360 && ($i-1)*($i-3)*($i-6) == 0) {
                    $score++;
                    next;
                }
                if ( abs($cor-$stu) < 0.05 ) {
                    $self->setMessage($i+1,"You're close. Make sure your answer is accurate to three decimal places.");
                } elsif ( abs(180*$stu/pi - $cor) < 0.05 && ($i == 1 || $i == 3 || $i == 6)) {
                    $self->setMessage($i+1,"It looks like you might have answered in radians instead of degrees.");
                } else {
                    $self->setMessage($i+1,"$ansObj[$i] is incorrect.");
                }
                next;
            } 
            $score++;
        }
        return ($score/$n);
    });

##############################################################
#
#  Text
#
#

BEGIN_PGML

a. Given the vector [`\zeta_1 = [$complex1]`] and the product [` \zeta_1 \cdot \zeta_2 = [$product]`], use the image above to find the vector [`\zeta_2`].

b. Use your results to fill in the following table:

* Make sure you're using *degrees* instead of radians.
* If you use any decimal approximations, they must be accurate to _at least 3 decimal places_.

[@
DataTable(
  [
    [['', headerrow => 1,midrule=>1],'`a+bi` form','modulus','argument'],
    ['`\zeta_1 \cdot \zeta_2`',"`$producttex`",$multiAns->ans_rule(15),$multiAns->ans_rule(15)],
    [['`\zeta_1`',midrule=>1],"`$complex1tex`",$multiAns->ans_rule(15),$multiAns->ans_rule(15)],
    ['`\zeta_2`',$multiAns->ans_rule(15),$multiAns->ans_rule(15),$multiAns->ans_rule(15)],
  ], 
  align => 'r | c c c'
)
@]***




END_PGML

##############################################################
#
#  Answers
#
#

ANS($multiAns->cmp);

BEGIN_PGML_HINT

* You do not have to find the modulus and the argument of  [`\zeta_2`] using its complex number representation.
Instead, use the modulus and the argument of [`\zeta_1`] and [`\zeta_1\cdot\zeta_2`].

* Use that [`|\zeta_1\cdot\zeta_2| = |\zeta_1|\cdot|\zeta_2|`] and isolate [`|\zeta_2|`].

* Use that [`\arg(\zeta_1\cdot\zeta_2) = \arg(\zeta_1)+\arg(\zeta_2)`] and isolate [`\arg(\zeta_2)`].

END_PGML_HINT

##############################################################
#
#  Solutions
#
#

$a1a2 = $alpha1 * $alpha2;
$a1b2 = $alpha1 * $beta2;
$b1a2 = $beta1 * $alpha2;
$b1b2 = $beta1 * $beta2;

BEGIN_PGML_SOLUTION

a.

    >> [@ image( insertGraph($im), width=>400, height=>400, tex_size=>800 ) @]* <<

The complex number [`\color{blue}{\zeta_2}`] is [`\color{blue}{[$alpha2]+[$beta2]i}`]. Let's check that the product [`\color{red}{\zeta_1}\cdot\color{blue}{\zeta_2}`] is indeed [`\color{green}{[$product]}`] when [`\color{red}{\zeta_1}=\color{red}{[$complex1]}`].

>>[` \begin{aligned}
\color{red}{\zeta_1}\cdot \color{blue}{\zeta_2}&\quad= (\color{red}{[$complex1]})\cdot (\color{blue}{[$complex2]}) \\
&\quad= [([$alpha1])\cdot ([$alpha2])] + [([$alpha1])\cdot ([$beta2] i)] +  [([$beta1] i)\cdot ([$alpha2])]+
[([$beta1] i)\cdot ([$beta2] i)]\\
&\quad = [$a1a2] + ([$a1b2])i + ([$b1a2])i + ([$b1b2])i^2 \\ 
&\quad = [$a1a2] + ([$a1b2])i + ([$b1a2])i - ([$b1b2])\\ 
&\quad = ([$a1a2]- ([$b1b2])) + [([$a1b2]) + ([$b1a2])] i \\
&\quad = \color{green}{[$product]}
\end{aligned} `]<<

b.


* The modulus of [`\color{green}{\zeta_1\cdot\zeta_2= [$product]}`] is 
>>[`\color{green}{|\zeta_1\cdot\zeta_2|}= \sqrt{([$productAlpha])^2+([$productBeta])^2} = \color{green}{[$productModulus]}`]. <<
Its vector representation, [`\langle [$productAlpha],[$productBeta] \rangle`], is in quadrant [`[$quadproduct]`]. The argument of [`\color{green}{\zeta_1\cdot\zeta_2= [$product]}`] is 
>>[` \color{green}{\arg(\zeta_1\cdot\zeta_2)} = [$adjustproduct] =  \color{green}{[$productArgument]^{\circ}}`].<<

* The modulus of [`\color{red}{\zeta_1= [$complex1tex]}`] is 
>>[`\color{red}{|\zeta_1|}= \sqrt{([$alpha1])^2+([$beta1])^2} = \color{red}{[$modulus1]}`]. <<
Its vector representation, [`\langle [$alpha1],[$beta1] \rangle`], is in quadrant [`[$quad1]`]. The argument of [`\color{red}{\zeta_1= [$complex1tex]}`] is 
>>[` \color{red}{\arg(\zeta_1)} = [$adjust1] =  \color{red}{[$argument1]^{\circ}}`].<<

* The modulus of [`\color{blue}{\zeta_2= [$complex2tex]}`] can be calculated in the following way:
>>[`\begin{aligned}
\color{green}{|\zeta_1\cdot\zeta_2|}  = \color{red}{|\zeta_1|}\cdot\color{blue}{|\zeta_2|} & \quad\Longrightarrow\quad \color{blue}{|\zeta_2|} = \dfrac{\color{green}{|\zeta_1\cdot\zeta_2|}}  {\color{red}{|\zeta_1|}} \\
& \quad\Longrightarrow\quad \color{blue}{|\zeta_2|} = \dfrac{\color{green}{[$productModulus]}}{\color{red}{[$modulus1]}} \\
& \quad\Longrightarrow\quad \color{blue}{|\zeta_2|}  = \color{blue}{[$modulus2]}
\end{aligned}
`].
<<

We can calculate the argument of [`\color{blue}{\zeta_2= [$complex2tex]}`] as: 
>>[` 
\begin{aligned}
\color{green}{\arg(\zeta_1\cdot\zeta_2)}=\color{red}{\arg(\zeta_1)} +\color{blue}{\arg(\zeta_2)} &\quad\Longrightarrow\quad \color{blue}{\arg(\zeta_2)} 
= \color{green}{\arg(\zeta_1\cdot\zeta_2)}-\color{red}{\arg(\zeta_1)} \\
&\quad\Longrightarrow\quad \color{blue}{\arg(\zeta_2)}  = \color{green}{[$productArgument]^{\circ}}-  \color{red}{[$argument1]^{\circ}} \\
&\quad\Longrightarrow\quad \color{blue}{\arg(\zeta_2)} = \color{blue}{[$argument2]^{\circ}}
\end{aligned}
`].
<<



[@
DataTable(
  [
    [['', headerrow => 1,midrule=>1],'`a+bi` form','modulus','argument'],
       ['`\zeta_1 \cdot \zeta_2`',"`$product`","`$productModulus`","`$productArgument`"],
    [['`\zeta_1`',midrule=>1],"`$complex1tex`","`$modulus1`", "`$argument1`"],
    ['`\zeta_2`',"`$complex2tex`","`$modulus2`","`$argument2`"]
  ], 
  align => 'r | c c c'
)
@]***

   
   


END_PGML_SOLUTION



##############################################################
#


sub GEOGEBRA_PARAMS {
    my $result = {
showResetIcon=>"true" ,
enableRightClick=>"false" ,
enableLabelDrags=>"false",
showMenuBar=>"false" ,
showToolBar=>"false",
showAlgebraInput=>"false",
enableShiftDragZoom=>"true",
useBrowserForJS=>"true",
ggbbase64=>"UEsDBBQACAgIAEGACU0AAAAAAAAAAAAAAAAXAAAAZ2VvZ2VicmFfZGVmYXVsdHMyZC54bWztmlFz2jgQgJ+vv0Kjp7uHgGUw0EycTtqZm8tMkmYumc69ClsYXYTks+Rg+PWVJbBNgRQMFJJLHiKvWMnyt6vVSvbFp2zEwDNJJBXch6jhQEB4IELKIx+manDWg58uP1xERESkn2AwEMkIKx96uWbRTkuNds/N63Ac+zBgWEoaQBAzrPImPhxDADJJz7m4wyMiYxyQh2BIRvhGBFiZXoZKxefN5ng8bszv1xBJ1Iwi1chkCIEeK5c+nF2c6+4WGo1bRt11HNT85/bGdn9GuVSYBwQC/RwhGeCUKakvCSMjwhVQk5j4MBaUKwgY7hPmw/tcAr8PEkL+gGDWSONx4OWH3y7kUIyB6P9LAl2nkpQU7YzQzHX0z18EEwlIfNjtQhDZou9D1/M0JRYPsQ8dq8zwhCTgGbOiBqdKBKa9qR1gJslcV9/pVoTE/tKe6XM6MhSBVEQbwGkgCGRMSKhHDWfPiIw9Jsa0lR4DIZJQgsyHd/gOgsmsnNrSqBg6D3Q6u6lXrVUTRipjv2jOwG6GOCQx4aFWWuCManHu9AznvOjb4tCYDwm5fWjInXfI6yCj7Sl/5VW2bi22yPUMXFP+qlDxKgLFNf+bRHrMVcatd8Z7Zbzowe1adB3D1nmlZI2KZSjz/zqREaOYkWyP4BnlJcQbIxTQ3Xr5RRW6cxTkTm3kOQ4LTw1p8MSJlDnZst/84i8a6tXL3E/oxJEq3RPq9mwP5D++YDKqLUa1zstmGKQ8UCagzNB+SZPnqi1abecY1ij73LcxdiW9nqUkUS4VXB7mcunY9RK6/7dji1Sx/M7XXOl9FjHuKpce7YmQ+FF39ZU/JpjLfLO16Enr7ZbgyUs2895tdno2m8etu284KSyR6uR+oMceVs1XLz1au4A3XO/YNtwimq8ksntKc1IOvaW37setOvWiguu0V2NsdE/YrZ7144mSx7eZWOYI7/na9iFyRZKNE0UkxfxnWxY2iSpz/H4uF/boWnvsPsatN5Vey9jUQ0v+jRz7h9ofHYQ6yD22mV8GvLA9uS8qSsToSIhPdNKspxkInp+Az7cXVio4tt9Y6NjDPo5GhNuIKwHIHKM2cUzjqTN7N5EhI0+Q+XWKbLVprwee0Axc2RZXVvHKtUXLFm1beAWgeptHY9pYR61K+vzD0tCut+N5TYHkTRr9F6TwPB2RpBIa7uZy4TyeDQ66v5QsmHaDULDOT9Z7hWQ01C40otpIZ9p6I5wZK+K+FCxV5CFICOHlazrrxmMaqmGe2Ol7D2iWu4v9YSgSOhVcFTBAPgmumHmft3DAscp73Jfy1wVf3S06Yx6xcjJeWak0gD2/N0o/Hu2tsksVoTMj2Gm4vRbqeS2ni7ofvV5nQ6KoV5voQrSxEDZZTpCzsRftHm22ihnuqpiBk6A8hG05e3aMpa3kn0VFuQk6xZNB4zJLqgc79GMiSGV5kG2lglDvjaU3OM0ooziZLN/pYIQVycoE49EIlS8RThLwQEcpIPVqNijij9RT9TZfsdv2Wm8h4Hzh/cmhhDZWVD7QtZUqHwlYBAOq2XM80g3sICj/jIOnKBEpD5cXu70AQ8f2yPXQ+kIwgsvw9XkuV15OL6UX6wBtvoYczKWCIQme+iJbWBJfjkxUlvPmxgiVl8Yr5s0uK+XZ0V2hzkHgpu8yz5ZmaLPyBVVz/pXW5XdQSwcI/6AQMOwEAABHJgAAUEsDBBQACAgIAEGACU0AAAAAAAAAAAAAAAAXAAAAZ2VvZ2VicmFfZGVmYXVsdHMzZC54bWztl0tu2zAQhtfNKQjuI4my5MSBlcBIFy2QFC2y6ZamxjZbmVRI+qFcrXfomcqH7MhJHDRG2qJovdDwMTMkv18i6eHFel6hJSjNpSgwiRKMQDBZcjEt8MJMjk/xxfnRcApyCmNF0USqOTUFzp3nNs7Wouw0dW20rgvMKqo1ZxjVFTUupMArjNBa8zMhP9A56JoyuGEzmNMryajxWWbG1GdxvFqtos14kVTTeDo10VqXGNm5Cl3gtnBm0+0ErXrePU0SEn++vgrpj7nQhgoGGNl1lDChi8poW4QK5iAMMk0NdsZScNazY1R0DFWB3wtjFwfMzQyxhVra+Da4wD2SJ/j86M1Qz+QKyfEX61dgoxawjfeV2PnY7ktZSYVUgdM8x8iCJekJRuMCO15VPaO2FJEk/Eg2SAjpkzSEV7QBhZa0ct6+hS6MZD6jb53QSsPG1459LUsIPVnrL/jcE0bagBWHYKRrgNKXwuoTr1TjRe/m4wJuTFMBMjPOvgrQln7eCXKFd7wswb07IQb4FMTSApFKW8UTP0qTePe7pH3F1sTXG+J770ho9vF2qoqv0ShEjILjKA2mF0wWTL5FArcizFO7Z4FrquxLZhMx1z+MW60fqU7XXHdEH7nq2x2hk95BQide5uShyH+ppPvporYMds3fvz0P239FjCoDmlPRwX7pOh5y7//r3PeDtPkFdPh99PUdfnYPPIjfYOABpmTgEXq73aHy18LIpFSlRuuwCYStwT9X25QT6o6gdpSIPAMxORCirJoZlEqKe46dpnuUvRblIV/OS/GTvOf55+TR8ZC1x0M+6CdZP3s1LQ59pfeTFWC2y/zgyl2W+X+WL2F5u6Cl32fbpX7a1LtMyWF3kSR7erOMTl4NyK+4QDx5fXCN4Y7QBHOXbhO+9EaBRv1gToI5DWaw97bB53XFGTfPS6kXamIvvU8dgG3XrqrZn1L1PvFvOQTJz+7fcefyHm/+IJz/AFBLBwgpvB2D4QIAAMIMAABQSwMEFAAICAgAQYAJTQAAAAAAAAAAAAAAABYAAABnZW9nZWJyYV9qYXZhc2NyaXB0LmpzSyvNSy7JzM9TSE9P8s/zzMss0dBUqK4FAFBLBwjWN725GQAAABcAAABQSwMEFAAICAgAQYAJTQAAAAAAAAAAAAAAAAwAAABnZW9nZWJyYS54bWztW9ly28gVfZ75ii6kZkqKRQqNnTblKY0fElfZI5Wl2M4keWgCLRIRCGCA5uaaD8if5CVfkU/Jl+TebgAExH2RRknZNthAo7d7z7lLA3D3h+kwImOe5WESX2i0rWuEx34ShHH/QhuJu5an/fD6226fJ33eyxi5S7IhExeajS2rfnDVtjwD61iaXmh+xPI89DWSRkxglwttopEwuNDuDNPSHZ+3Op7jtCzfYi3m+nrL4r5pepZx51meRsg0D1/GyU9syPOU+fzGH/Ahe5f4TMj5BkKkL8/PJ5NJu1xZO8n65/1+rz3NA42AVHF+oRUnL2G4RqeJKZsbuk7PP79/p4ZvhXEuWOxzjaDEo/D1t990J2EcJBMyCQMxADlNRyMDHvYHoALL6WjkHBuloIeU+yIc8xy61i6lzGKYarIZi/H+N+qMRJU4GgnCcRjwDBDQSJKFPBbFLVpMcV527o5DPlGj4JmcwKaGRkSSRD0GQ+jkV0KJrcNBaIecEceFGoNQm1hQ40GNS0yss6lFTIJNqEksC0oLq6mD92zob+uEUqgmhk4MgxiUGCZc2jaxHWK72NGAtk5HDqbDga1hOXCYWGeacMg604LDwDMYyFbDwCJAo/LMxtYwvm3g8mWl6RGrAxNhhe1SYsIa4NrVCYxo4vBUCmHpBP9RYuHwhksMj8B4IDeOrINqxmEe9iIO7GNRDvCG8V0G1KquczGLuNRdUTGHhp7BX2gRfoHmtg4EV4yAO7p+hocDh4U3zpugWE1IAAEdZDvDgqoCl+s46pau6nRTFYYqLFXYqo2luluqqZJWt1QbyzxUzFJIcxchvZqQFIUAUHD1sjAJrpvK9WNhFZeOupRU06le1Hr408EL0InjyZMDZTL3konWZhXZaLdJqymlq9h2ysM4Ws1pWLuIeaB2y0mpXZvU1s/kP3ksTGnuJKfS/WEzOg0rfOrZqeEdR8eWt/Wclt5xjyMnLr6c07PUjC7Oai/O6upLvZ0qaVE+OfqHOsVK/Rum7J6XQblbKIHkA2xbuA/BhzmqxewQ1ySOUQVoB0NoEaVdg7g2cZ1arD7DaO3Y84CN4dprBGzbq0VtCNkOVroyBYD5MOaqCG5YZRA/K8L4rwthHKKuNQ+8sEAcihICWYJ0yGUEhlUYVQw2bAzDBrhsCP8GcTAErAjHkA4meVjpdsCjtNK6VGMYpyPRUJ0/DMpTkUBrFslkr2gfJP79j5Wyi5E4y0V9WEid5vmYSqUa6do33Yj1eAQJ7w1SgZAxi9BXyBnukliQkniOJoeTmWGXj/woDEIWfwToy7zsp9GwxzMiTxMUUg6C3cnyFJJaqomfJFlwM8uBKWT6M8+gs+FYbcswOzZkxBaQHqx6VtwxjXbH7ni041i66VoejJj7DElu2G3T1j1KbdoxHPiDvZbfsx01NR/fcCFA/JywKc9LdfezMKifv81/TKKgUmyahLF4w1IxyuRWASwhQ5ku437EpSIlxpBY+/e9ZHqj4pOjxrqdpbxSca//JomSjGS4QhsaFGVPlbINLqxqpcs2umxRjIGDVvdBOtlClj1VylaAsVpaISgtpdTLWcKcqOsGCSU/LrQPnEUaGcWheKcqgJOhfz8XFvsoBpRabA5Llw77dsj6Ycyy2T5j4/ph85CLz7IVnv25OrsdcMFAmjak20AX17Xh1+h4nuLxAwZ3C9sq+ewnUcTSnAc1g+ieNxp173kW80i1joFDo2SUq+aKJnKNo5xfMzG4jIMPvA/Gf83QAQuQRDWdSxNwPxxCR1VfoMaQUX8CzajagPczXmpWLUZhWqyS5GnGWZAPOBcVssq86s2kMOXyu7mfhSnaAOmB+7/nc56DHhmEj6AmEcqaw6J99GQAk0CIrrMkgBqS3JE3yTCN+JQUcJEWkUZBbiec3cMGeSQGSSY3fUxAT7k1nCTZvVzyLZ8KwnrJGO6UA33/yygRr8bgsZIsVxcvyZiwOCATcjIJxYCkxfTp6fe/o/or+fMH5AZQ7y6cAohMQMxBgsl7PyfJEMVlUZRMeNAGQ+ERUlMKyCM+hC0oEdJO1cQVX8dywchHkvT+Dvce8Fld8DHugArlx4H0+pLEV0Ufpm7CEHPDNdyaaUsnwKJ0wCoYIzZD74oI0rIG5ry6u8u5IFNJ/NmF1urUbr5PAt5w6SwGjqnlCJ6ieaB5pZwrCxWFb8JHFzPp2mpUljzKcaaWIaeC3y/z5VX+RQzAeGG/Drbq1IbEkz+GQcBVAExS5ocCRqGut8DTXJm+zzLBczDSIpQIuL5G30v4FBZ/VfQDR4wSSTm/m774bhYWNq6Q3ITp5DiY9hYxLV26xFY67NWYGksxbdkKVPr4oD4jTOvQ+clwiMYey7z1o4Juni8x/UI7OWG9/GR8Sn5P8GRyevqKsKyPNS/kCdQU4IxE2SlVExXDb+JIuoEjZfq1iSTBIkmaEX01QczlBDFMidtyqzcXLNdr640/CvVO2+241LXMjgnQeYbHW7hD3sAE+8mtu/D0C+Y9hbCX49PYEo+PmhrkZAosOCMz+D0tVyuzN5UdNzEvblRDLEBebpJqfmEBTseTeGLRU8VKRPV1Rr3C2xaOQC4V80rVx6rXLkT6dYr6NFfURCpqsrOiPj0XRR1PT02n84HNHnoc/YzooCdGkScPXcsYO6z1Lhk0qFKKcvjDvcujpBUF/VreWhRyEFDmhJcqqVpIhBf43NocaoySicYaD+N1KgAiDG9vY8yu1WIW8/F7zlPcgV3FtxmLc3w900zE96XBpwUaTHaiweR4NDhiJnIEnJ8/zIte8XruFVPpFdOdveL1c/GKO4f8p3Od14tZ2U42kx7PZo6UmB1kLysgmS3H8JkbVZMHb8LMj/gaKtAFKvjreQAw4jv2AmZ/P2vbFuK1pnZ4bszDPo+LRxyETHW1l9dl5y968WXClMrrmVQW+UJVtewPu78snJJL1fFStbg05L6RXJqy36WlCrtaN/8lbuTj+Fg4vAv9ze4Rn8SWHpIqDOsucJNrlN33wstVyYzrbpHNrMXs4ZZ5yYaZbtowq+ct+hqfae/lMzfait7uLFrLbuby1V6W24ve9o5hMjsD6i0BdEdEv0K6ClLHasTO4wN8mfkP0AU0JLLKUzJYRxVa/+L7ZwSz/b8t7hjkS4stYE9ZNnenRa8j7R3If/7xT3Ly738Fp4tcOcLzTLd4Sm06B2VNSwlnVa77t2bc3k7kCVKxZWRdy1ag63gpXcd70XV8TLqydXTd99nH/x0dj+QAfxtyKmrWCFl6z4dEPSPpUpqmknDGrjwtuz3a/pLapvKmurEFQQ/bZe4foMGGnuke5YmeBQk+FUax21Hvn09UQV4Q+dQc3zKpmhfzG7PGjfBUlcseG+H4WnOy9ZSrvY/U11HuMR8DqtfqD75g2MhC+UFRzrPwbv4FGmyW3iNNrPJ7M1crwS58MuvlSTQS/MbPOI/L7+DlDkxXr0dNazsYzQaMF6SOY7oKx/qNcAcYzSPB+FhPpp4HhLZ6fenY2yFIVxvieBWA430NkW6NoLcOwCO+lXmmIFK6wgy3DO5L885VmyQVl+l+4XwTpAeE88q3IqUf9ZHx12C+bTBvcKKXJBFncYW12oNA5xFfsKD1Nn+QPdcfpNQ+3rKkCVnmWnMvvyhVn7XVV9P4TKL5Ed5uWukdTyu7ZBvrtaIv14r5ZFoJDtPKus3rkjRtn2dyNaUVSZF7DCpdb/vFDQZMqxGd/yrCIS8+2dS2CLfWDr75fyhjCvN37JZ/ftivHoOVmAeFYFdZCjU2I3VVvrrRd391c7XpIfNiAN33/c3uUMxfzehHezVzXv/eWf4HiOK/3L7+L1BLBwi8z6WtPAsAAD48AABQSwMEFAAICAgAQYAJTQAAAAAAAAAAAAAAABYAAABnZW9nZWJyYV90aHVtYm5haWwucG5nTJxlVBvfF66TECB4KO7Bihcr0KLBobi7tEBxdw9WoGihSItDaXHXYgFKseIUtyLF3f2mv/+9a91Pk6x1MmfPnH3e/byzZidaXVUeD5sSGwAA4CkqyGgCACDUR2AMBAN1qN5kD/j3zV1TXgpQMUy9DQDAAYoykto+wD0TzADtnqnHb2RPpBSkttrYYQwwdgU4WAPe+ZB9s4dXnWbte3GT/Wb629XkN9ab6733obnh8RCFSnq4Agw8vvs2+X642Xtw0TJ576lIv9NFplmN/+HCJ37rB3fvAdTB3DWspqaGc3t72+mHB/crz/oeN5eUHk3TragO6xFaJQODT4zIQT0qIaczPnlrLypBBxd8dZdvwgS/7zg4OLjmG531ELHDfNzcI0kfPuTFSpz+FdUe/zuUoVOjN2x0LvIDXT0rcqAhdP6F0672k8bp3ngardSeg8UWro2NjYGr0gozUmNPWtW70lLO1EUcTtKDq1YZ1bsMJD85F1xSEt6bjIjSl/9bWFhYVFLiBva5LqmqGh2kzpot/jvSXDwsLRDsp6ahqbmKm0l2aIMA2qW39o19IvkXg6aGxge+i+dXXem8Fkc/qZxZB6VR0YSWDkszBzutNZVrKZKQkPB8Uh6oWVl28fSs53tgP27EKhcmYHmQnBzyokNz6d1kdltdBmfnqSjMf/nyeix+PcUq3r75CRgMPGSssp9+FksnQbvw8leWCJkoeYaYf8Ml/5P4UAQsl9H8azzX86SK4jWhYa9ASA7pGIB8cITcqgKt5LfKQ9xsfz0M6gJx2ZSwm2T9lSFEgI6OfrX9LadOEpnh3E7hoq8Qgh7yVhpmj6aep3u6MFkD5jL5Pjmu/VRFdeSJj5rIuKVreAhs1BD6wucI5tR+/Vb0dq+WXNCBXFNTE97hu/737+qwyOGPCDLTqNMRiY56Yr0BesgVFBQ0Wt1sIbHiyAjX83Waq+tPyjFyYBM72LGd6NbU0soY/kKLL5qoqVJyqj6UM8XMtTtVroUxMVYgn0KpX31MupmXpFthvO15UapX2XD6ZWosKHC9j3Awnkak7tZfS87N3qqfo7xOwvmXbPJS3upQKk/S6/LOFzMpuCoD3STdlK/Z59+vrr7BafRkj8ZraB7SlXVgUvTCLZkx0p7mxqngTHjPEFPIqUYTzLGzlNTR4wAwxpv4ZUlFJDOkru2vp8asDIWCKcOQW2MFJk+VVFUxHvmZwbvsTaQmRLj40tNOi9ClfKmm8FcgdOAltEgQLcUuTnSR6EebvrgAIMoBs852LXvliRLqHJP08Pnf0Ce00SslfJZ0Q8xGXwW3EsnVODxJ1/9aFtR1xXpF00CJKA6YZHSs0PkIrHLfx/6Y/8mSa1LWOZdmYVjieqePXwKNvbrAguS5Ys4yMjCF+90XAWeavbm7yjg5FDDCw+lFwX39L7G1dKBPRDxcBhKcla16+ob7C0vN65QUpBwyIP6Z2HPNhMSU8Ue5r6iFA4ijaUS6NG3lfCw/pjZyvZGVfWnQkAzRcKhjZoL4RAE6fjhsjGSRP1okROIaPv0+El+fcIf1cbMRlzWpFAbtg+j3MQgJC1PyH/rSmjDM2dSexCvuu25nca+kukfC8qRZPAo++Tlog3WtSOfzbu4jG12nZpPZgrzPtsZojh3nV2pasiOaP+uRIUpO1prsxnl2fogwfBI+GQ6TsxgPJ4yi9MHDg0QNNNALMoCCj1QVYrCsP+B2y5tEREuzBX17tjVVz5/fJULPNf9hM3SaTlLGg/P75Zbk4+Ko4NMnjcPXuA0qdt2YHj8UGzuixf+MMbnTBdx32xvO0j66iWzaFLLH0oyndKmF97OX7kcH4IDwGNRljZ7x6wqRD5ExXu0BYO0JnrCNmKJo5VA4AHa0Vqr25i+a2uvpKTxy8Sahdp/zwytLf4j49Vv0oPtlEUihwe0PYJqdQnpMgvlAndFY9a2Nw5I3uTz/5Z/hQZaVUz3an7r2kYH9UTIBrQF0tD6raA6LZyLDu8bRzMVUGv1aAV9XoEOk+npKGHlc6I7Lz+l8Td6uXg6J8DzmEMCoX/qcSUh8oA0tYifbXHzAvVFrfXKbsM7AH4hZQSFG7Lzs/lKQ7uk+UENdnTU+0EtkIFLlx6BzR7Sr9YMOzuN2m1698wOrSBVuxiV1r6fzSOEGZrVcAxoj2QE9XElPgXLd1JguZkNLCSDnEDv/dIikLms87ze/XLK9JByqHg5/xArunh+RyHcN3k79RC6xK9t6eyzU5ODAjS2CFhzdSpRjpFm/+fvLacfZ14wPogEM7852iiLEjke5fKweKvs4n7/Wv72/sqPi2Uh9XLmWP5fsjoxJY2xT0+OceNPyYLP8qF1mevbsalPCre/0auROrcrHLFWwm6CYvYercSFU5PEcw6/uT7OyqFNOIJu4fi3Fy+Z1lxKtjz/P9JewioQf+EMF5Dl/+7jZ0YnezL+hrqjK1ZKB/m074SuK7lwkix1fHYyErZCC5sSD/Xqv2t6/9+aIrpQ9McBD4BWxu7V6RdVx/LajqZqtSXZGXW0Vg8tArP2+/g+W7SzbS9X4d88WVmY5RSv4TX3+Su6ZPbx81rLOoEzujfE9KyglNXvzciNSaNya9SB4T0pELfynuf4ApLctphTYNjwmHDxfIGjpjWizrSuD0/8xOe6Srj06phFaGY95u/KL99ufMiraBlnPgISWublrd3WeZy2XGrmWdA+njNxo6cfbXRvIXI5R7pc/P8hVgEPeQV9UdR2KPM2hAyTiJf6V/VDPKYWS+9EAiIeqJXOr/8eMduw2Iv+Vbl+z3BqVJ8MZ13qtj58L+Ct+SSjHhn6g1pLe7vupda7OWny0m58ktpzv2Njpc1OgNva0muWtaRBNtv8or031t1KPC74p60nis8JIHPFAsh7/tNfn8QGi932c4dEXy/ifjJtC4JXEfoX3LRYlt9MbctcZHK+qdF5NgbQTQxIdiByerTPTSSLFwLpBj3TUjCR39mYBobNkKmNPaEzus8VcHTq0tufCufX2ehmT/Q27B6tvvTUPI86qek/XVxczHl7tO4Pj/5LC6OWWhatWs6TxaiDxr7ocvVW7spy5frZkuT0/f+xJjhRIV2WAgFi5AV0KVKRShxk53g/GGeKVQJ52t4C38H5eiyhjHFAoghsGJe2eCxUlJtmKHNUmo/vpm7b4NZr01e6eTYF3lZoxHxDT1jf7193xXuD02qaJQAOWHdvCbVQ+53auwDrmblawxsWzTP8RHL/VlZuMhpiKpuSfjPOuI1SnY4d3kaPckg2PzZF+iN6SXCAioeXy/Mu7nccUn65iilW5UQLx6mB+SyYDCAj4DuHEb+adl5Wy5Se61hxhMdJigBNwdYMDVc5yVnCw/hH7GbeBOn1ShPBlXWTHCwvmXxHbW4REMq5SR0cxOaNTrWuqn05CEUML3z5WOM5fZzKM4fW2+D5QvjeXMuMDrIvOe1FHLRVo1hUlqynLGLbHJZOun41tfOTn7aR0Gynj6CFqJhJ202IBuS6GbP0QprudD/vJDRzlroqVO2l1IJUb9p/Jk3fXPDANZMvlHQS+7Vb85GCBbFAmqatPpyJJ7HlyDTmmKWG4tRQPTQ23fxu12I3EUv4cbfAL/NSThdPhBYVOfwP7pYvDH301ToMYFly+xF6Nbhqeu1WW/tcjyaqvc0hstXy5ylROAG5VWu/SzO7I/HJeYFto/cwLdGB6X8vvcuRDkWQzaC47zsYz5hfgqyhewplnM13BVfFJFEJHVzEV7E8Ypv56RAh33XieTQLFG2N+iaWmMieTzHY8tgD2o3shM0ZWBuDa8GdEw9nK1R+0Pvoxv56I17sP37XfjZG7/ibnvC8jJ0ug0+egw4KC9yDpRe59SphJkTNhhoOyBZiADlMCTv8c81z8U1W70/nmEAUnyqoDC/NdUODHnOfY6S6MoxZVVNH2MtGq+USkcYKvnuv00Rcs6Twy0ykXCz9WPP+RA3fzOlhTMFbp7ZFw3Fqp7APdeIHas7H27mQTLqHH2Jv4ZYXFn4Uf9RQE+rXKt9yhfjvLntE+t/0Py4hCOlDz7Xl9UbJzfqAmWeKtACOrn0eCBAdb96KEFrqWFQlNHqEptcsOF7KKjSkSKPr2ziM87B7vTNrWH+LziVK7C9pthvl1185cnqnZrh4yOvatW02EaO1nXC0Z79NmKlWFv0ZEDIyMY/7ItZnme41wDcyggE23gCYblgqfO+Ac3egLbIbzx7NPacLMqluTaFiq9MNCJBuFrEMLscFBdSAh0VPBOSFBd+sGt2lDQkihIHeFfSmnSYHp/Xgg8v2lp6iw2L7pZocr/tG+2DD521Fe5GNgstfh+GvLAjqPoKlRE6ZzRcE38sodiyp8hNyMhYSPdWSIeI+961ycODXq6YbCLqUHTNHpEAz8hU1TtIXHPEbVGBPHAI3DO0m798XO+XwsYeBfmKZ0YGLNXocvG6WIDWIeK0LPKWuqHB8vE9ePiFFDh02PGFK0H5RpD4rm2et3v7x6a2tmgEjynbO0NqADbEXeVOa2+jzwT4PNCU0UJHzt04RaV3cprfxtE9kHmZKWs8ywqMNJIrWlIF9+kGacv5pdmJBX8lfudkbaPK5wlMIqgWd2MvHubZrfm5KmadyJJWY4JGxPbMfXN1OVJY/fpKcXVbrt3GFNqHdDxQ6Ndriz46Z/I8553c5j68IzlorHricQXC9+El737U6SBVZ425UdJBxIP95vdwix8RQ5fQq1XRX1CQ0yuSFMVGygxdPQ0CAUdF5aUH9+e9raKuqwrGqH7V2UBL56IBR7u9qMxGLinCB7bsjn9dApIDyd8xR4i7/tSnV6LQy0+JJIeSdBJ7dTJSUfy8qRrvEx/r6ZP4652VF60LfWxgDBrsi6QuhjsyHnYNYMJRvOjObd/LIS9xgUF3tvu0XiPMrAJs8G4c8JJHB/OGdvOJwLO3Asw0OcBDCjdUghhsvgwsVJs4QqHE+N+Pe9CBMq6XtF58+2Xltr9m7LDfQEu8u8Z6h/LFcKIb8ZqrzCrFo3uDfo915U19CMJ69hOYoqkl52Sr9mmbBc/Uxg64sA3x7brSSoeTihO0/8OuMy/Pw14HB5VziSzT1EomXn/jfL5MhSx6SnqNFWq+8pj/0UZ8DVn8j129ZIrpkNIZq296Ztfq+RCODVqdpnWSEZjgPPH+QBOUGObPnDgZ8Y7/aW++iNsB3j+gBLXw18XbMfHJiLjfwmOr5pq2AYME99IEATCGP3Z/W8mqVCEBTVDCqV3QUzxIXr7N8gcHi6bap+H3Ka9z5ZvSRJTMHzKOZNPNIWjF7GhkK6cidY/kjfOL/E7rrCqey+SiWCu11ggkOL0LZfpj4VeG7k2LOu+rivfRfxjuJxnSCL9rdvrUXy2tisWWJq+tut2MOZ3u05Hi3uZxV4F66vbjlV+HIKHtWyvY6qzcbgjofPbmk0s5kcVV4cf8sbmry04VpG1nzMN2BH9NqOmt8uy1M6U0V4j+1PXbV7kMwSf5zPbnp4K/RlMydYvKLtwpBTvwrpyYQ99OvXpPBjURHr3ckgt+fPWomVMWt2rEoNUR67Oi8Sw4YcmjSDYhYXv/dBi1+pXDoHp1xuBpTdjakqB/g9YRS/fBupcmym/8h+GapSmSqQ38mSkFCkVR3Enn6xenjV7eqIltSRwNOlSxkQvY99dUbYW2+Mhaq0sJVSyJGgBFXEKFNRSrI1GmAHSjD6ZVROumWuYucxnkJyM04NPDcRPB3/roneQyUwK87Y9j3XlxDefcZut4dvWC4qq8/3hyeve4ke1bPB88K8lOCZyvK0+pOqnnq+g4lpP8nrh7MP59+zK7k77p7wKxZejytxvz5YaK5/89bS0nw0V7JuAgbx3y6a+PLli6GrHfpHcuEZL5ZxQ3Q0WSv+lA+d4CET8l9MO+bFI/VarGRTrzCUHSXnhMKKgYrmumPflfGwubm/vhJW2lx+/e5Ua94QK3KqfxnwF3lKc5JyF7nTefTrU+3C/y6UCkJM/UVwvOrHNpUmW3LXqP3zTPkpt653b0cpPVnWXLr3vrdPCt2v3Gw/+7xOSvE0bkxDh1jjcM/4WxENpczYobJAuLuExT653uhJbb+4xdegj7QCk3gqQcmpxMOj1zT+Ye3Ok5W7N+v3vTlBN2xPVT6HGzm36exuVXY8+HHdhKBhuowrsnQ6CK4d/emOiokpjYlzv8q/pbnsJ5utHJyPAx00jQQ5d6zE3w+27DcqbK03XayUvPCldOmPqdfSHivyOY+MP6q4p572NIkY5QPsD/78k+WHKy4o+LKmYubJMwgEFAofZBx9rTA4OBr7xOcLcOa0UJJPyPrxHbFqGkVOoti847vcD81R8QXxZcWW+3fTi2h3r4jURFTE31CRzYQZm9Tt70O4XN1h5K9qD+VpfcH8b38hh4SWy2bBLpsSln3QuS6zYi5g0+K1+0UGHuNnhTA4/hi1XqvxkR03pSLFFJ/c0qqP6WZKho7zU6jpgbQqHW5J9fHzFu/I2NvnjXueLBE/CMFgoGSuLyHrCEv33/uhjG8J638/YOT2jdI3ZDnk2KwdXSJbyKaeaRCrhB8Gl3pdncj3XR1SbB59fBlZFHaARkseEFC79ikiBaQ74PNpD2h35N/XU4PUI9Oiu5EwYr8X2TisO8T8eY65udnQkuszn+K18EJhu8KX3z38ezs2Pcm1dU7guxjudv1arJVbHs0PllZWH7qTvSpQnDUXqC5WrdB2rIw7iz4zyUfsCQgRuWMSG8qpnIhmIXg5fmreKXdr+/zSoWw/vl0L6nfT9s7/O6cw3J/BtsG5LpxGHun1dbZNxFK023ee0nzV3FtTkVoNmisc+OdztM/qI8TM8O3n4zvrK8uL52ovFIwlYVWEEItFF/MkeTSA9GUIbMffwTKvn8k+z16zKu/0r/VfzRX0mhMhOTz6MKTBy7pamIqaiUrA7HESx/geMBaFNOwiuK0zn244pH1YN7wl8I6mEjXMLPqS0bFM8dP5rfHbbw9EVPZHLoAqO3cG5m55O0leYwq2DyTv5LWx0B6wjIMSGGNpmCax3qrkutNsX+1cbmZiEjjovDygomTJ3KSuXGoRtqwUihM9A4dYcefEYHYsvFfxJ8hbv2wm57cQvvGcj+3lJl/FoormgIBhoKnzN+IzkY9oiQuTaLt94OlqxHDQE4a+hvK/Ha4gm3zVK1uLTTVbv2XsuzAM9H22BVXP9dfBpr2yB6SUxuQiwSUj9JXAhFHz1VG+GqNSl9qidyF4GYw+DMPtgbejHcHBdc7LtBa98b+9IJONIrnx+c+pW6MQ/F4PeeAHHYrhrI6fPUTbeLPDEoPNpwkwc+fhi6pTliC5bomEd58xDNZrTMEN3KZR2ck5PxOH93n9bkxeBAMSO7hKl56fqEMhHuDQbVysL49yNGHVnQSi2lIxMTEaOlaZ/NMYeWVluE+TsQjHcFOXkjLch1hZWccY5LPk/G8vJ1V76dFbPZSDOswUXc5OOn9Y7xEwk4sv3GTmV+VRi0zZCi3GGuHt6y195fmVk19qI7xgHB9DV+qyo5fe0cxDXUnArBOaIvVEOf3CnsLvfPS6Nwv+gO2ll8yJXvRs08DWyEibmIw0HzY+qpEaOcYv1pGYbid8Xw2pu86TXC0PSmZRm3p7V6pbLvXmDcNXtZz402ERguWg2/Jk/8l36dVvlgSNeSeOI5nRoy+YnmruRlb6zbvS+fMNVHwj+6SIZtQxKe5vpkJyqiePYKNxPczKW9Tzzn6XT3FILveCnudwJqRk4xJmOV059byT9dCDl/1ZgY+ywLVW4lhZdWsTgAQ9nk7bI+VTILmd5fYcWUS4QIe9MQF7UV3dx1htjo/H1kuZkezXFEQXnYfPHhvWL64z/U273aeXwg3UiGE+vW8CPKvEdM1chQcGBwU4jf/sPQvWs13E0hrKsVLPimR/WEh1M/r2XsKi29PMpLzRO+DQsab+9wJLEUPKVdWv3+tS33P34YZ/J295SDbUTTCcjwjEvPbSp4p+o9yA+oJL5JlV1dnTjvVv2JZLnw3M+G5r20YuQ12vm0+XJ1w467FjNPc1thiebFhv5me9xsvOISU9vmleotf7kHCTuUCN76VO6wtqbetYy/67t+i8HClmptPtq5wU1Xb8CTLM8hvxGuCy1/c6fBzILMi+ONnjFT3KU/sWaU1QtMYgJkajZxYoIOy1FzUi8Whs7+MLP7qYd16LFzkYfQ7kbj0CfVXNMsTE3FxxOqaBZnhr45TBeiU+9TfU37zS5rqKpHd7cuGVo/rMz0ww/NgK5MIwtGqvXlU8ufRa71P/C95IpDsDUWgQc310yi/CsOpKsXsFu+NDuTvbSfkvHXh+oV0Fka5l47zQ3Q85jnSA9DfdyOHEpmNzB22o96hggRLJndKrV3UhPMGlDE5Niz2RmIDZW4IVc13wRbkZOg4eUl03OvaMkXxUUGPed/WVqvd1y6td1XS7B86wAVEf6Nf3Eo0Fd9klqyMzR7f9ggoLTqWwg6va2tp6Ue0j6+yBo2KmlFrPIQn+2yb2Y4H2gQHzlDmqO2n/i71ZGKQ1Is5s7o9Qz+sN6SmBi0DH0WAkFgTPlucl6d0e9eZoAZpL30uFll3enlAi1RWK72HxGoGzRSpST8TUahbtlh2YroKVYiUlFYJjhbnIPMZp6ia1Cg63Z4MBflQOc1NtmwrpHrYvNuR3wvyiyu9nCeABAaKCDjO5Q3dxY7PHI82b2TGYtD7wU1Kf7cI0SrqW9vhHoaNh5X2DjJIepu7AufPPh6SwActMsc2VcEe1phXeqNaroQY2e+qRwbMQYszw+uhynYblx0tg6PVzCY4IXYenDxmOje8OeZGy0y/L4OVqwme0UfhZDp9pP9MEqA1i6L5/cat89u2n88RXlfqT+h7bHW9V1XjB2p35Rmctg2blkSBqH6FIo3ZqYzebF4n18nht7nLPCiqDsdJam7FzlgfQ+HBov7txZgffqhjfXXPMidjG2rdchHdwxbUsB2jmvHrgjNwRTeUy+lqGu/DyW2Eh29rI1d5yBx0hIeHWZDHRsl/WTDLb48+Z6k/YmoYdkX36gsUcmsXl6tfZ7FxoIQdcD4IqtUbn5OPZuMjIl/mk34uBJQ5d4f25md6WDk2rp5QLhebj8peitQ8zpeIeohLd5mqbEE9j18hu+wcsLCyfw7ZLO8+HhnQHTGQYouNugztdIridq+ngMH5TOHuDusGscfCxO6h8leG6T+qweuhRS4LbcE2BO25/kvlwA5ZZBgkfd1D7BjP8GsMgNEN7X4YuLwlzgW4Jf479Ta2EDNFymWY+n3luP7kBx6jfVzb7heb0XbKj4tB9R+SLgIsyp4mcVYLvYsvxnvj4tq+D84nbbQgtscjk9gev7qzv+7773+ITTVq889wnAUND0F5rFahIwI755Gu9o6Sw7KexqYuYL1OXaGJx9STujtC+e+2/P9+dtuc36i4wTJN4uCoQBPjslje6XG+A3wT+CLpKtiXNaEsDM+/+zPrSQD6Djo6V3Cd2+PK8sE58d2UuR/MFSR+GrkOlYYDbd9MwWgu7ss7J9z1AMZraYCoJ/HihyUXK8TczNictDRzXF3ygnGUBvLQR2siTyRioph+xWUAL+eHj3bs0d0KvUAQdI/L2bnaG+3sIxmJCzIz+mBgtALc5JjX8UbGBqC7A7UpuYCoOiEzkoCTl9nGnfbOm0KMexsBmM/NbAITXuIdV1b7ZUBQxWNNfcADYT11z3xwRWX+DGzh1GOyNq/0uwE7J4PJSYNq4oPy+NsPMO0oM4+Qm38jjerMO85FYbOeztQ3giHTA4kb9dW9zyoLVl8R0v072aLqDl8FpK2nJr3+S4ROqvpx15dIzlV1JSmuX7Tjt6F39bWRekxkrkuND9i1nsrvUF8IviawkXFHtsnVvMkv7WdrlLwowL604vgQsf+ACTDwZx8iUpTQNZpwC/Z4Ec93qc6N3cE8Me1XKHHg9CDo59TrL53rtqBaUhdUVXdK73DFT4Sw/PiRbchWMkUyE52yHM4+Y8YtVBwhGAfDE1WiFTG6deSpkoGAQYUjUDJyz3sjKafLNH5fPbL2ek9eVyq4cdDnxH5vx9W/Js74qGeJ95XCnoXQS+8IsXmz38H3hsyvFtUBQDXO4jUy25cfjXa0pjzpj3nepXOEoQoVDbVTBypY+Xog5tbWL2r+/ltVO7CYBiKydZdkGW1eibab2Y+RNWJDmnPAPZVqQw6wifvTbUd4Lug9hDGbVhNH0eEiBZcJRLURYC9Q6n4ugXfVWs3/lkNPQiAJnQXUQ0HnAlXDlVy0jjprpiBHcZSukmabxPNnemGxTntkxqUhlfGjpXCn55OLTjzxLTBcxrmZ+isWdXXquu1kLyriGrZwiRIP+tej5V40/WWfcTP+Rvzsj8mdQ55TT0Nb7huaNo+Q7Sx/LItYLtsSGKFy5UrwQN7Z/w1Am6Bbigtae81TRMyt49ZbZxgsr4wvQxsuWxLrvZ/C9LU67TRwxGtce+r+7cCQIWkBcfobjNXGKfMcptwHYwQr3yfN4aJVJ6YqP1Z3+GwQFB+AAlyaG4k3Sv5xOu7qcamEAPAE9DYDC3h0f1k7UlDDUzO0hgNI72f3lou+JzAhh87VTBmHcirU88L9nvagBVxig6qxAaXyJj1LPAU7ahgfiJwHpG3Iy0Qz/izqXFph9bO2onqPGMSL5b4QLbssUgTdjvsx/lgIOVU8AlNrMX+r+NwDj34AXzaONhCKpSSYgwefrRqgwwSFvZRAjmq+lJZzCGecFLH6W83NQAHzUK1v5zUek7hVer+Ui0UhNEwx1qugGhQ6mr+QaFX8UfXjTtOKZGJGHCrJOFKTz1A4RzOBdnqzIrLQ0H/45FjED8AHQOxOue3M/Hw9vPPiOwMiXvHj8+JH4x7kooQjfqdiKvvzkzitRn+Nir6WGAM0HN5zxmvLIPf3kv3t3+Gu9CdMPzn5Hf4TJX1X9d6OC4gGV3fiwyqLxa/9yBlm0FfqPkp0p5LxFBIJ7wtocVn9C9l/k2Yu2L/uQa0UaqCmVkJT+QVBtyQaP6NSV0WzNuwL4bKaW67loBvtkKHnp/1sgQy5QQBXBrT3BV4r0sg6/+s//TpvZ0u4/VXya8+JaPOD4gWorldp602Ntn0Dgu9mz5cesglOCaCR6mkl0Gl0rJs3atgtOXI3Lbulj7ebvil3WX5/45W9zPHs8L1yaof+ZOh1Yotp81cB2lC59JSOMEkBAb/KwZuf/vtM7MMbsr/L0QcGbLp7zpIy63e7BRsacPfRVW6Rh8vjUt1Wo7bHiYMIfi16GZ12+uBd/PadnkjNdltCu/B5QtmwAEf/ieue2bveQLhC/kMDk3RbUMDRKFa6pmUS61J5sAX1jzUD2vxziB9FsNenSM0NeXIq/qrExURbxqCtZKtb/ljWux2FMQaCqlEyDHhcj2aRkyZXubYY9Y2reNcs24xa7gDeVaz2VMit9cNtLc7GbRzaMUMGER5KtqVcc3PAlxL6paD15tf/4Z/PvzNpgxP+WKQZRDOkr+NSRqAmW4ufn1zA1NXV+yszMXFRSMv6wGzcTHllYWNjX1wfXTy8IO/+4mKP0SDDYaNmi/HPjbqe0cm/E70KCaOJpjmWH92X+YFE0pJcA5q/9q+0N37DiA25bPuNxZ6navAEu1TSN/EXh7CENGSkxee7wQgeB7YacvS7Or7fCw5oFECb5v/RcXGod9R/zvY12bvr3BSaG2/yhmuQxmb8s/ACBuh9T1AP0rucKnvuMTZ6dW6mr9HJ+Vv76dFBfj9bKyqqoqEj9Z0duLvy2nJjUkCPvmeRn6HAHf4onMv9aJPi3VdWoLr8OuY6ubui8sM4dBVWQhLSUVF3A1dL8j/uezjBMQsoS1ZDM1OwQmBY9Z/NObKb2FNDigOBOuqCggGrPo1ljlUxIUHBq/jApy8zUlHauzjYsKgrin/6NbF6zo7uUrtKVY9Zgk2WFmyzOlaOCV3OvaPB4TmLqz99v4nkr0k1HDta0JpHZIWqmQdjVdUzqGZxRPJw4Mr0/oF813tkUFlXk1wO/HrmcvYqVLTNJhobw6AL1b7MXTqDl4gP5ND7okeGCzfHaU7wd8r0JdM41sZTiz8Sb/vyVfFfCdwR/wcenSRH87MnhlTpzBx5f8ObIhroBBurS0SZXV98UnrhjDhGx6ZAfO6fP77jYFbSq7tlEsAqKfxC/NqMp7up9Wi1xRBCXkNDYb7QpJLRwt+q1BX92TDz2gVF+tD3Qf7/HzoGUV+OgI5HwA++zQa2UFprDgIa6VAivYqDr755k1oVYMz3novMFUpfnjiy6VtqrN64mdbk7O+xVh9hZnBD6eQBVz09mGH0nLwCC74R2K5n619xjbwaIr0dW5u6iHnAro/Kw49Oby8kp2mX1ibp4p2Ou2kI+eXx6JAt/L1+Syf6Ndpv8ggt/8eoIf2lT04uWIVnpDcdfp23+l+zJy08T2APSO3a8i16ezDS7f6sRXpFkUvNCrZ49uQ9DIXtCWRmnXeMTKMmGTYvDoW5yu29O8Bz3jGRMVollY/+3G7FzsocdjxeCXzKOb2swTXsppoWnRuMMnBeaCFlyAmrEikw0CW7rj/P3+g7TBexTeRVaDoS2hR0BRwO/+Y+CY7xncjougmNc/EZYF3g8OQVpC6WLX8VbXNRtp2fXUtDYuInXefM10LxwSnBuJV8eHPN1v/pm3yBCG4nemcvpEj6Y8t8eLCcELm0u+4Ncri6n9IVzOLVKKQEeB2LEWf4tPz8JkH/L9ivZjVTNu3jZFDPK98uoBkn7cCgRv/JVK6n+9uVpe6dO8jM4xmIKmWNnGsfF6jv+m7m17+6E723Cw8fHQ6JTgQpLyjWnVArPKk00NsOFaZvec77+n9Cwg97yxa4GPijnmU6xbxlv1OKDDX8DzIdV1lxGWgZ3WEvcj5rT2iPoJFp6XiYevPme+V5oyadA18cBc+cXx2miv4JZwSrEGs+FdSusa/7NjcLKkKff+H3w1zW4pxnfT6zfU3GOuENf3dLLObsUVcgcdIxOT1dWzE/cX1MhPr5eC3mCS21Zf1/Z1/mnW/DZOu4g7QjhU73/RUYDoihFqHb3NAbjaBujbofPu5xO6uV0HicYFxD3+bsXFUPyRdLfjHaFH1K7rbbm1YS6atMHq4bnrt1Df4grbrkLVrHtRzGKNWP6vZXiob1WITBM0kXwmlDLBVB/Lwmk+sU2NX1r5rPGYfzwN8Vw/Uygo2vsUb7ALOlPOENWVgD2f4ARDDljqO9Qj4FBFxia2Qcm31Q75tm/QKshc36hOUSeRzMarMDiv/Xr4fjKdZZyb7aNpFLClZuaQ8/R5WeV6YWPw7ymAzXmTM5k6p887uwGicWzKYeAu1Nsh8UqkN6CS4pYwHocN4NgmY5OVpuDGz0qD4AIuC/wWtXVkYoEUHj5iaGBHQ/pWPkQu/uo06JlVadXbhgxoSaxiB9go28x9Tn0eFhkxAK5HSxdrIPD47lg65XdEKNEF5CNi0Zjeu0DQGY/LuSFazS+0f9mcKWVcUC2UsJmYc+k9MMo4aFeLn3hfxcpDzkaK6mTflVkRWijM9EXVaTvmEHq/LvkeuRWN+61Y/7iMp9NWX5FQUXXe+4cWTq3B63kuR/Mqx9euvM+zs5n0h6PdTq2bEu+DfNHQLtvOng7350xpnvMQS0SvWiW8zgeDqc3yMciREEGbOkEfVVJjlSZ+6bt/4MyNrCDf37BXHJnynfHDXGdt3SkX3VsAZ9xNpNHgqH+Sge97Y4Z1RaPwLkbvegzxl7XNRUh2M+t7W+pcMmmDWacg12IzPTaGgMrAwWOnqPURWfgTI0FN7tC8VixvPO6tfWbLh/r23B8RusPOSFz86T2VlzA/3Krjgn01i0wW/ad5mJW1rf6Lwoh2Y24AMp7jPDcy2gW+8dfDAIC9SHDZiOs7UIGRyLz97/pJaFJLZdjD46Emu41DmknXP+xXDA99JWh3suZ3m87ZTpcwCviGk8Ikdz8paOhfwMe9Mn/YIIRLvLFI59r+AP+fJGuyVsJtkxcgGFR/gzr0Qq2G6fxLRxJmqRQQ/If+wFBt0HMYRKX3jZsiXD9qRAEJ3ASGLuzYs6HsfdJKrKjD6osMNll9o+DnoQEaX3Js1e82+03+p7jh9S475u5TtsoGVa5+g9ER5Uh2qPiwRl+EzOjAloXSXH+kD4AUizz408+kuSCXIdCA4j7f0T1mSBkXNQtSxUvI1GNEaYEEFcX22OB2SYSQBgY/gPWUSbw1Z4730XM4SzTs6QmXygUD366suf17Hobrhg/P8ac/3+xGC5L6DLp8qXZWYFZczIMYQPYAX5jpq0Oy1VvvAPXhXskfftvoCcMKrekMe/1j5lvGUD5wKNqQ6wEZgmgLndPjfvPXvZe655qmNAmJJ56WMqgcLeWJ0DXg/DXI2HVy+d/ylU+HD+YZQPYQ3ABAXU1s1YXnSmW8lo7BaD3+uIhiAh87eZY/NhXW6NYWcLekpfiaFX7LtFqfU+lwADQgvq+Ea6+5cp2Qj0XGoCuXECcpYiKxrQvGy0CHPId2qf/wNduUE9kxO/2mmFwV0skdeVT2t8pEpMudOAhAxJP5zC6ZmS5pUiQSqgaaCNYJDbC+sPFzZxiv4TwCdiO4Wjv163LHN5ynnUGO/VLu647M8ECPgW5OAHXJH9uRvVsApnj6KLnAIjtQn+JaarCjAAoJNdxvsHgfczucqWtYMb2j5c4XV24YpqMfKXcirKyET8lXujHaeaXleFqaGicR49+sWa/1tMB5/FwYX9oi791pCA3DopsymZktKcrXGA6Ko8uLSlRNzEp3xMrLysDd+TQ08czprfRN36Q9/VtNro6GhthlI8HBzeqKEj9cQ3cVdBWZmg3aqRhYUFn5OIa6jNsET+f1Kr22p8z/LOErX4tUtVfJEyQcnd9v37JyiCVVch9unnJ8Fx76C4+uNvV6vb1++34HfXDxmL2Y2e8Q2hapxE5ukaB4j7f4WJZRk1tbSeeVgKt2E785A8Wo0Yn2cOlNkKa8rcXNW4brzJdf1KofKw5WOfQLUOm89olNGtnRdo9iCYutPryvyOgMxVKVPFx25vhJSYmzlXLEVdZ0fTWPFBfW11NTFcwEiAm9B5eUGyRFoX+edcu3djfp2yGMNf+03D6zsJgEuxwf6NzhCh66voNeZ77XrtDmrNxR1S7gQTFsDONznQYBLRlJdguK0igI/RTZmYZE8ED3uXh8nJ5iJ0FWN2kgzN5R/QuIrhIhN/JyVJ7JOAzI0EKv1BQsfbppuhAQiRirx06H33YqsY+iKylq7DW8T0SKTkR0xla1s3IXul2b90eaYsdhU5loq+F9k8czT9x636lTZbzvatexvRKfaijvcZXMiAuAgVtIl9OhJBdHINh9ynJ8+NTKDpVLdlJ9+W/RgsWuPFquzspBoTwjDfJjupK0ksiCH3u7y6u/j6wsRbzdDuw97Qe/i3wJhPgl5AURRJCCgtbIC+HGwXa0O1Iow2JpMR1SS+IzBlxbEGhk1zy8zUfnOnkVt5xX6yeh4QREhXPbzJ5UEXVmz0uRjKTjt3hA42ku42SJj+FXNNwlkuajwGuXLlojctt1++awAx+WoQQrJjehjYKheUWHBGv/c6xMJHRb1n6VkIkeorOskhj/D2zn6StIEUOGfhzQJBxsYkqfWPnklDXuYUMSNtD2KneOucBaOjD1qAJ1t9f3n+uKTDuJZGnG4lSiLumV3FZwc1Olxu9hYRjqKKAtIYWsXpEdYUB8nSjESillKw/+QSUfyB2zizPZS+EhCKcYS4eHNf03Kd4BjLYVpC+Jwg+m4wEV57fbxKqV+Nqr6FpmeAQ3XhA6f3sYI56qRwqpvIvfrYU+oxEWWhaPkUpxK/e/yu1Q5HwbV43Xu+FzzkeD0OYAYMbTJSI1G5TxeG7WoFs+lwmotfs/8iMALikRc1gIrVbOvNdianeREgw6JqpDM/xd2Hw/yypqBSMfalg68nsI2EiM+CxsltYycWHUb0kh1DLmtr3n+IBY/9Jo0vh3inBqN6/eJpWkzE+yKkFdNWtUqbmY4FUZBD3773zu5stc4KJQs5G1lGiOOLnrmS31G7Pj1/GtAb/31TPwTPdOXFXCMI8hAwcAwBU39/r4Q9kq9MCo+U2acqgxhQRAC/2cySDfnDghSFuVgdrdWNQRaRursSTOLbeP5BlgjfPzSURjbj5CRhEAmy4OEyRKkwrDwq5wZ4Z4KIEUD0pv8vIQuetnoYfE+Va0OeEwuH06q/87/btQ8avMbUXpzX0KQCTQGV7vL9ARHFTfarqAPM/TWcwfzrdrjMJgGCBdMGWWDCkKbogrjI1uXlE6ioKMwCjDaW5MMgLepSdZnkB2n0qhMfaBynjU7JFD0FI5YoSvGilz33PGeFJBxwDNCXolMb0E4FmH+Z3SnSwmSAvoFFMh0GKbpmVaNBmPgYYGoCREuhFQ/mH1UVgH/cJqtCG6IYiCuc0EtVnNw8B0hWx/9KCMA7Q9lYCeEuaHKj+qRUfeIOtUG1ECJL8nRL14orqPlelvEhbk5gEAnacqerU1NcvCm08OmMEc+pIQrDzPiBQqQNjZkZXSnvm+NvqYxchSLQk7VzMl4RSCrmKgfKbBuMTE4tRs1esGYKONcRkZELZkmhsWlpaVFFjdb2TRYJxIVTFOpuv18w8ZvBeUxL2WPIZTJUbTUd9PGB+6XdWZDWSLXYg6yTD+XfIt+hLCodeieUIo2jhneNh4oCy9mZeTfmKYPy8MLN2qnN1uoqh0ZBjA/tck2tBYV2dNol/zM6NVafO5h/uO9nN+VEIGzJbTwfUewnWzytsCkUsM8ItWvI5wwroOrgTiXMpoFBG5FvAbE4NBAvNL/acJVpWF5Uqows3J8c5lBs4sFCEESQKA8mVDuLX3OaFBQebp5q84oJcoYNaRyScc8yi2KVuaP8V6jbhY4rorrlRoNTJuyvlwChYKJwLfCXPHnkzGSoLWpCGf/a+dMHu0w9QbrovchGkSz6VBp2IJUrd8oDo/wAecbBAcXB4KoASaIhkzht+RLrlrtsgU4yxOSFXtKAUj6KR4crz8qB2liyFdfl/ESWwl/bIFH4KvuRXAMJ7o1C/D4XjMCLtmh8mtjHZ2JtJK7A6RkxRomMWMaQFsNjk7sUSIvibFANbCYdbRVtM0XEWuLI4rH/VQLFPrkYVuhPDBZGLDR7MGhN8IAJjdgaCZ2qX/0hY5QTjwJc/135kRkwZt2AP/AwC+alpI/X1UZkLX4RUfc+Z/hMOBnIzoLKXO5TCpuzXBrwzBM5R/1EPaiMHfpWAtpcorQBVR6AEixhBqc65gYddtJPaJcgP+pyXc9f3meIc0a+aCgB9LWkzGx1VgFVkXF5jy//3mi/0t2kuFPyRHiatZDWKQbpWYHReqwAqndQ0/IeDrSPBy0/QYyMKYSuxKLnMPE5skWNIHE1oYB6cn1Q4Y/oEM4cjLq+9aRwhiUqScDzUIqmH2SNd8aJz3dHKjJKtYsyROh3EM1SQdA5q467Z3o5EWC4pkAJQ1CO00qSXnINow7FtcbNFkpigFuE5LcqFdlgD1a3sRAkU2SA+UQhWRzGwcv50ApznGmsr2Onedb0YBXGTcET5PaYJPGR1qRGDz60Uitq1rGChUQOjEsQqY0i81GmKhOOp5t3V/NzfQXU3O2GQVfOXZ4wFEnFfTp2Fc7s6icCWgBUswyfTjOBfX4Y4B09Zo2WcmEA9n4zGLr+4H3d1LNszYSpOMUA+kgH2NJQUoDZAlKrkosMJPnLOpr+0Vc/92LAxYGBiEF/6FV/yTe+PVCIpNydq0HEbZZqqJhbIFI6SWVOAhv7zWLgkhnr6rIto9B82EOnd4dhgdckJ3y/fI001IqU7Ujgsyvp+1sgQZt2vCXWRg9l0v8NL7eqS2lRvJgK3xsJ1ZKu8AnxdfOrTXmI+a/nAKJ/PYPEySa0ZV4TAymLjtpbX21EgNVCcjEOvInzWoQJd9dUr7NRnJoYcxiAUWMP44rGK2SNfajfcCNDtJCUrgTmd1HV1U5nFxBbj0uMSamwnniq5urquUTxP5PvnOchV/JXj7CbImz4YetKWXGcoPx387rnL+VvHibiBUCOeVsxuttZajVHze1q4rzste7SwU8KMdz7RasvSabXV+WeB1F6nTAeKdL6hGqe1C2FfdUb1TX5nUfV7/OrYklP1Kd30wfp/qpLyw1GgcbbEZu1nXOr4uKbOXsq//pMbmhOyhxrwips+U0gJtCoRZB7Mrk7/iqUSnrno/UL47oLv7aIfx5KvN73XM5N3GbRB6PWbD6b9qunPSYGOOY9qz++e4Olb0H/0WJKHxwLaEopdyjBDYJRuF6LJXiJGbRhTRDq0zSHw39nnF/qnfxfCvacNbbBWaPzTja8gqEzveNSeI7U4uthkzI8mD4lBpTuo8cCobP3ncIdoR88PG0+KyHAZBlH9i6Rb7Z6azhxH1PY+5copoSLE4AMkpuX2QFf2R3q9T/M9YqJHS/ja7+L1qn3afz0Elmi6X46jrFABXU6ScjghGCat6YtliBsyqpA7WEQ6aMpHCp5kHL0mPvdVwA+FMxKi/A1jSLLlkmzuXtMV1tGz/K7yHihr/4e05h5kWqyYJCq1snOZXVPokXMImC9Vyh+O2AzCkCMb4d/crRT8H8vAIT/pYSpOay7J2qZysNiA/FHJUMpqFQCVCkrAgiLgAllBJMrzRrHctWYk8HC85//q5ajdAsGAv40ArAXtLjdfRgYT5DDmh6A2jpItRSNqPvEFy6CtnOoQ+4BqFrUVZFZSjF7t9UKOKHUbJ2Q9N3p1yKPyYzBgUPvF3kGHvTtfqvpE2EothNxMnRaMqshB4ShhuiuqAalM2p1YZ9x/s9FBK1+6L2mdzehbGVnXY4ipCkXgAQ1AVfT0y+KRABHVXUmnMHqMEHE9IYZEkhBaiw30lLROZ9SdQSzrmiAgTvt3HEq29ud5ALZa4olGQtb+rxNl78v4BrYmGzmUvpQSoBRmMjJe9JQGXJSFgqkVkdBoJmAkoCOqRk1gpaRNgFQ/FWBoUd7QxRchO3ZhdzMdNYKfGJ1GgnttJwh2wGfPJ8rgUupHSWo0Kg208zlaU2CJbSSsL0vC63lIwMaM6pLzGnzB3vfKUuhACdl/2GnoGsLpJSVxQgTL1kW7JrF5qJH/AeyXzpAw69eAqkehTkMcylH2G5abRWUe6K7SY3cCrQhFvEOVCTTILeVHRoUKMDBYDgVVD8b+DCL4FUwSwPVQEwXokTDo7o1VRVrggymmLvkNNgi1QTgE6N2eY4eIE1l5wwsQyTJ4jMipB7GfozhPT7+860wif4hCtM4+6Nsz1NIA56fgF5T6HgoNGDMhH2LkGKFvB9B/kfj5bd9RxRUwnQWCQwUdZszZODnPAhVlpSRzacElHGHq2TTRWxVW5z+S5NnY2PS/exQDNAhAHzZmeSIapT2rMhV0CJRPigRg/3/bokPAR5UMQZSdjKEayQ35WFxcjPfDZs7CTjvjJzJu8l9D2/L/EGRTTpGKbGNrrCBVzP/i+bNqquJ0of81thUnifmrqw+YnYv86CodlrYy6nSk/terNm/nVKsz9jOOigtlokpmOCrNSI3PRda576ywcgViV7/gqkLWDzFPCF6g98fB1CVAhS+RhMIz8gpQCAz8A8KDYBdRf085oShaj/HKxL88c0wR7hM7PYPFKImkg6JSSmIyqW+LEJRX52HvTszEUu8HNlmRFWyH5QmaTYRghn18wUOYCISbT/kjP94g3RBqT4hRfpteXTb1YU8h/3vNFCMWFAwCotIXEm6Iv8J3QY0JHjSiRTrDxWQJ6HU0Evi30fPuH6RC4QBUAUcAys7g7DXVp9wD2F+SlAx2ZXLPmiymuLOCs3j6w6YS/nmOYBolLCygkoB6YicZSB7JiJRX/KwCJAgDwCBRbtpLIQg4gua7J42sLIl1CFBj+ZyUz1a9LWEjDwZlRVELTN3OcMD28M4ALSAgAfBEHfLREs13G6hjyVYGG4WiKrQNhN53aQcaFVXoGPq6vEJfM/cGOnpxnTZNJp2NBQJGokIAVQwUIvcYkX9Z4LVYiVYhuneaVWOYOjyVwAjZxlW3V2Ng4Mg/tzHCm0xcDgEuYQzPvc7TjUIUg2daFBoBjhE9yQwoW/FvZ6jIqo/fCoBM5GNyAB0VyJtwAme0M3GlnPtpMrV/073/J7R1q4UzwiA5CaCVB26rOgc8txXJuyuk68LB+plYEnr0FMQHOHr6uWvQmbCoAlTT+OIGr4NLIRdnGJ8ufv9L4MObI7KtytCPHa//FenXsUxyABgS/49fkp7+MdHKZkr5y4AXv0dw6FdxQEP/ngt9xrTu76GDVh1g8ExQ8K6yIOnYf8EsjgkRNKKQf6+3zv9jP8cipje5MJTE8rykmOooyUg4wUu6z3zQsWVwVGKEoSFSK4D9y5tQ85bqBpNAVU7CInEww9qOP/sIGOUi0f9tUxAx1hOwA8TSGeMkcC6sDL1ygY9EUTFaQVn5R12HmCNR23mwHjHItqPuwe0VFjpQFjUrnLCqq8ETNvpHZpxzIglROZ56Wdlt8vhAgtNYmOf9h/JMxPFpxf4C0aGEd0Tl+z2xZoY5UNtP/4QcF5SfwkVREFg/QslVTBo5c62XEeAcd925K45EYj5EKrJ4SC2P8fNsdg5W+Q3ajWBvW1eHleLJXw+IB0uUihFOp0EhXBwc0jokjqEVkezs7OZ8n9SNjY3Vyr6K8WUMsXJxRfTkHyC9Uk7ct3VOneNFtVNDdpu+fzekUP3XQ7ykdUGUWC+I+8eCnZCk8MAbedjR4gmlpY3njsFRp1CmPYxif4Vv8edRXHqkOsH+0vO4mMNYaT2CoPCulGPs2aEaMVnmmxejnThS0+Z5JbgLSyeHcJLgZxiNv5VLTlYeQud33t3gqUTJBRUlLbLG74jehQSPLrX5cxl/OickADT9iX+vaU1poq0dvrVlEx5eNNb49+/f/MDbi3OSWsUR7VPrPyhD81nC2ykiIiIqRsuiqgePTqjijrY4dzxJk9osK6pIdUNBNIR6mFUdsr6xQS8uvuR4HIlLAcKJU526O2F7JrHJrn32kqJ0qlwy8ujV3nxjJWlOnnTkm9O3fA7qMuzLebmhYTZ6rTS2GozbjZm5RWC7YOuGYX7B1Pgv2jKifRB8fDkZmfAk6ishaiKigo08BrSUw0QV09++V3jBdh+pvpSVcSr5+fmFHX9+8nriqwopJa9AhGybLcrN7QYuTUb+84DsMRgTpf96rXdy/3UF21dlxsW5nG1pDGaL+VNmtHZ6/Mo4vbqxi9lIu+VKTbLAe3r8omsghUPslZbVs+o7gcJGZvaGrI0QlkcLXqHEz5RIjDfP96uTvIvy+94YCPS+uXZ416P52V4jKTJvhn0e+3/N2Qu67hnl7wWdFvqS5f5/RnsDtqF6MtAIW3kKKjgEky7hBfjj8axzI6PTGhZIzy6fUcQcCp3/qJAWGSt3TM9gPkq22ZvyXn5jvnb7PTnnQyNA48Be/u7LR4avsNE3/5xRT25a7D1JSDwL0EQO9HI+hCtnPLoWB3iQ8T3W7uGJ2NlvjPQC/7U5nJF662yeO9GAC1Cc5t0nbHxAqRnwZPxMk+kMDSPzcIT5ySw6MEcaVUPF9ZgZEiVgPAe8K4OD/Svvf1krY4ZYjQD6pUMmDvIVS4S1kg1H725juO0i6nfHG3XTl8FHZcG/Y5CdJB65BBSoYpMn3rroPOno4VRLeQuj/0698P/+C6CjUw5MgiJzlEi24hZV4eCCZ3SORjlCJCSqUtMEKy2dQe0dhFvetDx7O38NSJu28mMMCVb8SQEZDzjaqdtW2HnMeAdnLdkYTFmjfgNszITgADDwGG2FkYT+zRUdiMVY4Z4IL+8aIBxnA1WXTZ7KMkMz+0CXbSFyV6nMxeDq2iDudrO27Lc5Ygh4erFPpyoirHp6uoHb9vyEDXJFDWr4OrgJJgpBoQTc/MbY5NXObX33lg8AackOCvIF6PPg5JbvSlVAzXPoBnN+iwSLDyVEd7/f+tzxSbUJvBIW9O+NHeTuvwf1u4nY3XApmHB03USYTc2riHjVsv2MrzC0q8zn6EYS3yBT4lyJR64muiAGC5KjWXRknsAw8qhMoqFW/NZ5IqRB33jvc0T4P/ua/K92JtMTY4TEIz898nCi6XdW2YIcZo5OvEziw+IdCxfBEkBm5ufpCvmEESn4z/JC0O4bB1y7xdnTqy1x0x2nRgrYMHZ/yZcp9KWhnF8RHLW87Ook2cAIuA+uy4lnPMPpE8H56k8cfWLO9NAJwvTjvDXf8lLjOQ14y1U7S9iDeOCNTUIp2/7t/6nuLaPiaoJ2Udzd3d0lwTUE9+DuFjy4Du6W4E4guLszuAYJLsEHCS7B7Q5v+K6sddf5fQ7/it67u6u76umnunr3YNsQBsLjN2LdOohhTG4Q09jure2oQ0XMyxGwPDdXPrncowUHugsTg2gwX/Mm56+Geb4oKyfeHScGdLF+MAryH5MK7HBhDbe8BnXimZgPaZLT5iD2J8777L2fhddlRYq+2WeEE9I2GsXzbLqEN/5Dkvf785GK/fwpgD4c4eMyxC4CqVS/M5JsoSr+uRhl0O/NlCnvl1+sDRPKDrbHK4/7SyHFB3KMm685MdrX+JSWIlTRmBpDQXBeCbgQAPnnuG8Z4ngEMK5KhPA4G0D5EGl1aINs8+FnkkCm5fO7CAJRReBN+uQGycHvU0gu9D9Uyaezp/jCRasef6ZDhW46FU2UD7O9DYKq/VAXn6grbzwqj1cd+8fVPW0yDoKUemEhHSkpgFyOlIy+UGSBaPQHevOEESGQP65f0uAgv5oTyAQHP7pu7HnDQW4PE+QAYkSnYTJTexfYacneMZzxP9x0mbfgfwiqooBMnweWArTIKYJFyDrVv9r2544wT68PM327lmnLf89+yObvhU4bzB1J8vJX4ZdZ3cGPdRytUprXqwX+H6eyg5H6/eOhwgTVXUt3S+SQ/2Bbpn08qVlZVAI8Fc8oBbQB7sIlp1mSY5vvSi2Jbtqej9Vye0DRorWiCc98seoydbOAGZYcNKmNzrhDNRt2kP8Gnw2iUs6R7lHyy4+JF60nwvLHpT5yHyUiCVvSI3mvLeSqGdkSsHl5koIjB0/STyYEwwg1xBv11jyOejdW8O6JNvxZpSzKoqIB6FhfJBNl6+N8N1ZLk0S+ohEdgJhS0ENQCZQ/mE+Z586//EEvnikpJvUjf2J3OwHuQ12LSLzw+LP5QS1OBoeYeTOU3z98vaLGXoHJuuOYaor2Z899PZRFAm6XIVybC+EdzGc2o6EbuKJO5I+f6AmxlPbGT2kVuw0aVQmdkFkzE6UeCWyQBGaw2ianCQ4RLNeRsL7CkM2h821zbgw7HGxofXUmH99P/JzlAbyF3DjzXp2KJ6WLeef6fB00TSEm85pvQwU39Y549psqBXW3HuTkSzx0oNms1nwdYh7tuVIDzMTDM1trD+bCXlWvDSJFhhT3Jf0f6gwMrBb4yhUtezgeX4Wy6K+6hEdfXr6IoqLtp8fe7vd3zacDWRVek6p5S6/+Xi/BRgjHBHW0GPcessPdULzMfx7m0bXd37EPJDAdJ2qI2nB+ZWT5ASq54FsFzogqpNpD4D3vSJzIwsHCY2vI/HX8V1UJfty0SuLYzwBpltfbPgLxxLoXMtPxAj7W4cDMUnZnBLgYFiC+Q4XzLKUAiNpAWfMHNM9wV+pKMrkC8kob5qrteTc6bF1tJLaHHUUJ2RtI4f8WOxVZmCN6uWb6r9JFzj6/u2yuvDv8fvF6mI814pV2Y/iOCHBgRpxKDp40CAZ49OiU/RX58ssvTjzAKO+bL4/TjbgxRqy5nevIDU5TUgosQ0DlEzqLB46cAIBY7oIV4RbllaWml2GIQyEGfpxOiwcEwUGaSwzBnXnZdOf7lAJ/QAaKqn3+zFST0SRMQgMtPk70Dam48i4xNiDh9cABzuvrpQcYJuDX0SSsMc5+nGCmjo37Q7kD59UQYToeDSXlr7+a1HLMNP867uLlPso4hz2G0S6rB7GiF0/QBYv+eeyMrD3DgukIFdv/Si2tYmhoyA7c2jJNG7c8WgcC1zwVvbQ1NcPAZCMk9FeSDKROCaw6o2h0RQbABfS9DmVubu71KhYnJlbWuTULDQsGFpZZEPum2UOHpPnV+uVn45VMWQ7csG7zuN7UhdStJOk4HgF4eQ0EpmhqVUOMsAvTOiC1850xFRG08dZU+9d7fkIQ5nlK3zDPX3KJF2fLkIGkRKnezzF5A2dHBRLLQv17lRKw55X5gVIOuWPGaLQa4UHuEhqkLpEaTbjJyckbTUqFj3d3dwioqJKSsBptUFvPVNzcamMSkD6xhXK+dlzlE33ktzM+WFEREQjEdGzA1t476V8qQqmRqygoaGjzAoBkrxLQ+ID6OJD/QkUnrKyM6ZMHPGNfJDw+JSXl6C23o4c/u0PXxbX1qlQC5pOZZYUnFc6uKTISOpzZ8fY532XrJ5+sMDbPPxv7/gJ6e9p8Py3eXVaJX7u4M2SN9fTrBtRRPZGH87MEdeBl2HCoxMU9/YUMwqvtQQdz2AlxpzVLgZ8RKMRU+SrvjyxDXIUkuOGtObyoqUp8lXVIkh9bK9VGjOZ8PpfT6emRkhk+CNcP2Bw0g9WqK31VK5yOjm75rOQmUzoxW6DURTElMbEgO3StCFVRFmcKTPiSRoE+Nw6Fj4RYc14lu5P15PN8pd7xLS67g7K4gJy7oYEBm9PSBSPVe5tGAPTt0TLWRymp3izsSEX2/s0srYPJSI0H+OapC/EUB0d8ySd206t3R5ZA1qeOlYYZqzshBG2v7Yfbv0dD7/Wg1QhRCWQIVWT00Jlj5cSQ9hJRVn7FhZI1KPT4HtA2XbMviad/DOKMIvBhp7rPMcMuWPsR61VyL6yxpMkhUiH76SqOKwfMIvcKmOKPjuxc96cgVVVVY0u74GZKqozwxNzInntzM5Ffk0tnkDDLN7QF9z7TziDRY1Ulh+AlZQhB5DJbjITWRMunPxNhwaqzdx6UlT7vjkeRzQbIljYIGav5oGvDOEXtJ2aG7fLjglOQ6tz/QEPNo53l5W8gkMrNou4rQUOmNoBOxa55dK0oAv8IgX6ycSeIZQP/8ziRV497EsBChEYIjPiBpH8GeTb+3vv7EsfZniDY/8qgQ+Qyvp9qi9e792E5tptnVOZ+/Pwg9dmCn/Tlhriy/XYvMICnJS9QbAKSKMb8G+u6DAVdD2X47iRfahP7GQ+c+6ngZ7wLNEkiquHC1+a4oOqQu1Ntpc2RZWCsqWKRRWajJ5ieSaar2NMq18knIXjYnV2fMRtMeR6FvnRJZ7uYCb5kUIRPNeStRPandlwN6jTfXbaKtmvOzPvrOkxFoHPX8N9tQhw8/rKqe1YQeU1esmMkQAYYqSciEFojnJUGHdnF0y8c0dYEWG4arSkUnznMft1gvYYegggeOnRW3+jz+ksKm5W2xYEgJPXlyPfGRTsu8XNQJ8E5ggWocvquvARikxGB4XVPKgoS7zpIzAwzojDRmR6o0ORm3m9Fd+WrSpnQlCBglvE9950s74kti6Py9DmP+ES1PdYSrrlvcFOhEQV9Kkp3r8xTNzdA5/IGRimS+f8FkzCBywgNp9BZvVawKoCxjIBPXKNj960p9kbHIWtolIofRS+RZaruxTPvtLS+QkLpwcntndr/CM2lmLJDiMA0rhdnJUSEyqFOUCb1F4ecilZ/esCfOcpFgbiUfeg6I+NCoTkO2LVfcGnKKda+ojrm1Y8EXeOpJ+FSdPiwf1JM0KXyQU1QSAl+vT1kAytQ+gRIDGsB8p+JDMJc/049GJMZFmmFFtZsuTUQzrMZaLU2sFP/HMgZLCiws38gBQdRpfKVQ7wbHsOY+cXv0E18+ebHyHVnd5SVa2qXr3u3sgfEXtcaSlBV0QTg5ZNfQ09c1iey78LBNGKEXsxCDRGPaAlOv1B5ecYpSs0pX3k7ue9fyrq6Sw7vBrM8fUXRk/WFevFJiRazE6iOL0841cMGqk44jk8PrChtxk160BCOqEIdVB0crwe2maHM9ZpvPFYH43yJbQY9s8NnliBBLuHoGFv2NxWI2k6Zpv236sC/hd5hdmdDXU8PCeLd3YeyCRijWb3Xvj8AzHiSi8hQz9Hdodqd3o1PgGNy4oTUVlhom4eTcl6hQKq64ium1hFlOAhDSYM0RKg18TyG1PkH2avJWa+Ak5H46kQlCsDz5ExIJspFUMkLK59MRsDKXSSU7nXDtjpQZjmFI4PrwK8pBzP+A9wURcCIwiHWA3CmekcNTET0KVRcEAOE7N3JYjgLwlR/Nw6zuwkBLqbTCURcOaSyuEj2FNqx9LiHDYqzR51fjnk3U4ir+iLpcyIwi2MHYpWBtoCDP6EoEEKSCq/7OgEQk8/pi1Circ3+9g8K+/AKXslwEOZpQ652xUHi3dU/dkVb9CcDRGOzgaxJc8/n0Fft+l8wVHRMfajWX88YYUAVbS92+OIPCKAjcr7LoyMCyBkJXQOMCpRHJVtOLi5rz8K9CN63kD/zqVcIA87D5M+rF9Ylw7GHNv/Qkw/InoemPcRnOi0hly67KIuCuTujKZYRhkqMmElcXyf7w1b1n4nU+vhDSAVAPJsMtCNb8U7AkyEJDAvNS99S5shiwCN0vbKx5sOK5XNW61dvw5L1WuWAGPJYZZAuxutlHY9RgdDwyy/wKPvKe4x/+0/vLG5Nd0g+zhg/T02nfm2jRJ0EAPGNOq75HFpvOEEnetFcAhXbA/gjW6nWm/eN3SHSZ7ekWiaDoalbneTPU70efvJ4vbEUnFB2MNbfPr5ucttfTxo+Z+mfzg8npUJaTkB3tX0CdAdDNCHy3cQSYQq3PkUKPDxCe0edLd4LAuJwFRRlLmzeGVPwBt6123SLfUB8TYyiwkTV2uJPXtt+QojDUbcOvDz11ROse7hQ6/2ak1MVCgv/qzTPtdE5/DlTv4UtUek3yqDyul8orSggch0P4GBh3eTvc01LBDF59m0xKOATDItGT4aL8MSPh0hUK1xM+KwXyIECQpgHxOGlKnQoz60TTta2zHnskq9bfXBpSoW/PhpwA4TDnmWHnj9/1A1YRNKNfz0AF0iEoFC83jjwB3i7IQADr2FyLboXBKhqTb6PbwJU+xO8s064FjXewKgoTCSVr3zwW8/xx/sgEeT8LJm61YFc1Z3VnFnklggpFvv5m554Hi0mg0BaicXk8+2K7tMZL42yK1Ql1XSFaEjTTrjUQVigb5RxY19VCSVet+MLRNxpI1ckBTqg4preFuDBkKSO+5UGRjUTyW93S4krkhoBhgvKvg11gKbplF+3YvyRmsoGR1ld1zDuywOAkAudrLtocrc7qRGrmVVGcov2pZs9OlUliysJhxAVphFE5gcO0albaOoo8vACR7EOBiIP1J7av+XZ8xpPR7+DoIgPVPCaRa+vpzAcdExYXH45y/V+l2MTU/RLgry8ERAcWYkpoDzFRvVziEHNoPCPw6dIegQYOqhNmSMpvM6/PputLz99Jn799Mt5tsskjYuplP+a9huTV/QK76Keu/FshY0kHMOvIsAP8LN5cIe1NScR7IRDCp8QoOQwjHH93y9Nc7qR/gD4XCx1sTvNc9/TdAU1YZDR2z/n+1wVBgrcVeQJhxKWW+eKHguaC1yhZdwpz1YL39nNDBCtfAjHi8Drxt6ncdigwiCRI72psuD7ZN47BeyIi+xr36wdgzfbjb3Z7Gp3oHBcwZ5ia2U29DxicyVwEvAekN4+l6ffir0R6TYqyg3zLzaArpvt6fE73S6AbXjeoxyB6JbopmmjyEMK/N1nLiH9RI/t+GKLQjtd1XGufOAP9UEmcTIsUHRdMqvCN9uzLpH85HEpQXllbzL7RUdJ9XcXYoFlxB0HLvk08D/wcjp24zW9Pkpn+5rmPVD51w4OjSSvUvRj9Dx5f33zWkB+A+pinTC01qQ3hJ7HiXh1k8nfgpMoN2LNjWr87PgnQVn5qROYinUflEP/B3a6WPLwoC2mj7TXNP71nHbGnyofy8rKnB5vnYudDtnnGRUcUBBAINDGpOhLf5FiptN2LG++/UpTWPEXRgRP9MAgagH2lefnwS+WVuLYOdOpV+XDaBQe6GfbLHXYnmPdzjvIbeFAjM/aaUrUjYBqKffnrJxUNNO2GdglzCQd5N7ND0FiAdRns9w/LMqD2Z5X7PQfeffC0bvu2Mhc0erWaf8KEV523cKVlpYS8Dpkka0dMv1Ubh0l22tcMKw3ImbverrWE3p5vNSuMYbEwnbToyaEIMDMP4QS8LIYz3ivDfQXnp6e7oaGIDbwOW5skr/Dg+mIE/OKnOBS54W2xgKQwWRdkdTcns6SJ7oDKYLd4acoxAosGpQSWjU9fmMrQC/PyEhOJU2s09m4Wj2V9bZlGSOnvt5PZu0Mw9ayAXW01km8l+1B5n+7CU9DS0u5AQ9RWFio3eXbultwt9ETzOceE5MW8l/02PRMPhDtbZgX8twNSrVGDHC/qc0lI4P5SSyqevqYDjAX5LqlbgTyCgB28X/liNoYq/4AjL0eGAXzlR7XkmG/m9XiEW63A3V7DIy/W9HsuiJ3oETv0IqKitlr9rBwH5PPn6SVLCwsIvKe4lh+IDGxsUkpxJiZNLa1EYGjmOWY3Bsrno8d/r0HL22g8pKSKZBjhLJazn87umztI+ZBZTcRh0kvEhJzMbbjlvWaL26d0IqlekrE2NPJIPqCzjyuJC3B6rn3oR6SlAnGRDFxiOkRaJMf5XKb+BodyZ47OaEEFXm2O4tAecXHS1UkjKv1mGR4ca3EAkvHhBkTS+kS4+t1lzNF66WgyVP6l4n8/IqV9xlsDqxoj8TI+W4/ajgOL5QqCRSRrU+a2mym25Qq6+L4pY12J/0IlV2HYgiD88YzefW0cbv7D5GKc/3brrecilucNj9kuUzw83n/ZYiKinp3+hGbfholKMIoaB3wstqF9jU5udCthD5LrE06P5aYN9/n5kRGZuoG/uXvjHIal0XUdQysdp05Ip1ixo/VS3TT58e7FtA3eQxS0srAZrdDzYfzSPAbc0RX4Hjn87lWlX7kZl+4rShVyuVeLqDpcrcHKbWPYW+50RaTkDCypYXP9/F8aLHWLNyPFVhTo3LY4nXW088EICb7mXm60uwYz6rXFDA2yLc4r9wisfH3UYidXXm7z2a5wabgfKtmcVHbJzwpOzsWBDI/JuSFsUswGYvm3LzlWXuktozRsaSbaP7xA4ZU8NuLG1WV2GAKYJ+HX9B1ijeufWs8mFTPgMNoYpKHYsuKhRvYcH5+zjE60sFf3QYy77D/o83WhYaIWMvrpC+1O5JNNzagi6cn3PZtJgLyNNSCxM4dv0OoD7XYHqLntLDyJ8seWMdfVQaCx6Sma62tre5CsW4DX7hSFJratlqy4u3PjEH2iYnbFy5suapK/gAdTT2Dc8EergvBof5TjAVbCTeNDPygrXws7Q9jCb/5DeXHiJgm1jW1stiNVog8PA6Z0pL8CgVfPqrqUEpdy/SPkXfF9YZg6pKt3SOwCgdkzbEX6N2sMZx4tDcvzm/1FZdluy/bNPZ24nbLBwwW6HUOpN45+Z2sA41s5TtGFoaZqaklbKPtgKjb3I76Liu/mx2NHh6OUXOkyMoOUy8/ZYezX+51gfKj2XpjCLkxPqJaBc8YOWi3fq1//3tw9GL+PerHqWJlcp/gI9YuHpGxFNa+dKscm/Y5AUXNp/rr4LEFDOdw58PK5s+kbeUtLC9T3yUMhC900w6E1C/3lDzXu/z213NdZ5Ai4BGEIP5MD/jU6xakpaen6+7ZHu7HTna4ddoYECkRLQ3pHWXdqmul1dxHhjuPF7B/33P1ozjF6UHmtU15RC8YypZHQ8zv7B5/j0wPzbCgmoGjW19pjVVPFNlKarJClWwPi87h+WGmVK2YqNVnLyU2Hx6FwJgc5BNeHhx9cL+Tpb63Wf7odcHr/+ipkMYBOSnyYLexDHFmFwMP0yQOORmBz24YUw181s1cPY+Z4OPZWiKHHqr30+oX5WvXKCpHMERQez15XMG6sYpDLeXpsWsTnGK1LISwsRSJZKqI4/HMYx8RJWRqqmmjy4fbkb6X9GfmLCR/IvdOK+Qbj6rJ3qTXfwnLPYSa68j324yMjLp7kySzLCnebHrRxcaNJ6CRxEYAXth7smg81l5H4FNRqZNOJTPNvy9BwdhGDdXxFE/f0YE7JcUFUUOpup1p1TdTpDjzfPfVPrtD6l3rv+8CXr9MJYM0vNkDUIPX0Qgp4QU7fbrXWljjuxMoJXQNI+hkupelAHRNsNuI30QT5YhRAL0DlZAQetz1tDfh7OvR39Ve4jafXNk/6E/linoTidr/lODLhnNfEfbIl84X2CavhpAnzsefyI12THEQ4yBtETm+F3SekBXu6G3jT3+8vZe/dX8QGYdeWXcAwFihzHrgvytGWRsuZzATPvH5zJbjdcVK87F0CeoAld/2YrHpr23HA9BIRd/hw3qP1WaPIHs1YeHSSfMP//AboGEQhQDS2fShw1Bp8CyBOI9e4MEIg98TVsf5+QLouoGEtD2+dyZ+TVvZbt2nv10GvLBHGO0CQxQLV51O3VJvcxJQZf4O4V+eAWMFVRR3t1PHvWGw4co5eezEYqRE2QTqiNkfD1FBeBp5fQkzZxvjAIlMjsM/ArG7y0KI677PZAptB5wPJ3TyboQ53luIoFT1fLN4bA6febezH4Ru5QOxcQhSB0PVm3HtRr1N1HGNZ43VRiGMC0KBIRMGTk59zXwWiwQ42ujcNmR8QpMWA9I3m1mZZOJU2a12/F1kamHZEbTtnc17yy6RQsrn230LZgfHN/laL4ERjEuBiMUo4mzH5DeCZWyDhVZGsR67XcGf6uOBfHBOlx7PAZjYVlueMlRw1wcR4Sf3yJY/n+MppS4Ndoufgr4z7hVwUEWu7dkTCH+GmbdqUpIWGoDXMgH9wa6lFkSg/3sYCBDxQg5TpK0O8/8o5CLSFtok9rJD+JjU9jueZkHy6croZuuI2Iq6VdqIRr1w+8ik0+ln/HLO6DczvtYZCN92ZnLoe9ceemWD4Fre55kdVcIudX4LWq/bzUb6s+5gl4Qcrc/rX4QstLUeD+da96dLoUX8TrwVWv8wTvbv2gdcqRXujAYSiXQYTF2ofs4Q9nY1T9u5bwYsnDAj17OYv9h6/R3Kvt7L0U9wawltcvL61l/Q0sAavKWllyOEnVK74PejuJKtMWsJMwXgXzdBbVhFxW7QNqtRo9VnvUJ1Xl6TNcU/ejE6ufJubc4ycWfZtrrZTsDB4Yth2ZiYzz3r4Tx/zfNNH75e29VJUacWPLKiYkru9Yne9dHSjLmIG6fxWsNY98WkTbE2x4WDvEYjmcuST7Mu1kOWvhZXfO/IUKj0+ekRovwCp8ehTdkoK3727NpFM/xaSSL8S5jGt1EP62doQONGyBirhSXpbUyosu4H6vxOQ5LOyjMBvVgTC9kqC4xG/k746yaFlgse2vQfxcV6w+zRUGt0Y4R+Zizxk3ome37HjQs68zvqy6edzt8vSEdXgctBaWsp+rYudyvMzboYwkufTJ9YfST8dzdXudJYZnz47zdCM9H1rAVLB0zlmhOWJ+vOPVvprnMsclX1iBxlmb8A2Z0koDGi0Js1vpc9YtwLxLLKqVhU9Sfyxsx+kDBLdOqhML6ZBJn9hpf1yJPgDneecfIc8LsVuPL46bflX1HkIi+VNb0CT1vFinXa+Z58PUpYx8I695p/xA9JKo4snPnUB09i5eG5jUoQuo5fQE6PE71ZtB5znT1Nqz1PhKg3LmixZzIG0BGG+9JJ3pMxPxU0XiqaYEbG6y+HNHo16r/RhU82xQ8vbfz3R/35w4kB0pfPXilmL+6in0T94L7O903xstwHZTE3duifyKvE4rORBG6IoG99LR9qz/0pbZpx6b4d3zmwn3XDLPVphdu0n0Hfe4+b4Ov+vFydchIVUxN3fhnaxjKxQ9+um7zzgkuuya2kGcbOzKkqQdeavJ+TjkCbrqp5uNdOO6Cg8AYBxbZ6xewvyU8/Z7JsjP1nHZLrDO5fTbw5UXM/OD7pnOfc7p9h8Llxgq7KoybY4xvh4aiL1PBx1JMWzH8d3cyBnIq+XV8Rl0rHk5VQVH+hnsiI7FT2i09nLD6Ip77PS0VfZk0rFcmShxhXM9CZV1X/BtIazwi3APiGzmo59usE+xj7KWGGDaww92n/xWt3v1KqdcFS+2USk/aFNOfJnRSuidWe8+5gPO3R7PzBlCYTOk+EWrkZLazcSac659b3hfKnXu0oQaNO9fYWCZxOj6p5sFnZ3e67cr5RrkMuawbXwzCmn2TCUI7PXbAu/hSCBKLfn0odt5dEfw6PeBPldnl25socY8zop4QjtXV+TWnHq7zH0hPpyDwqbjNWVIKVor3VXNlOWL5jouKBtArX3MxEX5t9vvtTCQ3fU/wA+KbqIqo0WDBoM3iyPnIyROGZMH7lUvgABMQa3T1LRlBIiPJcl98B9lJt3JsLlFSkPKQRUMmtUYxpkd8fpbseW8fYKWc8PJzy5mqVk+hXpbmf7wvmzPuS3cu5VOtaVHXDSjA77QHX7B0wDycnP8Vy4Lv1WBD4PuizHu9SA0yPGOk3v6Jr5bWVoeSyy5bnWnGLsoq2/JjwKFufvUxYv8JB+UjEb3gvm3gUXg0UUy11CV/5SSVo0Wy7qca2TlkVLIqOa/arXIs6BqWtS+RbzGOU3ONVTAp3Bze/fwhCGUvfS7MH/Yaihg1594vNdFEhq0wgm8x3u1G6aBaoCxNcC1VRq64eWPh+0gVVMXgpVUhcUnhSXFxpVSccVUILnI86g1+hJGvFWHFhkZIKcXTnU5pcFbtcHGQzkxoUE7NAPE/re4F0Hl6T/MDE3w+7VRzL1XQkGXP763bPdkXXFYG8vPib++n1ab72Q1IzRWyR6NS6JAUGnV/4m6JAFo6DDhX6h8D+XwYtje2fnMYTntPulb7x5agHSjJ7tBGP7o0VVuWE7+xGO9dV6M+uBh0bdcbdKa/FdNI/vvQ1ficwgKeJ4llkpc19Yg2H831pfC9p3SvjZTz1B+hlJmh7U10VY/cCgAn6KHWRi1LbixncCHqXUgzRN/Qc/6Tg7/XUfKMpaC8/8seMAd/4Zu1LobvRNYUZwhb7QdBRgYCgjZyQXmfTM1/NEW4Dlk5u4Zb+ZDROTbkKW9Vv8mN8R6vQsCZyVVvi3ZhOpAUPm1AAi3YErKN/EWvX7WpzP5KBqFFvtLSomu0TNSAnXDUxdWPdZkj1izuyboj6sYyL1HfX3wiqf+ssiJnGMNY313mi97wiuDOhnfk6/lCjkN1AGd4dWf3ibccgNbyrCB2cnk/pq5KubgCYbD6cbrdrrDh7HrLfDk07TqzMFC9ch0jRb3ywX1V7QlYjuRrQya70ObgQ/NpJni3QvwGV+NuMcj776T4KiIeDl73iOd9YM1EUOuAm0crn1Xb7WzaN/8Ej2gD2nGC28iKdo0q9WdyW7cFcAu7FNhWNYKGedt7+EZPZsBN+yPTedGtYqSrp8NKryX4lTP7hqqd/MaUFYA5o0DuHlbupGiImr+4rz+04urxvUcvz7f4E9H/iAF047PtsOtkp6sUkDXDXsre8t3wUL4K8JwPVQ4AuS+iFLuputXQWvlGLUjUcY/jeHwv0S8MdtkFeiEvZrzQNzNSLYl0dBf7FPaXy87JIHt+s5LjCXDf6dr3G6zLlDOJjIWuDs1pg720nvjzry1//s/xrdbTzJl65P59UNaEWi73rthFvkAJrEyh09aOd0q/tdy0qc9moz+lA+9mI46Z27cZdmXMbwORzuG0VeYSqsJTxzDzt4uWc/hHOXrav6Pfwk3KH8g3aYIfuroe1gdv00+1I0b892zeKHz7OWrs6ENtkFqo+Jk9WR0YTh+3XmDY4eghqAtzQHcjfwaW+OCMJvNuQ803Fgrh7x8X8gfMdJ6Z9QTl56cMwu08t7zng7IymVvDdOyJRuzsob4sk9MjAnohQ+DoMQWecojWf9uEHhTR+R1y6+MN1BoFVd8EtXOKm7d311QCfkzSIAn6slf4ZYNnpz7z1ouFOogv/2etgxopkd8cDS7qn06YUZbqk9dCUzoeN0WH9CoP6pjJG/IbDoNIzP6gEp3Fs7YsFlo7HfomtbfvBIdsmY6eH0WJVTQ02McMQNxmMyW2GweU52zJdcxSebXaN078vH0RQZERGag2GRAgwaAChPFUmvy4UOWKBsOjrPvBB/X70n+qUKIsw31M2nbotnq+KOt9gX3JsjnG9iyYJMXjAo+jgdJVUgQmL3PvV0EqhnpXqIac2OToilc9XfVwlG5//1CJ7e1QA5C3h1ZeDAMkGbHGxsmDdDFTzd3Q5SIyrtvk7otiEpYqsBqS6GozdM+HHFBr56a0rxJq2nQdVLV/iEUlaSNLu9IXtAl9K48psuytn9qIBAmXf8ssUfBxEihijjsdqBD1O+ved400b2JGvdHXoWG5vkaslL6XzyLQU+9Ez6e8LrCdH3zX1D708VeOZPd0f7y/WUsd8d15MeAxecTZ4XqIFdDrnQcAybDhC3pVooD1jyygrL1ShXBqhCrHfDZED/gpFkFN72vocfYHaGGqqs0ZSCYDCmYoM8HVta+YRe/78zXeZGd55ggKTQQzyvzJrbygFiAzD3PoBnNd6xSBPIDeswP9/FaZwYFvg8xivBVYU/ntY0TAAUSzxBQHjX3lGPgR2QNHioR3363tI77wgg1skq+bcw1vpqsUhkGBaxExNqfrrR6E7ku3IwF3+FtCSU79GKSi4kpGernPqm/MzISXFErf58IrmY/ZYVGb4LP/qTKRp4YFNhHeqlkHYwe2WZtt8UlKS98qfE5TW1tb5mBgUD/xqD35CvZTvX2SUlDQJDdfG5W1853IpdbzmLEr0tYbdsvQXSklAbbpV9RqoW+k0ms7yjhqbyTfVFS0HkypmkZMSEGyFslmR/zTE0AZrD1UuukT032DkBciDZX3TLYV/gxHYjWUMhaH9NDKv8k9lfbC8gYVtlf/fUCHBpIBHoPtAuSnsX3Ub3yCwIVLT+LD/K1WhqIFJg4DRBNS9fxvbVLDc0t+f/jb4gbTdMBQ194mV3/9V10LTDdNNSzfx8d/UILBA0kEEzuXaof6rrpsWLLNxOtP966uYMSIHJELFs7ntW3VsYNkpMpJb9V9foeTBs2V8Ra8n8a86J9kAxAB5hf+F6nSpu6+aB7DBVqR8+uG4FEPqUz8l1ERd1eZL0fshpq2CmVkSEmkfXG9A131w43exYBg6t/fbueF6ZhsBVXw1zKSDaH21tpdyIVzw1gnUVkViajjGUFM4Ptn2aCiBF32e7NwEel/qfwHdOJSO5cXV98S1L8c4f08VKxo3fK0EdyoJ8YwEwqIBRQFyDkq85fRZdnkQWPcD61gZFqUM2dIXPKIQMCgrTfb1u3t78GJqizDgGQjoMRa7OVnFnKujrVDKwDPn6knpeWkygR2MBUKG4seI30+uU+xt+JPwOW10FGJhY9c6rLZdDrrAgM0bBraCDWVpaWn+508F/F1XQEpDwztbcfHAB8gDHS5GxtIcfqFMFs2yHFgnExuwdcUZ6ep+VpSVdf5eeDL+LTPTdnx42MnRBD/lW60FIVM4CGQefrbj9bDTtP4H2OPCo808qqQzZCiooMCrR55xlfK+uqXX5C7D0ZwlOfhgMqGi/YBLvYoZJKBUdaiO9sWAdWGpLGx2Vt1WVfVr2dnog5am/pVVSW3/Hj+EIscfs0Ti+IBaLx8cfHyJE9QSfoKgoF/VZhh1ci/oPnXxeS76Yxn1Z8KvJqsLmwYxK9eHiDMVHABpWb9kI8kU5xt+rCPFYuCG7v7DrhuNP3UgSD48XrydYW+MD7X/K9iUEl/niaY9OxYNJtCMUO8jN5PT5f4hSUY9y3tpGohfaNu3rxYupgg2wPeStzlQsR8p3G9dncUA3tSlrOWyYKNQ7HbHql07GLh/8xSGPBWwmXzQDyv6Z9v7GmDR0Cfs6s11GMAt5mGkfmf8Z7veuGBxIcek882WkcBaMPRbWL/ZpiFY3G+00mZ/kz+AO4LkTJjx1pQWWPQ+JpyP/t+/aQ67labxjNzc6hqE14f+UiWkpUW4Hc7/KRSogAcjiDJkObKEu3vTe27umVD+1VxaVtZxXRWVDf8A7xQFObkyK7RJuujkMitIfW1dN8Ky7dSUFF21T5/0rUt96kG7u7rcXFy/9RZDC83tdMJGKfn5dSw8y7UXGGKr77lpmKvbUJIPkVrjEg8FkoP1tcbcmsoo37+f01VX/8HUcfGcjKv1RPiZ5dwLoTt3CX/sTYfbV9zI3dBheUOhXDBuPJZO1r0Nn/8rJNJ74sn/Q5mzYjAEAg7q7d6GKxcMebeIOQOf/pnFIyMY4nLvWTLfAJMePHz+QVa1bwAJQOKAPCsW+d+56Q3awPzfwvtTXcI8lrg+jFR9r0DOZoz0F67ODwfBE9KWkc6dICJ4oFjVoRHG12VJDPyyINRv6A1SaanvAeswmp3evN90Q+8qffHAq+o7JUg1WVlElpQ2wPuk5GTxZ0Gf64o6TU11dbF5WgYZA6X8umcsKTdfo6Tvdag7OzvLVooahIWFhQkWJVHsZmNc6gOpT7aNuqRbxLFSISEhw7JK4R1HNPqifrTXOY+4iOPbDV4gD9n5+iZxiyeL9LLth3wUDfuEQ74l0bIteeLWqpRDVQ2L4WNJ59plCcWyAzY/u+ADLiIJ376T1fbK0NBQHuxiwM7oR3dva+uSmT9rBIyHFNtjuDlRlOzs6Ky6DSqYutGlAl0sH19NcSxWT1wckoqaOvlvsP75SKyrdlqGUd8Fsj1XOgYJScWPZy5rf3+RDF7HqoFU5nOy1TbjvjDyDRZJx5lC+HtXu/B0o+PeGWoqY6T4vjYJSQIs+/5q4PNBQ9zT/ZXE589lMzZrBKVgKz4sAVvx+c6gkYNQX19fZ8/xu7AYdlFBN9o4JXxX5x+PaLy28wpL8e+Tn5voytgKavt53Y6X5yw/Ha5TS8dZAyewsbA+Zw/NIK40O0rvTWQn+3HOVxvFd/k9eOuO66Yqy8l9IjLCq2hDmTvfgRNgZVVMn6k1G2qIlUhlM/h8TK9Md7RULy1PtdJgsy2bJpC62mCTGKwsqYpm6PfkVjz2cTDKyITKhG31lofMSC+9f9LvgFPofgq7xdL4ONfNHz3ri2asemtVgqjr9clq9Fqnj1by3jH+w81p+JKdlL7+Dhml49DL3V6u8/kAPrN8g0Z1+I9w8ITI++bA7B87wm+rdkdDDVpyvmb8ZiJaV4ADCWu81jOjl8rWALN8to+Xw8vohP8/mL6/ZnZ21kr+DTdUmtNikOU3jwUSW7QjHsK4GTHYvAegDiZa1PSdw4ohYMtgg37iG2/2hu4vjeoLf6t+tf5Zym5hCLcNJyxYouiAHpwowdOO+lr11KsfELCuFOAigX0YzcA3r7IFe1UXnpXXm1eRInBATonT1yC9yUbIYC9r5BsreMNHMbCXoQkEauS9eZ0uWK4e0CH753Vi8GCvc51CrIB688IFsOzLOnHxBtUCGGDoDlYyxn2Tu/DAgGDrnCL9DxCmAsGAQOoKVf7Gdc8awbJRdL3QG0CggQFiPw9bG/ZNrgZD/a1u1s4bgQLrBiMGn93N9Lak+bKCAWRhiyX+Deqhwar+f1VXc3Bgk5OXP2ThFRUlB3vr8OhomYuIMh0ePj41PT08JSVlTkipm075XFnDc51xskhTWCmbhMVjVBQS2PcZXZgYGAIeSUo/B8ewA+FpaWmbTs7ujaIH8iuem0bLLs41fHXiRnn+I7mVo5MlDV8PbRNHYO781DiHJ19+FUhjgp2MqgNPzeJlmamyLWYQY52GEwS/ltwzd6cfqbEixDSmQqjXJP0KJr8WVygOq8tQm1qagI1+YcKkk7dXyNMMNbHqN9WzagZSZhGOUXFxVTeJrqBhzGMtZV9f387LNsG+UrViDGzs4LMzJw+J2VZ9S0s69XnAepef3PbYIFgVN9+7w5fFWrPRLPOjjg4hi/Md9lw8Vt08xUweN/95oXTP/SlKi6vmueh16MyQSyDTAi4YYhJ8vqlfBz7uOPJkYuLj73UEksUJO2/BmJqZYeLh7ZrQ7EWXG6KPfNJI/aW9dUkCsZcrHlf4Rr8bwPQ74MSQJ+5tBvTB1uOC/39EccA1n7FjEwctLazL3sRczuhuZQsmJqbnUe0Qh+u+imOoOckby0/EO9OFosj1Y+F9hEcj6V5xBMZqVuq+R7heMAJIfvgAlfHe1n3tywVmNbfVRD9YaADCBW92ePEFQb253AJshzvCxkZP8GKjLf6a92ElnOOjhZvP3/1N12324hGrGuD1su1yPcgivzvwf96ZhOlIgcjid4EjFXA9q+C5Ylnt9PHIEvKqy2nlS397KPD+9aFZLRWJIsfVNlxZWVkGNjaHYdOs1P+pBR/cMlpA9M6cYKYQwHipzmLZ93i5sS9fMvrAIIxbJY8RWZtJzonmBiemx7K/yWEpiXHRHUklDxKqbRk6ISNDt/E81Z7P99YJOqNXh3HqL4zSYxMSYpFKgy1xeAuPaovPrp4ieiAreMH+jTU1RwbhVQR8wiAltX+Zv7Mthx3J9A/6tc+3ZPr+uwdQJAIB1T1rI9SD3+dTLSs1fIFU/AeMapmzIAaJFmaOmiGCvxlc6A+GXgSl5WtDebbkQ7nFa+EOrK5Xc+7NXvrgOVRhyMts7AeNJDpfH+mEF02kcw9tdtbzplx8QiHm6V+iofgLpeB1pfb49Kw6UdzO7z4WFO70tO7fvTSNFALn4y+7hY0eJ20JQvC9mDY4E8rIjN0yjVXOc0VHizB6Af70Q5BGJ/tQPObwUXPeLSJudnyh7Ty9bhDH+gMUnZeOgbi4eLYiGzt7eMeEKC9v8/wd6OqlWCm7zM3bAHvNa4/0aLG2zyQFvErleGiIKymhcXFxGfNklu8mIwaQCLiRh/MapXs6dOAWx5P3N4ygb8zBS1Gf519ZmzRsFef6C/Cu1osS8icb6n3y/xBRes8Q1X7OW2EUDuJCpos/Dj2wnjjcdyCcKMgW1bTwvSNkYmQ81xWdrfCTlIRlwYuuY5YcHR1tkFfl5ua2xTtuvLlot1yHt19uEB+KI80eZF+sx9R0fgT6lctkDTNjXaZxPM4/fG0Aob4eZClWznVvy/LzLeb25Pvk5qsjnGwxPFladntxySgq/UpPCI3WmHemvktcVC7A35Gudfm1a4YtBOcil0wPaVrNWbwLuy0CB3M2xB7EUfVA6XiyOc3fbga7+ERESB5Vv1KZS9m7bnu/axcs5gJenm4EYxIdQaM1asVK++tAwNw7eZ2dZKivrALGoOGvGts70f4m0r+eipZqBg1R+s6PI2JixMOj20yHbsbOnqqohNxfYOyfe7Hhz6aztBK01OPGHNICRA31tuOCcQe0u0DcWR8Q4OIJgtlqMZe/PLF14lgmitPQ+iT4XcQsYpuUa/Fz3RxJmLsK6fKiz+wmv3Q+XHDzqZuLex7RetpNsybA09ehW0bF+sPWFENsjhdMVcpLO2t//4wZjk/M6+Bq3d73kb1DldSoKm16Zsa51EmNk0TUX78d2gJMXzRChsAsotHnU4PDqq6sL6Hr5e44mp8dg0DU7nimIKtjPzR6xGkV5CUvWbeisN6m07JHxCkvlpGCDxxgf0ZDPiyL1frlyZlnNdG5TFSizqFsTEZJyX0fWlnoM5uetjZCutj1wVz5Lk11Z6roICj+x+KidvVAAzoN77NBu7uktvZV27h/W3G1xk11Re5/PEVNRcUPyA/JH/9Iz/uzEfWLt8TpFD33cftp10MHjPlcmYZSdO1cTxcZnLxFtvnS9bZFQJTL0F7jeSF8/3iYlInYi9JinT07G4kna5Gwq2CXlwlVVeY55oefJgg8Ucl/Hk633IMxvQotInBOL9ELCyQf1jhgCDE/o2FGqtRq1+pSfoFUzUGcv9A/ppPatvILWY8X2VPLCeZc9d4oWxDpEYyPk7gbUU/TM0Aq5yQw2TVOwsPfZKxdAYz98hXivsCiOcWVjbkiRjl6IM9iYviQxtDc9N2S2DZzcvWcaXvMSHTifmS1bV5/InnSZtyC6MLQzVfPgOSegktp5ZB2rIpfHeQbXm2EPjDeyYT9zT0VpR9CW1VVU7W4XuIxQ7zzDPy4juTcK3HfLntEkyGey1XTe9IgvMsH+KmGyZE7zr08NkbFHBWIn9xPTk56MCiuyU5lf//xTpnk5WfZI727X6XROSSWYxbth9DtC88Vdgbo6enoGhsjtjoeQGxvcGZDLKC8nGyQsFaVDWPJa2lTsUaXaarwycy0QD8M3nEfVc2ar2vdGUdSUpLveowoLXggrOtiOBhZozHPr/mjNkYY9pZT6CNsQr1j9E/ys0duR1OLQR19Xu2jLVIuXMkBK0y6awanPLFGMQxTU9OGTqMv6GlPwEVDalXDTyE38+Sz3/ZjxtfNkH/KWBPrdPIZvZKoEkkxHagz8JJl8nh77nlQWkyBTjlLAqIFvPd4GjNoxjdwVR3XMfB+WZlvxiQnN7qr9vPRd+7LZrDxoxS5F8FlzVyqginp/glw7E0IGTVSsBUxLBsw5AE0NYmnOMY4S5zlecaIG/4aYvi5Gh4PgKsCbuPq/6r8kl5yomU2HyHCNiebWepEr0Q+NpRc3xV4jHRcguZ3vlshPVVslcUTfWuCksXqIUy39yOADiNUozpmqFPLdcr/ysB1oDLtnRkenAGBEvnrInyFxCBo4DkYQ5jQDJ9Cjza6nYGMzzbLM1EndaH10mAzQ8JEP0np+XS1QDr4HqQY/3yz9ZA8GUmalxxc+DuIZ4sgrXTw5G9IXNM4x9lhbbDpbs7+EuPo95JNXoOgu2g0yj2FOJ6umystfgXgwU1woMYKyeDcdRGDLT2oEXXb3xHweyXCx6ptfX55dOu2iJ7j84FpkmJGeIzEUkSUTXwaw8paVZFVD55aABQqtkoTvFjOV3vR324LBotAALCp/GIA36gh90en/3YsL8nEiahjmDc4Tmz40LCoweSROgpwxh/Mbj6vJuw1TtLCBn1Xp0G+7nuO/YOAihpor6H7pW3Kn+vLDd/S7YqsbhSS09ZvLLjE30vscQQRLkhH4QXw15SHw84Pqr6KkM0wCpiq+2ud+LzPDn+cKZHeGUC2j6d4YNIplAj1rYkcQG9V7AxE45uNpbBeF3A6/OajoaHhi7kcyOwxEIwwratxlwgzcMZ73MGnYGkQpMoHuW6Iyq/VZ6L9fOApP+gUz+JgCBd+KEd3+cHyoK/PqcIfQqsfJeMMs7VU6phLDJ2ET9LTHTKRxve0VOa8TAs5d3ob9cbS1cEm6DSjyLvZtOLgUtLRMXjXhy7qSrBQb6U83An9TBAWDyDAj74sdABaoIhoN5EJApaizpdpA+CEbk6VMYlw4/EH0XrmHdd6FeMnBquyBTl2rr9qZyZqaRgUlIZeC1zNWHGeJlp43g0UO+3NEgYclwTVsp60cxEN8YjEnicaDlwn0TyelnBSHUM4tNtGNVXwORIv3qHuAjsJ3rtLqSf34b+RJUckBnkIc1p6paySp/Ly3gsacS+dyJVgex2Jq/HArKFviu+pqPrdP+S12Mt5Cla8bML//ewF5xHXCbXKThhAu+4wmCU5djaF4YVqj/Hc4fn++Z07jEguL/Jqv8P8gohQwF38b12NkB6CoOyR07yGfRSyp23tx7sR/pbzwkICJqwAy57kDhfgajw/umE4LXJR/mTexwo54ZC7XEC4INbA0oup55lDbTPop8Pjn54VFTp2TtBbUKXCkndEI6bXuWMI4WVwAE/ua0ysceVx//fP2d8iJnG6S68bzBKm7mdJGSvO9TurZekbkphjwiU+NwHlLDECGofBGf3fa80ZAzhEkV651QOXXncJK7sRwzLODuIYZJ732WX7+OjZIS/+C4TBbMRQBYzav9kpCCbXbJsp0ZH6Eh39YkfYs8VDW24CdfQ1CajYaYV49T1OSZJWOZG3LKE5zrnEiJFfPv3buAs8YM3u/QDaCeBCTNCCWgz7LhG+dVBe3dR5dWgU3CKurF4+OvaYpdUfsMvK5HElmMt2jRuRCRUazltuX8fTpfPxeMEqM6j8RmGWMAT7bQfKEEPFG6G7y795zZs5x+9ahxUzoZkqW9in6cm8620jSaU7HEpkCdrYdjKBUqLG++rzit5z4f8UKUMZrkBjYGF9WbP361wtSqRXjjq8lOgkTTl/i2gVlbptEQN8zUt99U9Hk5kHFB31u1T9cVs6BI3HM+GX6EaLhysN2vEEPU4+nq512hfc/d/km/eV9lNoZ/XzOK4Jvf5e/OXepAczYkklKlfT77NXzk+2ylrTdEui/JYI+gxlKA9Bow920PNeZPYWC4ZcTlZyOM8gNXUICKgvr3cB4JBNwKe3wA6HvnUOFeNsmTIALmlEh/F/Iux6cIQt6gERXfwWz0t024ZDXFxcnN1ux1MzMyNf5LysEYFQQ7a2ts5+ck/mgenYGtIY31tEQce6kQwj5nC2A+1XpJAO5VUJ/KXtl6mn5baTDH1hDsyYhK3RZnkNY+iUc2Jo5VMoqKh65xKy4v8nItGF6fiCsHF7vh0RGdlTY9L3e/l2rG3kAWH29+3Y+Rf01zwBp+VPArvfLapxW29dzgsIhzKUhQBVu9a3dhWfe2nO7G5RCoOj7xzYzem3sUSyhgp7vTt2R9FTiBBEeq9Zob15VkPusTWC2iE0//pt4rEO5DXvUzq3VZw7fjckNFwLdsYYuPYPFJycmIfzlQWwr9UMS/+ZFzS3stJ7JcvZp2ebfZrOMV3zf6qUMJRRG6LoYT6/SMWTHdcsbg1EVxr53xX2tYR1yMrJfR4+Qs3RqlAn7lCQwFf5LzEmhG4rJycnkXl88fiWGEvdc9yP4WpikdhJyMqq4OrwlWqcvn3ldBdlYyqtVQwMDDVFhYWMx9Cvm352Gq87ao0lJTIanjtPzMxV99yz2ls+OzRijjSfkPFYphxo7aJp0ZnfedZEWYC+fHlvNa04Q/Qbhe1/nB48US4QQpXm3t4CKWwGBxh0IiJrm37ZjkrHMM9uxBx6+n5n+hxnHbc3z54bU98F8C9H+h2LIclXC6pK2E+hRHIQ7L58CQJ1TQ40IHYvmbrtsHdi3oIo3c4J/0WXSQiPokarN2xqmvwEEOA/WUmlj1UfjIP+L1BLBwgE+TWUl34AANCCAABQSwECFAAUAAgICABBgAlN/6AQMOwEAABHJgAAFwAAAAAAAAAAAAAAAAAAAAAAZ2VvZ2VicmFfZGVmYXVsdHMyZC54bWxQSwECFAAUAAgICABBgAlNKbwdg+ECAADCDAAAFwAAAAAAAAAAAAAAAAAxBQAAZ2VvZ2VicmFfZGVmYXVsdHMzZC54bWxQSwECFAAUAAgICABBgAlN1je9uRkAAAAXAAAAFgAAAAAAAAAAAAAAAABXCAAAZ2VvZ2VicmFfamF2YXNjcmlwdC5qc1BLAQIUABQACAgIAEGACU28z6WtPAsAAD48AAAMAAAAAAAAAAAAAAAAALQIAABnZW9nZWJyYS54bWxQSwECFAAUAAgICABBgAlNBPk1lJd+AADQggAAFgAAAAAAAAAAAAAAAAAqFAAAZ2VvZ2VicmFfdGh1bWJuYWlsLnBuZ1BLBQYAAAAABQAFAEwBAAAFkwAAAAA="
};
$result;
}

ENDDOCUMENT();        


