## DBsubject(Calculus - multivariable)
## DBchapter(Calculus of vector valued functions)
## DBsection(Parameterized curves)
## Institution(W.H.Freeman)
## Author(JustAsk)
## MLT(pc_circles_in_intersections)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Rogawski')
## EditionText1('2')
## Section1('13.1')
## Problem1('37')
## KEYWORDS('calculus')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "Parser.pl",
  "freemanMacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$c1=random(5,10);
$c=$c1**2;
$b=random(2,7);
$a=$b**2+$c1**2;

$context = Context();
$context->variables->add(t=>'Real');

$ans1=Formula("$c1*cos(t)");
$ans2=Formula("$b");
$ans3=Formula("$c1*sin(t)");

$multipart = MultiAnswer($ans1,$ans2,$ans3)->with(
    singleResult => 1,
    checker => sub {
      my ($correct, $student, $self) = @_;
      my ($cx, $cy, $cz) = @{$correct};
      my ($sx, $sy, $sz) = @{$student};
      return (($cx == $sx && $cy == $sy && $cz == $sz) ||
              ($cx == $sz && $cy == $sy && $cz == $sx));
    }
  );

Context()->texStrings;

BEGIN_TEXT
\{ textbook_ref_exact("Rogawski ET 2e", "13.1","37") \}
$PAR
Find a parametrization, using \(\cos(t)\) and \(\sin(t)\), of the following curve:$BR
The intersection of the plane \(y=$b\) with the sphere \(x^2+y^2+z^2=$a\)
$PAR
\(\bf r\it (t)=\langle\)
  \{$multipart->ans_rule()\},
  \{$multipart->ans_rule()\},
  \{$multipart->ans_rule()\}
\(\rangle\)
$PAR
END_TEXT
Context()->normalStrings;
ANS($multipart->cmp());

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$SOL Substituting \(y=$b\) in the equation of the sphere gives:
$PAR
\[x^2+\left($b\right)^2+z^2=$a \Rightarrow x^2+z^2=$c\]
This circle in the horizontal plane \(y=$b\) has the parametrization $BR
\( x = \sqrt{$c}\cos t, z=\sqrt{$c}\sin t\).Therefore, the points on the intersection of the plane $BR
\(y=$b\) and the sphere \(x^2+y^2+z^2=$a\), can be written in the form 
\(\left($c1\cos t,$b,$c1\sin t\right)\), yielding the following parametrization:
$PAR
\[\bf r\it (t)=\left<$c1\cos t,$b,$c1\sin t\right>\]

END_SOLUTION

ENDDOCUMENT();
