# DESCRIPTION
# Simplify exponents in an expression.
# ENDDESCRIPTION

## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Institution(University of Missouri-Columbia)
## Level(3)
## KEYWORDS('exponents')
## MLT(Quotient_in_power)


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real",y=>"Real",z=>"Real", "X"=>"Real", Y=>"Real", Z=>"Real");
LimitedPowers::OnlyPositiveIntegers();
Context()->operators->undefine("+","-","/");

$a = random(2,12)*random(-1,1,2);
$b = random(-4,-2,1);
$c = $b*$a;
$exp = random(-3,-2,1);
$m1 = random(-5,-2,1);
$n1 = random(6,10,1);
$k1 = random(-3,-7,1);
$m2 = random(2,5,1);
$n2 = random(-4,-6,1);
$k2 = random(-12,-9,1);

$e1 = ($m1 - $m2)*$exp;
$e2 = ($n2 - $n1)*$exp;
$e3 = ($k2 - $k1)*$exp;
$num = Formula("($b)^{-$exp} x^{$e1}")->reduce;
$den = Formula("y^{$e2}z^{$e3}")->reduce;
$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1x^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01x^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);
$const = ($b)**(-$exp);
$texanswer = "\(\displaystyle \frac{$const x^{$e1}}{y^{$e2} z^{$e3}}\)";

$frac = "\left(\frac{$a x^{$m1}y^{$n1}z^{$k1}}{$c x^{$m2} y^{$n2} z^{$k2}} \right)^{$exp}";
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write the following using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


COMMENT('Desc: Simplify exponents in an expression.');
ENDDOCUMENT();
