# DESCRIPTION
# Simplify exponents in an expression.
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(Exponents_Division_Mix2015)
## Level(2)
## KEYWORDS('exponents')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real",y=>"Real",z=>"Real", "X"=>"Real", Y=>"Real", Z=>"Real");
LimitedPowers::OnlyPositiveIntegers();
Context()->operators->undefine("/");

$a = random(2,9);
do {$b = non_zero_random(-9,9)} until (abs($b/$a) != 1);

$n = random(5,8);
$m = $n + random(2,7);
$k = $m - $n;
$num = Formula("($a x + $b)^{$k}")->reduce;
$den = Formula("1")->reduce;
$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1x^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01x^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);
$texanswer = "\(\displaystyle (".nicestring([$a,$b],['x',""]).")^{$k}\)";

$frac = "\frac{($a x + $b)^{$m}}{($a x + $b)^{$n}}";
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write the following using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


COMMENT('Desc: Simplify exponents in an expression.');
ENDDOCUMENT();
