# DESCRIPTION
# Simplify exponents in an expression.
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(Division_exponents_multivar)
## Level(3)
## KEYWORDS('exponents')

###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "unionTables.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real", y=>"Real", z=>"Real", X=>"Real", Y=>"Real", Z=>"Real", w=>"Real", W=>"Real");
LimitedPowers::OnlyPositiveIntegers();
Context()->operators->undefine("/");

do {
  $a1 = random(2,4);
  $p1 = 1;
  $q1 = random(2,9);
  do {$r1 = random(2,9)} until ($r1 != $q1);
  $m = random(2,3);

  do {$a2 = random(2,9)} until ($a2 != $a1);
  do {$p2 = random(2,9)} until ($p2 != $q1 && $p2 != $r1);
  do {$q2 = random(2,9)} until ($q2 != $p2);
  $r2 = 1;
  
  $dct = ($a1)**($m);
  $nc = $a2/gcd($a2,$dct);
  $dc = $dct/gcd($a2,$dct);
} until ($dc != 1);

$e1 = $m + $p2;
$e2 = $q1*$m - 1;
$e3 = $r1*$m + $q2;

$num = Formula("$nc w^{$e1} z^{$e2}")->reduce;
$den = Formula("$dc y^{$e3}")->reduce;
$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1x^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01x^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);

Context()->texStrings;
$num = $num;  $den = $den;
$texanswer = "\(\displaystyle \frac{$num}{$den}\)";

$frac = "\left( \frac{w z^{$q1}}{$a1 y^{$r1}} \right)^{$m} \left( \frac{$a2 w^{$p2}}{y^{$q2} z} \right)";
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write the following using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


COMMENT('Desc: Simplify exponents in an expression.');
ENDDOCUMENT();
