# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Exponential)
## Institution(PCC)
## Author(Alex Jordan)
## MO(1)



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserFunction.pl",
  "niceTables.pl",
  "PGML.pl",
  "unorderedAnswer.pl",
  "PGcourse.pl",
);


TEXT(beginproblem());

##############################################


Context("Numeric");
Context()->flags->set(reduceConstants=>0);
Context()->variables->are(x=>'Real',y=>'Real',t=>'Real',P=>'Real');
Context()->variables->set(y=>{limits=>[0.1,10]});
Context()->variables->set(P=>{limits=>[0.1,10]});

$multians = MultiAnswer(Formula("x"), Formula("ln(y)"))->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


##############################################

BEGIN_PGML

Rewrite the statements using logarithms instead of exponents.  

END_PGML
BEGIN_TEXT
$PAR

\{
LayoutTable(
[
["\(e^x\)","\(=\)","\(y\)"],
[$multians->ans_rule(6),"\(=\)",$multians->ans_rule(6)],
],
allcellcss=>"padding:0pt;", 
)\}

$PAR

\{
LayoutTable(
[
["\(10^t\)","\(=\)","\(P\)"],
[ans_rule(6),"\(=\)",ans_rule(6)],
], 
allcellcss=>"padding:0pt;", 
)\}


END_TEXT

##############################################

$showPartialCorrectAnswers = 1;

Context()->flags->set(formatStudentAnswers=>parsed);


ANS( $multians->cmp() ); 

UNORDERED_ANS( 
Formula("t")->cmp(), 
Formula("log(P)")->cmp(), 
);

##############################################

BEGIN_PGML_SOLUTION

When one side of an equation is a number raised to a power, then that base can be moved to the other side of the equation in the form of a logarithm. That is

    [`\begin{aligned}
b^{\text{left}}&=\text{right}\\
\underset{\hookrightarrow}{b}^{\text{left}}&=\text{right}\\
\text{left} &= \log_b(\text{right})
\end{aligned}`]

So we have the equations 

    [`x=\ln(y)`]

and

    [`t=\log(P)`]

END_PGML_SOLUTION
ENDDOCUMENT();