##DESCRIPTION
##  Functions: Difference Quotient
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Difference quotient)
## Date(Spring 2012)
## Institution(University of Missouri)
## MLT(DiffQuotDirect_nonpoly)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('2.2')
## Problem1('')
## KEYWORDS('algebra', 'functions', 'difference quotient')

########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#######################################
#  Set-up
Context()->variables->add( h=>"Real");

$c = random(-1,1,2)*random(2,5,1);
$a = random(-1,1,2)*random(2,5,1);

$f = "\frac{$c}{x + $a}";

$num = Formula("-$c")->reduce;
$den = Formula(" (x+$a)(x+h + $a) ")->reduce;
$texanswer = "\(\dfrac{".(-$c)."}{(x+$a)(x+h+$a)}\)";

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      }  elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

ANS($multians->cmp() );

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ 
\frac{f(x + h) - f(x)}{h} = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle \frac{f(x+h)-f(x)}{h} = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


 

Context("Numeric");
################################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Find the difference quotient for the function \(\displaystyle f(x) = $f\).  Simplify your answer as much as possible.
$BR$BR$BR
$displayfrac
END_TEXT


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ENDDOCUMENT();
