## DESCRIPTION
##  Statistics: Numerical descriptive techniques
## ENDDESCRIPTION

## KEYWORDS('statistics', 'numerical', 'descriptive statistics', 'quartile')
## JJH tagged this problem.

## DBsubject('Statistics')
## DBchapter('Descriptive Statistics')
## DBsection('Quartiles')
## Date('6/1/2005')
## Author('Jeff Holt')
## Institution('UVA')
## TitleText1('Statistics for Management and Economics')
## EditionText1('6')
## AuthorText1('Keller, Warrack')
## Section1()
## Problem1()


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
"PG.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGnumericalmacros.pl",
"PGstatisticsmacros.pl",
"PGauxiliaryFunctions.pl",
"MathObjects.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;  # Change to 0 to shut off separate answer display.
## install_problem_grader(~~&std_problem_grader); # Uncomment to turn off partial credit.

$data[3] = random(580,620,10);
$data[4] = random(630,650,10);
$data[2] = $data[3] - random(0,20,10);
$data[1] = $data[3] - random(30,50,10);
$data[0] = $data[3] - random(60,100,10);
$data[5] = $data[4] + random(0,20,10);
$data[6] = $data[4] + random(30,50,10);
$data[7] = $data[4] + random(60,100,10);

$q1 = Real(0.5*$data[1] + 0.5*$data[2]);
$q1exc = Real(0.75*$data[1] + 0.25*$data[2]);
$q1inc = Real(0.25*$data[1] + 0.75*$data[2]);
$q3 = Real(0.5*$data[5] + 0.5*$data[6]);
$q3exc = Real(0.25*$data[5] + 0.75*$data[6]);
$q3inc = Real(0.75*$data[5] + 0.25*$data[6]);
$iqr = Real($q3 - $q1);
$with = ($data[5]-$data[3]);

@slice = NchooseK(8,8);
@shuffleddata = @data[@slice];


$multians = MultiAnswer($q1, $q3, $iqr)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q1stu, $q3stu, $iqrstu ) = @{$student};
      my ( $q1cor, $q3cor, $iqrcor ) = @{$correct};
      if ( ($q1cor == $q1stu) || ($q1exc == $q1stu) || ($q1inc == $q1stu)) 
          {$q1OK = 1;} else {$q1OK = 0;};
      if ( ($q3cor == $q3stu) || ($q3exc == $q3stu) || ($q3inc == $q3stu)) 
          {$q3OK = 1;} else {$q3OK = 0;};
      if ($iqrstu == $q3stu-$q1stu)
          {$iqrOK = 1;} else {$iqrOK = 0;};
      return [$q1OK, $q3OK, $iqrOK];
  }
);

BEGIN_TEXT
$PAR A sample of eight math SAT scores is given below:
$PAR
\[
$shuffleddata[0] \quad 
$shuffleddata[1] \quad 
$shuffleddata[2] \quad 
$shuffleddata[3] \quad 
$shuffleddata[4] \quad 
$shuffleddata[5] \quad 
$shuffleddata[6] \quad 
$shuffleddata[7] 
\]
$PAR
Find the first and third quartiles and the interquartile range for this sample. Different software applications may use slightly different formulas for finding quartiles than what a person would use by hand; this problem is coded so that any of the common methods that software or people would use will give correct answers.
$PAR
\(Q_1 ={}\)  \{$multians->ans_rule(10)\}
$PAR
\(Q_3 ={}\)  \{$multians->ans_rule(10)\}
$PAR
\(IQR ={}\)  \{$multians->ans_rule(10)\}

$BR


END_TEXT

ANS( $multians->cmp() ); 

BEGIN_SOLUTION
Note that we should first organize the data from smallest to largest, as seen below,
$PAR
\($data[0], $data[1], $data[2], $data[3], $data[4], $data[5], $data[6], $data[7]\)
$PAR
(a)  If you include the median in finding the lower quartile, we find \(Q_{1} =$data[3]\).  If we do not include the median, we find \(Q_{1} =$q1\).
$PAR
(b)  If you include the median in finding the upper quartile, we find \(Q_{3} =$data[5]\).  If we do not include the median, we find \(Q_{3} =$q3\).
$PAR
The \(IQR\) is found by computing \(Q_{3}- Q_{1}\).  Thus if we have included the median in our calculations, we find \(IQR = $iqr\), if we have not included the median, we find \(IQR =$with\).

END_SOLUTION

ENDDOCUMENT();       # This should be the last executable line in the problem.
