## DESCRIPTION
##  Find basis for span in polynomials
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Span)
## Date(01/09/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(2)
## Static(0)
## MO(1)


DOCUMENT();  

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"PGcourse.pl",
"MathObjects.pl",
"parserMultiAnswer.pl"
);

sub pick {
  my ($k, @in) = @_;
  my @out = ();
  while (@out < $k) { push( @out, splice(@in, random(0,$#in), 1) ); }
  @out;
}

TEXT(beginproblem());

COMMENT('this problem should be rewritten so that the answer is a List rather than two answer blanks');

Context("Numeric");
Context()->variables->are(x=>"Real");

($a,$c) = pick(2,(-2..2));
($b,$d) = (0,0);

my $line = Formula("($d-$b)/($c-$a)(x-$a)+$b")->reduce;

$ans1 = Formula("(x-$a)(x-$c)  + $line")->reduce;
$ans2 = Formula("(x-$a)(x-$c)x + $line")->reduce;

$MAns = MultiAnswer($ans1, $ans2
)->with(
 singleResult=>0,
 allowBlankAnswers=>1,
 checkTypes=>0,
 checker=>sub{

 my ($correct,$student,$self) = @_;
 my ($c1,$c2) = @{$correct};
 my ($s1,$s2) = @{$student};

 my @result = (1,1);

 if ($c1->typeMatch($s1)) {
   my $msg = "";
   if (!($s1->D('x')->D('x')->D('x')->D('x')->isZero)) {
     $result[0] = 0; $msg = "Wrong degree!$BR";
   }
   if ($s1->eval(x=>$a) != $b) { 
     $result[0] = 0; $msg = $msg . "not zero at $a $BR"; 
   } 
   if ($s1->eval(x=>$c) != $d)  {
     $result[0] = 0; $msg = $msg . "not zero at $c $BR"; 
   }
   if ($s1==Formula("0")) { $result[0] = 0; }
   $self->setMessage(1,$msg) if ($result[0]==0);
 } else { $result[0] = 0; }
 if ($c2->typeMatch($s2)) {
   my $msg = "";
   if (!($s2->D('x')->D('x')->D('x')->D('x')->isZero)) {
     $result[1] = 0; $msg = "Wrong degree!$BR";
   }
   if ($s2->eval(x=>$a) != $b) { 
     $result[1] = 0; $msg = $msg . "not zero at $a $BR"; 
   } 
   if ($s2->eval(x=>$c) != $d)  {
     $result[1] = 0; $msg = $msg . "not zero at $c $BR"; 
   }
   if ($s2==Formula("0")) { $result[1] = 0; }
   $self->setMessage(2,$msg) if ($result[1]==0);
} else { $result[1] = 0; } 

 if ($c1->typeMatch($s1) && $c2->typeMatch($s2) && ($result[0] == 1) 
  && ($s1->eval(x=>2.7)*$s2->eval(x=>3.1) == $s1->eval(x=>3.1)*$s2->eval(x=>2.7)))
   { $result[1] = 0; $self->setMessage(2,"span is too small!"); } 


 return ~~@result;
});

Context()->texStrings;
BEGIN_TEXT

Let
\( \displaystyle 
  W = \bigl\lbrace a_0+a_1 x+a_2 x^2 + a_3 x^3 \in\mathcal{P}_3(\mathbb{R})\ 
	  \Big|\ \ p($a)=p($c) = $b \bigr\rbrace\)
$PAR
Find \(p_1{\small \!(x)}\) and \(p_2{\small \!(x)}\) whose span is \(W\).
$PAR

\(\quad W = \mathrm{Span}\Bigl\lbrace \)
\{$MAns->ans_rule(10);\} \( , \)
\{$MAns->ans_rule(10);\} \( \Bigr\rbrace \)

$PAR
On your own paper, you should show that these two polynomials are enough 
to span all of \(W\).

END_TEXT
Context()->normalStrings;

ANS( $MAns->cmp );

ENDDOCUMENT();
