## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Basis and dimension)
## Date(10/11/2013)
## Institution(NAU)
## Author(Nandor Sieben)
## Level(2)
## MO(1)
## KEYWORDS('linear algebra','elementery matrix','multiplication')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "MatrixReduce.pl",
  "rank.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context("Fraction");
Context() -> parens -> set ("[" => {formMatrix => 1});

$I = Matrix([[1,0],[0,1]]);
$A = Matrix([[1,0],[0,-1]]);
$B = Matrix([[0,1],[0,0]]);
$C = Matrix([[0,0],[1,0]]);

$multians = MultiAnswer($A, $B, $C)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);
$s2 = Matrix($s[2]);
if (0 != $s0->trace) {
  $self->setMessage(1,"Trace not zero.");
  return 0; 
} 
if (0 != $s1->trace) {
  $self->setMessage(2,"Trace not zero.");
  return 0; 
} 
if (0 != $s2->trace) {
  $self->setMessage(3,"Trace not zero.");
  return 0; 
} 
  if (rank(Matrix(
      [$s0->element(1,1),$s0->element(1,2),$s0->element(2,1)],
      [$s1->element(1,1),$s1->element(1,2),$s1->element(2,1)],
      [$s2->element(1,1),$s2->element(1,2),$s2->element(2,1)]))==3) {
    return 1;
  }
$self->setMessage(1,"Not independent.");
$self->setMessage(2,"Not independent.");
$self->setMessage(3,"Not independent.");
return 0;
}
);

Context()->texStrings;
BEGIN_TEXT
Find a basis for the vector space 
\( \lbrace A\in {\mathbb R}^{2\times 2} \mid \text{tr}(A)=0\rbrace\) 
of \(2\times 2\) matrices with trace 0.

$BR
$BR
\( B = \lbrace \)
\{ $multians->ans_array(4) \},
\{ $multians->ans_array(4) \},
\{ $multians->ans_array(4) \} \(\rbrace \).


END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

;

ENDDOCUMENT();
