## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','magic square','basis')

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Basis and dimension)
## Date('10/12/2013')
## Author('Nandor Sieben')
## Institution('NAU')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context('Matrix');

$I = Matrix([[1,0],[0,1]]);
$A = Matrix([[0,-1,1],[1,0,-1],[-1,1,0]]);
$B = Matrix([[-1,0,1],[2,0,-2],[-1,0,1]]);
$C = Matrix([[2,2,-1],[-2,1,4],[3,0,0]]);

sub magic {
my ($m) = @_;
my $s1=$m->element(1,1)+$m->element(1,2)+$m->element(1,3); 
my $s2=$m->element(2,1)+$m->element(2,2)+$m->element(2,3); 
return 0 if ($s2 != $s1);
my $s3=$m->element(3,1)+$m->element(3,2)+$m->element(3,3); 
return 0 if ($s3 != $s1);
my $s4=$m->element(1,1)+$m->element(2,1)+$m->element(3,1); 
return 0 if ($s4 != $s1);
my $s5=$m->element(1,2)+$m->element(2,2)+$m->element(3,2); 
return 0 if ($s5 != $s1);
my $s6=$m->element(1,3)+$m->element(2,3)+$m->element(3,3); 
return 0 if ($s6 != $s1);
my $s7=$m->element(1,1)+$m->element(2,2)+$m->element(3,3); 
return 0 if ($s7 != $s1);
my $s8=$m->element(1,3)+$m->element(2,2)+$m->element(3,1); 
return 0 if ($s8 != $s1);
return 1;
}

$multians = MultiAnswer($A, $B, $C)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);
$s2 = Matrix($s[2]);

return 0 if (! magic($s0)); 
return 0 if (! magic($s1)); 
return 0 if (! magic($s2)); 
if (Matrix(
      [$s0->element(1,1),$s0->element(1,2),$s0->element(1,3),
       $s0->element(2,1),$s0->element(2,2),$s0->element(2,3),
       $s0->element(3,1),$s0->element(3,2),$s0->element(3,3)],
      [$s1->element(1,1),$s1->element(1,2),$s1->element(1,3),
       $s1->element(2,1),$s1->element(2,2),$s1->element(2,3),
       $s1->element(3,1),$s1->element(3,2),$s1->element(3,3)],
      [$s2->element(1,1),$s2->element(1,2),$s2->element(1,3),
       $s2->element(2,1),$s2->element(2,2),$s2->element(2,3),
       $s2->element(3,1),$s2->element(3,2),$s2->element(3,3)],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0]
      )->order_LR==3) {
  return 1;
}
return 0;
}
);

Context()->texStrings;
BEGIN_TEXT
Find a basis \( B \) for the vector space of \( 3\times 3\) magic squares. 
A square matrix is magic if the sum of the numbers in each row, column and main diagonal is the same. 

$BR
$BR
\( B = \lbrace \)
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \} 
\(\rbrace \).


END_TEXT
Context()->normalStrings;

ANS( $multians->cmp() );

COMMENT('MathObject version. This requires solving a pretty large system of equations. CAS may be useful');

ENDDOCUMENT();
