## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','magic square','basis')

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Basis and dimension)
## Date('10/12/2013')
## Author('Nandor Sieben')
## Institution('NAU')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context('Matrix');

$A = Matrix([[1,0],[0,1]]);
$B = Matrix([[0,1],[1,0]]);

sub magic {
my ($m) = @_;
my $s1=$m->element(1,1)+$m->element(1,2); 
my $s2=$m->element(2,1)+$m->element(2,2); 
return 0 if ($s2 != $s1);
my $s4=$m->element(1,1)+$m->element(2,1); 
return 0 if ($s4 != $s1);
my $s5=$m->element(1,2)+$m->element(2,2); 
return 0 if ($s5 != $s1);
return 1;
}

$multians = MultiAnswer($A, $B)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);

return 0 if (! magic($s0)); 
return 0 if (! magic($s1)); 
if (Matrix(
      [$s0->element(1,1),$s0->element(1,2),
       $s0->element(2,1),$s0->element(2,2)],
      [$s1->element(1,1),$s1->element(1,2),
       $s1->element(2,1),$s1->element(2,2)],
      [0,0,0,0],
      [0,0,0,0],
      )->order_LR==2) {
  return 1;
}
return 0;
}
);

Context()->texStrings;
BEGIN_TEXT
A square matrix is half-magic if the sum of the numbers in each row and column is the same. 
Find a basis \( B \) for the vector space of \( 2\times 2\) half-magic squares. 

$BR
$BR
\( B = \lbrace \)
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \} 
\(\rbrace \).


END_TEXT
Context()->normalStrings;

ANS( $multians->cmp() );

COMMENT('MathObject version.');

ENDDOCUMENT();
