## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','elementery matrix','multiplication')

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Basis and dimension)
## Date('10/11/2013')
## Author('Nandor Sieben')
## Institution('NAU')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context('Matrix');

$I = Matrix([[1,0],[0,1]]);
$A = Matrix([[1,0],[0,-1]]);
$B = Matrix([[0,1],[0,0]]);
$C = Matrix([[0,0],[1,0]]);

$multians = MultiAnswer($A, $B, $C)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);
$s2 = Matrix($s[2]);
return 0 if (0 != $s0->trace); 
return 0 if (0 != $s1->trace); 
return 0 if (0 != $s2->trace); 
  if (Matrix(
      [$s0->element(1,1),$s0->element(1,2),$s0->element(2,1)],
      [$s1->element(1,1),$s1->element(1,2),$s1->element(2,1)],
      [$s2->element(1,1),$s2->element(1,2),$s2->element(2,1)])->order_LR==3) {
    return 1;
  }
return 0;
}
);

Context()->texStrings;
BEGIN_TEXT
Find a basis for the vector space 
\( \lbrace A\in {\mathbb R}^{2\times 2} \mid \text{tr}(A)=0\rbrace\) 
of \(2\times 2\) matrices with trace 0.

$BR
$BR
\( B = \lbrace \)
\{ $multians->ans_array(4) \},
\{ $multians->ans_array(4) \},
\{ $multians->ans_array(4) \} \(\rbrace \).


END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

COMMENT('MathObject version.');

ENDDOCUMENT();
