## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','magic square')

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Coordinate vectors and change of basis)
## Date('10/23/2013')
## Author('Nandor Sieben')
## Institution('NAU')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context('Matrix');

$a=non_zero_random(-9,9,1);
$b=non_zero_random(-9,9,1);
$B1 = Matrix([[1,0],[0,1]]);
$B2 = Matrix([[0,1],[1,0]]);
$M = Matrix([[$a,$b],[$b,$a]]);
$MB = Matrix([[$a],[$b]]);

sub magic {
my ($m) = @_;
my $s1=$m->element(1,1)+$m->element(1,2); 
my $s2=$m->element(2,1)+$m->element(2,2); 
return 0 if ($s2 != $s1);
my $s4=$m->element(1,1)+$m->element(2,1); 
return 0 if ($s4 != $s1);
my $s5=$m->element(1,2)+$m->element(2,2); 
return 0 if ($s5 != $s1);
return 1;
}

$multians = MultiAnswer($B1, $B2, $MB)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
my $s0 = Matrix($s[0]); 
my $s1 = Matrix($s[1]);
my $coo  = Matrix($s[2]);

return 0 if (! magic($s0)); 
return 0 if (! magic($s1)); 
if (Matrix(
      [$s0->element(1,1),$s0->element(1,2),
       $s0->element(2,1),$s0->element(2,2)],
      [$s1->element(1,1),$s1->element(1,2),
       $s1->element(2,1),$s1->element(2,2)],
      [0,0,0,0],
      [0,0,0,0],
      )->order_LR != 2) {
  return 0;
}

if ($M != ($coo->element(1,1)*$s0 + $coo->element(2,1)*$s1) ) {
  return 0;
}

return 1;
}
);


Context()->texStrings;
BEGIN_TEXT
A square matrix \( A \) is called \( \emph{half-magic} \) 
if the sum of the numbers in each row and column is the same. 
The common sum in each row and column is denoted by \( s(A) \) 
and is called the \( \emph{magic sum} \) of the matrix \(A\).
Let \( V \) be the vector space of \( 2\times 2\) half-magic squares. 

$BR 

(a) Find an ordered basis \( B \) for \( V \). 

$BR

\( B = ( \)
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \} 
\( ) \).

$BR

(b) Find the coordinate vector \( [M]_B \) of \( M=$M \) in your chosen ordered basis \( B \).

$BR

\( [M]_B = \)
\{ $multians->ans_array(3) \}.


END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

COMMENT('MathObject version.');

ENDDOCUMENT();
