##DESCRIPTION
## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Institution(NAU)
## Author(Nandor Sieben)
## MLT(diag3x3)
## Level(3)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('6.4')
## Problem1('')

##ENDDESCRIPTION

DOCUMENT();  

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

Context('Matrix');

$showPartialCorrectAnswers = 0;

do {
$e1 = non_zero_random(-2,2);
$e2 = non_zero_random(-3,3);
} while ($e1 == $e2);

$D = Matrix([[$e1,0],[0,$e2]]);


do {
  $a = random(-9,9);
  $b = random(-3,3);
  $c = random(-3,3);
  $d = random(-3,3);
  $P = Matrix([[$a,$b],[$c,$d]]);
} while (abs($P->det) != 1 );

$PI = $P->inverse();
$A= $P*$D*$PI;


$multians = MultiAnswer($P, $D)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$sP = Matrix($s[0]);
if ($sP->det == 0) {
  $self->setMessage(1,"Not invertible.");
  return 0;
}

$sPI = $sP->inverse;
$sD = Matrix($s[1]);
if ($sD->element(1,2)!=0 or 
    $sD->element(2,1)!=0 
) {
  $self->setMessage(2,"Not diagonal.");
  return 0;
}
return 0 if ($sD != $sPI*$A*$sP);
return 1;
}
);

Context()->texStrings;
BEGIN_TEXT

Let \(A = $A \). Find an invertible matrix \( P \) and a diagonal matrix \( D \) such that \( D = P^{-1}AP \).
$BR
\(P=\) \{$multians->ans_array(5)\}, \(D=\) \{$multians->ans_array(5)\}
END_TEXT

ANS($multians->cmp());

COMMENT('Eigenvalues are small integers, lambda1 != lambda2 ');
ENDDOCUMENT();      

