## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## help from Paul Pearson
## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Elementary matrices)
## Date(9/20/2013)
## Institution(NAU)
## Author(Nandor Sieben)
## Level(2)
## MO(1)
## KEYWORDS('linear algebra','elementery matrix','multiplication')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context('Matrix');

$a = random(-9,9,1);
$b = random(-9,9,1);
$c = random(-9,9,1);
$d = random(-9,9,1);

$X = Matrix([[$a,$b],[$c,$d]]);
$B = Matrix([[2*$a,2*$b],[$c,$d]]);
$A = Matrix([[.5,0],[0,1]]);
$I = Matrix([[1,0],[0,1]]);

$multians = MultiAnswer($A, $B)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);
return 0 if ($s0 == $I); 
return 0 if ($s1 == $I);
return ($s0 * $s1) == $X;
}
);

Context()->texStrings;
BEGIN_TEXT
Let \( X = $X \).

Give an example of two \( 2 \times 2 \) matrices \( A \) and \( B \),
neither of which is the identy matrix \( I \),
such that \( AB = X \).
$BR
$BR
\( A = \)
\{ $multians->ans_array(5) \},

\( B = \)
\{ $multians->ans_array(2,2,5) \}
END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

;

ENDDOCUMENT();
