## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Jordan form)
## Date(11/5/2016)
## Institution(NAU)
## Author(Nandor Sieben)
## Level(3)
## MO(1)
## KEYWORDS('linear algebra','similar')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context("Matrix");


$l1=non_zero_random(-4,1,1);
$l2=$l1+random(1,3,1);
$l2++ if ($l2==0);
$D=Matrix([
  [$l1,0,0,0],
  [0,$l1,1,0],
  [0,0,$l1,0],
  [0,0,0,$l2]
]);
{
$P=Matrix([
  [random(-2,2,1),random(-2,2,1),random(-3,3,1),random(-1,1,1)],
  [random(-2,2,1),random(-2,2,1),random(-3,3,1),random(-3,3,1)],
  [random(-2,2,1),random(-3,3,1),random(-3,3,1),random(-2,2,1)],
  [random(-1,1,1),random(-2,2,1),random(-3,3,1),random(-2,2,1)]
]);
redo if (abs($P->det) != 1);
redo if ($P == Matrix([[1,0,0,0],[0,1,0,0],[0,0,1,0],[0,0,0,1]]));
$Pi=$P->inverse;
$A=$P*$D*$Pi;
}

$multians = MultiAnswer($P, $D)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my ($s1,$s2) = @{$student};
my $sP=Matrix($s1);
if ($sP->det == 0) {
  $self->setMessage(1,"Not invertible.");
  return 0; 
} 
my $sD=Matrix($s2);
if ($sD != $D) {
  $self->setMessage(2,"Incorrect Jordan form.");
  return 0; 
} 
if ($sP*$D != $A * $sP) {
  return 0; 
} 
return 1;
}
);

Context()->texStrings;
BEGIN_TEXT
Let \[ A=$A. \] Find a matrix \( P \) such that \( D=P^{-1} A P \) is the Jordan canonical form of \( A \). 
The Jordan form is upper triangular. The blocks are ordered increasingly by eigenvalue and then by block size.
$BR
\( P = \)
\{ $multians->ans_array(3) \}, 
\( D = \)
\{ $multians->ans_array(3) \}

END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

;
COMMENT('Diagonalizable matrices. Eigenvalues are small integers. Blck sizes: 2,1;1 ');
ENDDOCUMENT();

