## KEYWORDS('linear algebra','span')

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Kernel and image)
## Date('10/31/2013')
## Author('Nandor Sieben')
## Institution('NAU')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context('Matrix');

{
$f1 = list_random(-5,-3,-2,-4,-1,1,2,4,3,5);
$f2 = list_random(-5,-3,-2,-4,2,4,3,5);
$f3 = list_random(-5,-3,-2,-4,-1,1,2,4,3,5);
redo if (abs($f1) == abs($f3));
$a11= $f1*$f2;
$a12= $f1;
$a21= $f2*$f3;
$a22= $f3;
$A = Matrix([[$a11,$a12],[$a21,$a22]]);
redo if ($A->det!=0);
$M=Matrix([[$a11,0,$a12,0],[0,$a11,0,$a12],[$a21,0,$a22,0],[0,$a21,0,$a22]]);
$rank=$M->order_LR;
redo if ($M->order_LR != 2);
}

$b11= random(-5,5,1);
$b12= random(-5,5,1);
$b21= random(-5,5,1);
$b22= random(-5,5,1);

$B = Matrix([[$b11,$b12],[$b21,$b22]]);
$LB = $A*$B;

$zero = Matrix([[0,0],[0,0]]);

$K1 = Matrix([[0,-$a12],[0,$a11]]);
$K2 = Matrix([[-$a12,0],[$a11,0]]);

$R1 = Matrix([[$a11,0],[$a21,0]]);
$R2 = Matrix([[0,$a12],[0,$a22]]);

$kernel = MultiAnswer($K1, $K2)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
my $s0 = Matrix($s[0]);
my $s1 = Matrix($s[1]);

if (Matrix(
      [$s0->element(1,1),$s0->element(1,2), $s0->element(2,1),$s0->element(2,2)],
      [$s1->element(1,1),$s1->element(1,2), $s1->element(2,1),$s1->element(2,2)],
      [0,0,0,0],
      [0,0,0,0],
      )->order_LR!=2) {
  return 0;
}
return 0 if ($A*$s0 != $zero);
return 0 if ($A*$s1 != $zero);
return 1;
}
);

$image = MultiAnswer($R1, $R2)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @ss = @{$student};
my @cc = @{$correct};
my $ss0 = Matrix($ss[0]);
my $ss1 = Matrix($ss[1]);

if (Matrix(
      [$ss0->element(1,1),$ss0->element(1,2),
       $ss0->element(2,1),$ss0->element(2,2)],
      [$ss1->element(1,1),$ss1->element(1,2),
       $ss1->element(2,1),$ss1->element(2,2)],
      [$a11,0,$a21,0],
      [0,$a12,0,$a22],
      )->order_LR!=2) {
  return 0;
}

if (Matrix(
      [$ss0->element(1,1),$ss0->element(1,2),
       $ss0->element(2,1),$ss0->element(2,2)],
      [$ss1->element(1,1),$ss1->element(1,2),
       $ss1->element(2,1),$ss1->element(2,2)],
      [0,0,0,0],
      [0,0,0,0],
      )->order_LR!=2) {
  return 0;
}

return 1;
}
);

Context()->texStrings;
BEGIN_TEXT
$rank 
Let \(V=\mathbb{R}^{2\times 2}\) be the vector space of \( 2\times 2\) matrices and let \( L :V\to V\) be defined by \( L(X) = $A X \).
Hint: The image of a spanning set is a spanning set for the image.
$BR
a. Find \(L($B)=\) 
\{ $LB->ans_array(3) \}

$BR
b. Find a basis for \( \ker (L )\):
$BR
{ \{ $kernel->ans_array(3) \} , \{ $kernel->ans_array(3) \} }

$BR

c. Find a basis for \( \text{ran} (L )\):
$BR
{ \{ $image->ans_array(3) \} , \{ $image->ans_array(3) \} }


END_TEXT
Context()->normalStrings;

ANS( $LB->cmp );
ANS( $kernel->cmp );
ANS( $image->cmp );

COMMENT('MathObject version.');

ENDDOCUMENT();

