## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Associated matrices)
## Date(10/12/2013)
## Institution(NAU)
## Author(Nandor Sieben)
## Level(3)
## MO(1)
## KEYWORDS('linear algebra','magic square','matrix','linear operator')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "MatrixReduce.pl",
  "rank.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context("Fraction");
Context() -> parens -> set ("[" => {formMatrix => 1});
# Context("Matrix");

$B1 = Matrix([[1,0],[0,1]]);
$B2 = Matrix([[0,1],[1,0]]);
$M1 = Matrix([[1,1],[1,1]]);
$ML = Matrix([[0,-1],[-1,0]]);

$a = non_zero_random(-9,9,1);
$b = non_zero_random(-9,9,1);
$A = Matrix([[$a,$b],[$b,$a]]);
$sA = Compute($a+$b);
$LA = Matrix([[-$b,-$a],[-$a,-$b]]);


sub magic {
my ($m) = @_;
my $s1=$m->element(1,1)+$m->element(1,2); 
my $s2=$m->element(2,1)+$m->element(2,2); 
return 0 if ($s2 != $s1);
my $s4=$m->element(1,1)+$m->element(2,1); 
return 0 if ($s4 != $s1);
my $s5=$m->element(1,2)+$m->element(2,2); 
return 0 if ($s5 != $s1);
return 1;
}

$multians = MultiAnswer($B1, $B2)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);

$goodbasis=0;

if (! magic($s0)) {
  $self->setMessage(1,"Not half-magic.");
  return 0 ;
} 
if (! magic($s1)) {
  $self->setMessage(2,"Not half-magic.");
  return 0 ;
} 
if (rank(Matrix(
      [$s0->element(1,1),$s0->element(1,2),
       $s0->element(2,1),$s0->element(2,2)],
      [$s1->element(1,1),$s1->element(1,2),
       $s1->element(2,1),$s1->element(2,2)],
      [0,0,0,0],
      [0,0,0,0],
      )) != 2) {
  $self->setMessage(1,"Not independent.");
  $self->setMessage(2,"Not independent.");  
  return 0;
}
$goodbasis=1;
return 1;
}
);

sub mycheck {
  my ($correct, $student, $ansHash) = @_;
  return 0 if (! $goodbasis);
  my $t  = Matrix($student);
  my $Ls0=$s0-($s0->element(1,1)+$s0->element(1,2))*$M1;
  if ($Ls0 != ($t->element(1,1)*$s0 + $t->element(2,1)*$s1) ) {
    Value->Error("Incorrect first column"); # this does not seem to work, don't know why
    return 0;
  }
  my $Ls1=$s1-($s1->element(1,1)+$s1->element(1,2))*$M1;
  if ($Ls1 != ($t->element(1,2)*$s0 + $t->element(2,2)*$s1) ) {
    Value->Error("Incorrect second column"); # this does not seem to work, don't know why
    return 0;
  }
  return 1;
}


Context()->texStrings;
BEGIN_TEXT
A square matrix \( A \) is called $BITALIC half-magic $EITALIC 
if the sum of the numbers in each row and column is the same. 
The common sum in each row and column is denoted by \( s(A) \) 
and is called the $BITALIC magic sum $BITALIC of the matrix \(A\).
Let \( V \) be the vector space of \( 2\times 2\) half-magic squares. 

$BR

(a) Find 
$BR
\( s($A) = \) \{ $sA->ans_rule(3)\}.

$BR
Consider the linear operator \( L:V\to V \) that subtracts the magic sum from every entry in the matrix. 
More preciselly, \( L(A) := A-s(A) $M1\). 
$BR

(b) Find 
$BR
\( L($A) = \) \{ $LA->ans_array(3) \}.

$BR

(c) Find an ordered basis \( B \) for \( V \). 

$BR

\( B = ( \)
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \} 
\( ) \).

$BR

(d) Find the matrix \( [L]_B^B \) of \( L \) in your chosen basis \( B \).

$BR

\( [L]_B^B = \)
\{ $ML->ans_array(3) \}.


END_TEXT
Context()->normalStrings;

ANS( $sA->cmp() );
ANS( $LA->cmp() );
ANS( $multians->cmp() );
ANS( $ML->cmp( checker=>~~&mycheck ) );

;

ENDDOCUMENT();




