
## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(11/5/2016)
## Institution(NAU)
## Author(Nandor Sieben)
## Level(3)
## MO(1)
## KEYWORDS('linear algebra','similar')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context("Matrix");

{
$l1=non_zero_random(-4,2,1);
$l2=$l1+1;
$D=Matrix([[$l1,0,0],[0,$l2,0],[0,0,$l2]]);
{
$P=Matrix([[random(-3,3,1),random(-3,3,1),random(-3,3,1)],[random(-3,3,1),random(-3,3,1),random(-3,3,1)],[random(-3,3,1),random(-3,3,1),random(-3,3,1)]]);
redo if (abs($P->det) != 1);
$Pi=$P->inverse;
$A=$Pi*$D*$P;
}

{
$Q=Matrix([[random(-3,3,1),random(-3,3,1),random(-3,3,1)],[random(-3,3,1),random(-3,3,1),random(-3,3,1)],[random(-3,3,1),random(-3,3,1),random(-3,3,1)]]);
redo if (abs($Q->det) != 1);
$Qi=$Q->inverse;
$B=$Qi*$D*$Q;
}
redo if ($A==$B);
}
$R=$Qi*$P;
$Ri=$Pi*$Q;

$multians = MultiAnswer($Ri, $R)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my ($s1,$s2) = @{$student};
my $sPi=Matrix($s1);
if ($sPi->det == 0) {
  $self->setMessage(1,"Not invertible.");
  return 0 
} 
my $sP=Matrix($s2);
if ($sP->det == 0) {
  $self->setMessage(2,"Not invertible.");
  return 0 
} 
if ($sP->inverse != $sPi) {
  $self->setMessage(2,"Not the inverse.");
  return 0 
} 
return 0 if ($sPi*$B*$sP != $A); 
return 1;
}
);

Context()->texStrings;
BEGIN_TEXT
Show that \( A=$A \) and \( B=$B \) are similar matrices by finding an invertible matrix \( P \)
satisfying \( A=P^{-1} B P\).
$BR
\( P^{-1} = \)
\{ $multians->ans_array(3) \}, 
\( P = \)
\{ $multians->ans_array(3) \}


END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

;
COMMENT('Diagonalizable matrices. Eigenvalues are small integers, lambda1 != lambda2 = lambda3 ');
ENDDOCUMENT();


