
## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Associated matrices)
## Date(11/7/2013)
## Institution(NAU)
## Author(Nandor Sieben)
## Level(3)
## MO(1)
## KEYWORDS('linear algebra','similar')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context("Matrix");

$B=Matrix([[random(-3,3,1),random(-3,3,1)],[random(-3,3,1),random(-3,3,1)]]);
{
$P=Matrix([[random(-5,5,1),random(-5,5,1)],[random(-5,5,1),random(-5,5,1)]]);
redo if (abs($P->det) != 1);
$Pi=$P->inverse;
$A=$Pi*$B*$P;
redo if ($A==$B);
}

$multians = MultiAnswer($Pi, $P)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my ($s1,$s2) = @{$student};
my $sPi=Matrix($s1);
my $sP=Matrix($s2);

if ($sPi->det == 0) {
  $self->setMessage(1,"Not invertible.");
  return 0
}
my $sP=Matrix($s2);
if ($sP->det == 0) {
  $self->setMessage(2,"Not invertible.");
  return 0
}
if ($sP->inverse != $sPi) {
  $self->setMessage(2,"Not the inverse.");
  return 0
}

return 0 if ($sPi*$B*$sP != $A); 
return 1;
}
);

Context()->texStrings;
BEGIN_TEXT
Show that \( A=$A \) and \( B=$B \) are similar matrices by finding an invertible matrix \( P \)
satisfying \( A=P^{-1} B P\).
$BR
\( P^{-1} = \)
\{ $multians->ans_array(3) \}, 
\( P = \)
\{ $multians->ans_array(3) \}


END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );

;

ENDDOCUMENT();


