## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','magic square','basis')

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date('11/11/2013')
## Author('Nandor Sieben')
## Institution('NAU')

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);
$showPartialCorrectAnswers = 1;
TEXT(beginproblem()); 

Context('Matrix');

$I = Matrix([[1,0],[0,1]]);
$A = Matrix([[0,-1,1],[1,0,-1],[-1,1,0]]);
$B = Matrix([[-1,0,1],[2,0,-2],[-1,0,1]]);
$C = Matrix([[2,2,-1],[-2,1,4],[3,0,0]]);

sub magic {
my ($m) = @_;
my $s1=$m->element(1,1)+$m->element(1,2)+$m->element(1,3); 
my $s2=$m->element(2,1)+$m->element(2,2)+$m->element(2,3); 
return 0 if ($s2 != $s1);
my $s3=$m->element(3,1)+$m->element(3,2)+$m->element(3,3); 
return 0 if ($s3 != $s1);
my $s4=$m->element(1,1)+$m->element(2,1)+$m->element(3,1); 
return 0 if ($s4 != $s1);
my $s5=$m->element(1,2)+$m->element(2,2)+$m->element(3,2); 
return 0 if ($s5 != $s1);
my $s6=$m->element(1,3)+$m->element(2,3)+$m->element(3,3); 
return 0 if ($s6 != $s1);
my $s7=$m->element(1,1)+$m->element(2,2)+$m->element(3,3); 
return 0 if ($s7 != $s1);
my $s8=$m->element(1,3)+$m->element(2,2)+$m->element(3,1); 
return 0 if ($s8 != $s1);
return 1;
}

$multians = MultiAnswer($A, $B, $C)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$s0 = Matrix($s[0]); 
$s1 = Matrix($s[1]);
$s2 = Matrix($s[2]);

return 0 if (! magic($s0)); 
return 0 if (! magic($s1)); 
return 0 if (! magic($s2)); 
if (Matrix(
      [$s0->element(1,1),$s0->element(1,2),$s0->element(1,3),
       $s0->element(2,1),$s0->element(2,2),$s0->element(2,3),
       $s0->element(3,1),$s0->element(3,2),$s0->element(3,3)],
      [$s1->element(1,1),$s1->element(1,2),$s1->element(1,3),
       $s1->element(2,1),$s1->element(2,2),$s1->element(2,3),
       $s1->element(3,1),$s1->element(3,2),$s1->element(3,3)],
      [$s2->element(1,1),$s2->element(1,2),$s2->element(1,3),
       $s2->element(2,1),$s2->element(2,2),$s2->element(2,3),
       $s2->element(3,1),$s2->element(3,2),$s2->element(3,3)],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0],
      [0,0,0,0,0,0,0,0,0]
      )->order_LR==3) {
  return 1;
}
return 0;
}
);

$LB=Matrix([[-1,-2,4],[0,1,0],[0,0,1]]);

sub mycheck {
  my ($correct, $student, $ansHash) = @_;
  $sLB=Matrix($student);
  return 0 if ($s0->transpose != ($sLB->element(1,1)*$s0 + $sLB->element(2,1)*$s1 + + $sLB->element(3,1)*$s2) );
  return 0 if ($s1->transpose != ($sLB->element(1,2)*$s0 + $sLB->element(2,2)*$s1 + + $sLB->element(3,2)*$s2) );
  return 0 if ($s2->transpose != ($sLB->element(1,3)*$s0 + $sLB->element(2,3)*$s1 + + $sLB->element(3,3)*$s2) );
  return 1;
}

$D=Matrix([[-1,0,0],[0,1,0],[0,0,1]]);
$P=Matrix([[1,2,-1],[0,0,1],[0,1,0]]);

$multians2 = MultiAnswer($D, $P)->with(
singleResult => 1,
checker => sub {
my ( $correct, $student, $self ) = @_;
my @s = @{$student};
my @c = @{$correct};
$sD = Matrix($s[0]); 
$sP = Matrix($s[1]);
return 0 if ($sD != $sP->inverse*$sLB*$sP); 
return 0 if ($sD->element(1,2)!=0); 
return 0 if ($sD->element(1,3)!=0);
return 0 if ($sD->element(2,1)!=0);
return 0 if ($sD->element(2,3)!=0);
return 0 if ($sD->element(3,1)!=0);
return 0 if ($sD->element(3,2)!=0);     
return 1;
}
);


Context()->texStrings;
BEGIN_TEXT
A square matrix is magic if the sum of the numbers in each row, column and main diagonal is the same. 

$BR
a .Find a basis \( B \) for the vector space \( V \) of \( 3\times 3\) magic squares. 
$BR
\( B = \lbrace \)
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \},
\{ $multians->ans_array(3) \} 
\(\rbrace \).

$BR
b. Let \( L :V\to V\) be the linear transformation defined by \( L(M)=M^T \). Find the matrix of \( L \) in the basis \( B \).
$BR
\( [L]_B^B = \) \{ $LB->ans_array(3) \}

$BR
c. Find the determinant
$BR
\( \det([L]_B^B) = \) \{ ans_rule(5) \}

$BR
d. Find the characteristic polynomial \( f(x) \) of \( L \).
$BR
\( f(x)  = \) \{ ans_rule(20) \}

$BR
e. Find a diagonal matrix \( D \) and an invertible matrix \( P \) such that \( D = P^{-1} [L]_B^B P\).
$BR
\( D  = \) \{ $multians2->ans_array(3) \},  \( P  = \) \{ $multians2->ans_array(3) \}


END_TEXT
Context()->normalStrings;

# install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );
ANS( $LB->cmp( checker=>~~&mycheck ) );
ANS( Formula("-1")->cmp() ); 
ANS( Formula("-1+x+x^2-x^3")->cmp() ); 
ANS( $multians2->cmp() );

COMMENT('MathObject version. This problem is a bit difficult. CAS may be useful');

ENDDOCUMENT();
