## DESCRIPTION
## Algebraic equations
## ENDDESCRIPTION

## KEYWORDS('Equations')

## DBsubject('Algebra') 
## BookTitle('Algebra: Form and Function')
## DBchapter('Basic Algebra') 
## BookChapter('Rules for Expressions and the Reasons for Them')
## DBsection('Expanding and Factoring') 
## BookSection('Expanding and Factoring')
## Date('01/01/10')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Algebra: Form and Function')
## EditionText1('1')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## Section1('2.3')
## Problem1('38')


########################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
);

TEXT(beginproblem());

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(n=>"Real");

$k = random(2,4,2);
$a = random(3,5,2);
$b = random(6,9,1);

$p2 = $k;
$p1 = $k * ($a - $b);
$p0 = $k * $a * $b;

$expression = "\( $p2 n^2 + $p1 n - $p0 \)";

$fac0 = Formula("$k");
$fac1 = Formula("(n + $a)");
$fac2 = Formula("(n - $b)");

$multians = MultiAnswer($fac0, $fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f0stu, $f1stu, $f2stu ) = @{$student};
      my ( $f0, $f1, $f2 ) = @{$correct};
      if ( ($f0 == $f0stu && $f1 == $f1stu && $f2 == $f2stu) ||
           ($f0 == $f0stu && $f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1,1];
      } else {
          if      ($f0 == $f0stu && ($f1 == $f1stu || $f2 == $f1stu)) {
              return [1,1,0];
          } elsif ($f0 == $f0stu && ($f1 == $f2stu || $f2 == $f2stu)) {
              return [1,0,1];
          } elsif ($f0 == $f0stu && ($f1 != $f1stu && $f2 != $f2stu)) {
              return [1,0,0];
          } elsif ( ($f0 != $f0stu && ($f1 == $f1stu && $f2 == $f2stu)) ||
                    ($f0 != $f0stu && ($f1 == $f2stu && $f2 == $f1stu)) ) {
              return [0,1,1];
          } elsif ($f0 != $f0stu && ($f1 == $f1stu || $f2 == $f1stu)) {
              return [0,1,0];
          } elsif ($f1 != $f0stu && ($f1 == $f2stu || $f2 == $f2stu)) {
              return [0,0,1];
          } else {
              return [0,0,0];
          }
      }
  }
);



################################################
#  Main text

BEGIN_TEXT

Factor the expression $expression.  Simplify your 
answer as much as possible, and put the greatest
common factor in the first answer box.
$BR
$BR
\{ $multians->ans_rule(5) \}
\( \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \)
\{ AnswerFormatHelp("formula") \}

END_TEXT


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

ENDDOCUMENT();