# DESCRIPTION
# Preview Activity from _Active Calculus_ by Matthew Boelkins
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Compositions and combinations of functions)
## Institution(Westmont College)
## Author(David Hunter)
## Level(3)
## MO(1)
## TitleText1('Active Calculus')
## AuthorText1('Matthew Boelkins')
## EditionText1('2015')
## Section1('2.5')
## Problem1('2 5')
## KEYWORDS('derivative', 'composition', 'product rule', 'quotient rule')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "parserRadioButtons.pl",
  "PGcourse.pl"
);

Context("Numeric");
$showPartialCorrectAnswers = 0;

$r = Formula("(tan(x))^2");
$f1 = Formula("x^2");
$g1 = Formula("tan(x)");
$f2 = Formula("tan(x)");
$g2 = Formula("tan(x)");

$ans_str = [ 
"A composition \(f(g(x))\) of basic functions",
"A sum \(f(x)+g(x)\) of basic functions",
"A product \(f(x) \cdot g(x)\) of basic functions",
"A quotient \(f(x)/g(x)\) of basic functions"];

$mc = RadioButtons($ans_str, $ans_str->[0],
    labels => ["A", "B", "C", "D"] );

$ma = MultiAnswer($mc,$f1,$f2)->with(
  singleResult => 1,
  checker => sub {
    my ($correct, $student, $self, $ans) = @_;
    if($mc->{orderedChoices}[substr($student->[0], 1)] eq $ans_str->[0]) {
      return $student->[1] == $f1 && $student->[2] == $g1;
    } elsif($mc->{orderedChoices}[substr($student->[0], 1)] eq $ans_str->[2]) {
      return $student->[1] == $f2 && $student->[2] == $g2;
    } else {
      return 0;
    }
  });

TEXT(beginproblem());
Context()->texStrings;
BEGIN_PGML
Let [`r(x) = [$r]`]. Which of the following best describes its fundamental 
algebraic structure?
[______]{$ma}

where

[`f(x) = `] [_______________]{$ma}

[`g(x) = `] [_______________]{$ma}

END_PGML

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
The function \(r(x)\) can be written as a composition \(f(g(x))\) where the outside function is \(f(x)=$f1\) and the inside function is \(g(x)=$g1\), or as a product \(f(x) g(x)\) where the two functions are \(f(x)=$f2\) and \(g(x)=$g2\).
END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
