# DESCRIPTION
# Preview Activity from _Active Calculus_ by Matthew Boelkins
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Compositions and combinations of functions)
## Institution(Westmont College)
## Author(David Hunter)
## Level(3)
## MO(1)
## Static(1)
## TitleText1('Active Calculus')
## AuthorText1('Matthew Boelkins')
## EditionText1('2015')
## Section1('2.5')
## Problem1('2 5')
## KEYWORDS('derivative', 'composition', 'product rule', 'quotient rule')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "parserRadioButtons.pl",
  "PGcourse.pl"
);

Context("Numeric");
$showPartialCorrectAnswers = 0;

$r = Formula("(tan(x))^2");
$f1 = Formula("x^2");
$g1 = Formula("tan(x)");

$ans_str = [ 
"A composition \(f(g(x))\) of basic functions",
"A sum \(f(x)+g(x)\) of basic functions",
"A product \(f(x) \cdot g(x)\) of basic functions",
"A quotient \(f(x)/g(x)\) of basic functions"];

$mc = RadioButtons($ans_str, $ans_str->[0],
    labels => ["A", "B", "C", "D"],
    values => ["A", "B", "C", "D"]);

$ma = MultiAnswer($mc,$f1,$g1)->with(
  singleResult => 1,
  checker => sub {
    my ($correct, $student, $self, $ans) = @_;
    if ($student->[0] eq "A"){
      return $student->[1] == $f1 && $student->[2] == $g1;
    } elsif ($student->[0] eq "C"){
      return $student->[1] == $g1 && $student->[2] == $g1;
    } else {
      return 0;
    }
  });

BEGIN_PGML
Let [`r(x) = [$r]`]. Which of the following best describes its fundamental 
algebraic structure?

[______]{$ma}

where

[`f(x) = `] [_______________]{$ma}

[`g(x) = `] [_______________]{$ma}

END_PGML

BEGIN_PGML_SOLUTION
The function [`r(x)`] can be written as a composition [`f(g(x))`] where the outside function is [`f(x)=[$f1]`] and the inside function is [`g(x)=[$g1]`], or as a product [`f(x)\cdot g(x)`] where the two functions are [`f(x)=[$g1]`] and [`g(x)=[$g1]`].
END_PGML_SOLUTION

ENDDOCUMENT();
