##DESCRIPTION
##  lhp4_51-54.pg.
##ENDDESCRIPTION

## hcao tagged and PAID on 2-20-2004

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring trinomials)
## Date(6/3/2002)
## Institution(csulb)
## Author(Tangan Gao)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.4')
## Problem1('46')
## KEYWORDS('algebra, factoring, perfect square trinomial')

DOCUMENT();     
loadMacros(
  "PGstandard.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


$a=random(2,5,1);
$b=random(2,7,1);
$c=$a+3;
$d=$b+2;
$a1=$a*$c;
$a2=$a*$d+$b*$c;
$a3=$b*$d;

$mp = MultiAnswer($a, $b, $c, $d)->with(
        singleResult => 1,
        checker => sub {
          my ($correct,$student,$self) = @_;  # get the parameters
          return 0 if (($student->[0]*$student->[2]) != ($correct->[0]*$correct->[2]));
          return 0 if (($student->[1]*$student->[3]) != ($correct->[1]*$correct->[3]));
          return 0 if ($student->[0]*$student->[3]+$student->[1]*$student->[2] !=
    $correct->[0]*$correct->[3]+$correct->[1]*$correct->[2]);
          return 1;
        },
    );



BEGIN_TEXT


Factor the trinomial \($a1 x^2 - $a2 x +$a3\):$BR$BR
\[$a1 x^2 -$a2 x+ $a3=(Ax-B) (Cx-D) \]
$BR
\(A=\) \{ $mp->ans_rule(5)\}, \(B=\) \{$mp->ans_rule(5)\}, \(C=\) \{$mp->ans_rule(5)\}, and \(D=\) \{$mp->ans_rule(5)\}
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();
