# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Simplify the following expression
#     x^m * x^n
#
# We use an INTEGER value for x on the interval [2,20]
#
# This problem chooes m and n to be POSITIVE, m=1
#
# Last edited: Carl Yao 9/8/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions')
## KEYWORDS('multiply','exponent','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "answerHints.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
Context()->flags->set(formatStudentAnswer=>'parsed',reduceConstants=>0); 
Context()->operators->undefine('*','* ', ' *', '/', '/ ', ' /' , '//','+','-');     
Context()->operators->set(
   '^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
   '**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef});

# m and n are the exponents
$mybase=random(2,9,1);
$m=1;
$n=random(2,10,1);

$total = $m+$n;
$left = Formula("$mybase");
$right = Formula("$mybase^($n)");
$ans = Formula("$mybase^($total)");





##############################################
TEXT(beginproblem());
BEGIN_PGML
Use the properties of exponents to simplify the following
   
    [`[$left]\cdot[$right]`]

    [____________]


END_PGML

$wrong=$m*$n;
ANS($ans -> cmp(checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroPow=>1);
delete $correct->{test_values}, $student->{test_values};
Value->Error("Completely simplify according to the rules of exponents, not in any other way") unless (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroPow=>0);
return 1;
}) ->
  withPostFilter(AnswerHints(
    [Compute("$mybase^($wrong)")] =>
    "When multiplying terms with the same base, you do not multiply the exponents.")));


BEGIN_PGML_SOLUTION

Note that [`[$mybase]=[$mybase]^1`]. We _add_ the exponents as follows

    [`\begin{aligned}
           [$mybase]\cdot [$mybase]^{[$n]} &= [$mybase]^{1}\cdot [$mybase]^{[$n]} \\
                             &=[$mybase]^{[$m]+[$n]}\\
                             &=[$mybase]^{[$total]}
       \end{aligned}`]



END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
