# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
#
#  Rewrite the following expression using only positive exponents.
#
#   a x^-m y^n z^-p (b x^q)^-r = a*b^(-r)*x^{-m-qr} y^n z^-p
#
#                              =      a y^n
#                                ----------------
#                                b^r x^{m+qr} z^p       
#
#   Keep all variables integer and fairly small (<10), especially b and r
# 
# Last updated: Hughes 9/4/13, Carl Yao, 7/03/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','simplify','fraction')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
"contextRationalFunction.pl",
  "contextFraction.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Fraction"); 
$var1 = RandomVariableName(type=>'variable');
$var2 = $var1;
$var2 = RandomVariableName(type=>'variable') while($var2 eq $var1);
$var3 = $var2;
$var3 = RandomVariableName(type=>'variable') while($var3 eq $var2 or $var3 eq $var1);

# actually, let's just call them x, y, z
$var1 = 'x'; $var2 = 'y'; $var3 = 'z';

Context()->variables->are($var1=>'Real',$var2=>'Real',$var3=>'Real');

# exponents
$m = random(3,8,1);
$n = random(2,8,1);
$p = random(2,8,1);
$q = random(2,8,1);
$r = random(2,4,1);

# coefficient
$a = 2;
$b = 2;
while(gcd($a,$b**$r)!=1 or ($b**$r)>100)
{
   $a = random(2,10,1);
   $b = random(2,5,1);
   $r = random(2,4,1);
}

# total exponent
$total = ($m+$q*$r);
$totalCoeff = Fraction($a,$b**$r);

#   a x^-m y^n z^-p (b x^q)^-r = a*b^(-r)*x^{-m-qr} y^n z^-p
#
#                              =      a y^n
#                                ----------------
#                                b^r x^{m+qr} z^p       

Context("RationalFunction-Strict");
Context()->variables->are($var1=>'Real',$var2=>'Real',$var3=>'Real');
#Make it so that fractions like 1/2 in formulas are not computed to 0.5
Context()->flags->set(reduceConstants=>0);

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});



$totaldenom = $b**$r;
$ans = Formula("($a*$var2^$n)/($totaldenom*($var1^($total)*$var3^$p))")->reduce;

##############################################
TEXT(beginproblem());

#   a x^-m y^n z^-p (b x^q)^-r = a*b^(-r)*x^{-m-qr} y^n z^-p
#
#                              =      a y^n
#                                ----------------
#                                b^r x^{m+qr} z^p     

BEGIN_PGML
Simplify the following expression, and write your answer using only _positive_ exponents.

    [`
        [$a][$var1]^{[$m*(-1)]}[$var2]^{[$n]}[$var3]^{[$p*(-1)]}\left([$b][$var1]^{[$q]}\right)^{[$r*(-1)]}=
     `][______________________]

This problem initially looks quite complicated, but as long as you gather each of the variables 
together and remember your properties of exponents, it's not as bad as it first looks- honest!
END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 



BEGIN_PGML_SOLUTION

We use the properties of exponents to help us here, remember that [`a^m\cdot a^n = a^{m+n}`] and [`\left(a^m\right)^n=a^{mn}`].

    [`
           \begin{aligned}
      [$a][$var1]^{[$m*(-1)]}[$var2]^{[$n]}[$var3]^{[$p*(-1)]}\left([$b][$var1]^{[$q]}\right)^{[$r*(-1)]}
                  &=  [$a][$var1]^{[$m*(-1)]}[$var2]^{[$n]}[$var3]^{[$p*(-1)]}[$b]^{[$r*(-1)]}[$var1]^{[$q]*([$r*(-1)])}\\
                  &=  [$a][$var1]^{[$m*(-1)]}[$var2]^{[$n]}[$var3]^{[$p*(-1)]}[$b]^{[$r*(-1)]}[$var1]^{[$q*$r*(-1)]}\\
                  &=  [$a][$var1]^{[$m*(-1)]-[$q*$r]}[$var2]^{[$n]}[$var3]^{[$p*(-1)]}\frac{1}{[$b]^{[$r]}}\\
                  &=  [$a][$var1]^{[$m*(-1)-$q*$r]}[$var2]^{[$n]}[$var3]^{[$p*(-1)]}\frac{1}{[$b]^{[$r]}}\\
                  & = [$ans]
           \end{aligned}
    `]

END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
