# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
#
#  Rewrite the following expression using only positive exponents.
#
#    (a x^m)^p
#   ----------
#     (x^-n)
#
# We make sure that n>m so that x is raised to a negative power. 
#
# We choose a to be integer on [2,6] and p to be integer on [2,3].
# 
# Last updated: Carl Yao, 7/03/13
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponent','negative','fraction','simplify')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextLimitedPolynomial.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

#    (a x^m)^p
#   ----------
#     (x^-n)

Context("LimitedPolynomial-Strict");
$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});


# exponents
$m = random(3,12,1);
$p = random(2,3,1);
$n = random(2,15,1);

# total exponent
$total = (($m)*($p)+$n);

# coefficient
$a = random(2,6,1);
$totalCoeff = $a**$p;
$ans = Formula("($totalCoeff*$var^($total))");

##############################################
TEXT(beginproblem());

BEGIN_PGML
Simplify the following expression, and write your answer using only _positive_ exponents.

    [`
        \displaystyle\frac{\left([$a][$var]^{[$m]}\right)^{[$p]}}{[$var]^{[$n*(-1)]}}=
     `][____________]

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 



BEGIN_PGML_SOLUTION

We use the properties of exponents to help us here, remembering that [``\frac{a^m}{a^n} = a^{m-n}``] 
and also that [`\left(a^m\right)^n=a^{mn}`]

    [`
           \begin{aligned}
                 \displaystyle\frac{\left([$a][$var]^{[$m]}\right)^{[$p]}}{[$var]^{[$n*(-1)]}}&= 
\displaystyle\frac{\left([$a]\right)^{[$p]}[$var]^{[$m]\cdot[$p]}}{[$var]^{[$n*(-1)]}}\\
 &= \displaystyle\frac{[$a**$p][$var]^{[$m*$p]}}{[$var]^{[$n*(-1)]}}\\
 &= [$a**$p][$var]^{[$m*$p]-(-[$n])}\\
                                    & = [$ans]
           \end{aligned}
    `]

END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
