##DESCRIPTION
##  Differential Equation
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(02/07/2012)
## Institution(Miami University and Hope College and Cleveland State University)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson, Barbara Margolius)
## MLT(Exact_Solve)
## Level(3)
## MO(1)
## TitleText1('A First Course in Differential Equation')
## AuthorText1('Dennis G.Zill')
## EditionText1('9')
## Section1('2.4')
## Problem1('3')
## KEYWORDS('Differential equaltion')

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric")->variables->add(y=>"Real",c=>"Real");
$a = random(2,8,1);
$b = random(2,8,1);
$c = random(1,3,1);

##############################################################
#
#  Text
#
#
$answer = Formula("$b x^(2)/2+$a xy-2y^($c)")->reduce;
$dx = Formula("$b x+$a y")->reduce;
$dy = Formula("$a x-$c*2 y^{$c-1}")->reduce;


Context()->texStrings;
BEGIN_TEXT

Solve the following differential equation: 
$BR
$BR
\(\   ($dx)dx+($dy)dy=0\).

$BR
$BR
\{ ans_rule(30) \} 
=
constant. 
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $context->variables->add('C1'=>'Parameter');
      my $c1 = Formula($context,'C1');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer + $c1");
      return $correct == $student;
    }
) );

COMMENT('');

ENDDOCUMENT();  
