##DESCRIPTION
##  Differential Equation
##ENDDESCRIPTION

##KEYWORDS('Differential equaltion')

## DBsubject('Differential Equation')
## DBchapter('First-Order DE')
## DBsection('Section 2.4 Exact Equations')
## Date('02/07/2012')
## Author('Anna Ghazaryan, Fuli Zhang, Barbara Margolius')
## Institution('Muohio and Cleveland State University')
## TitleText1('A First Course in Differential Equation')
## EditionText1('9')
## AuthorText1('Dennis G.Zill')
## Section1('2.4')
## Problem1('9')

DOCUMENT();      

loadMacros(
   "PGstandard.pl",
   "MathObjects.pl",
   "AnswerFormatHelp.pl",
   "unionLists.pl",
   "parserPrime.pl",
   "parserPopUp.pl",
   "answerHints.pl",
   "parserImplicitPlane.pl",

 );

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric")->variables->add(y=>"Real",c=>"Real");
$a = random(2,5,1);
$b = random(2,5,1);
$c = random(2,5,1);

##############################################################
#
#  Text
#
#
$answer = Formula("xy^($a)+y^($c)cos(x)-x^($b)/$b")->reduce;

$lhs=Formula("x^($b-1)-y^($a)+y^($c)sin x")->reduce;
$rhs=Formula("$a xy^($a-1)+$c y^($c-1)cos x")->reduce;

Context()->texStrings;
BEGIN_TEXT

Solve the following differential equation: 
$BR
$BR
\(\displaystyle ($lhs)dx=($rhs)dy\).

$BR
$BR
\{ ans_rule(30) \} 
=
constant. 
\{ AnswerFormatHelp("formulas") \}



END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $context->variables->add('C1'=>'Parameter');
      my $c1 = Formula($context,'C1');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer + $c1");
      return $correct == $student;
    }
) );

COMMENT('MathObject version');


ENDDOCUMENT();        
    