##DESCRIPTION
##  Differential Equation
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(02/07/2012)
## Institution(Miami University and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Verification)
## Level(2)
## MO(1)
## TitleText1('A First Course in Differential Equation')
## AuthorText1('Dennis G.Zill')
## EditionText1('9')
## Section1('2.4')
## Problem1('7')
## KEYWORDS('Differential equaltion')

########################################################################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric")->variables->add(
t=>"Real"
);



Context()->flags->set(
  reduceConstants=>0, # no decimals
  reduceConstantFunctions=>1, # combine 4+5*2?
  formatStudentAnswer=>'parsed', # no decimals
);

$answer[0] = PopUp(["Choose","Exact","Not Exact"],"Not Exact");
$answer[1] = PopUp(["Choose","Exact","Not Exact"],"Exact");

#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Are the following differential equations exact? (You have only one attempt! Submit both answers at the same time)
$BR
$BR
(a) \{ $answer[0]->menu() \} \((x^{2}-y^{2})dx+(x^{2}-2xy)dy=0\).  
$BR
$BR
(b)  \{ $answer[1]->menu() \} \((x+y)^{2}dx+(2xy+x^{2}-1)dy=0\).  

END_TEXT
Context()->normalStrings;


############################
#  Answers

$showPartialCorrectAnswers = 0;

foreach my $i (0..1) {
ANS( $answer[$i]->cmp(
  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( ($numOfAttempts < 2) && ($correct == $student) ) { 
       return 1; 
    } elsif ( ($numOfAttempts > 1) && ($correct == $student) ) {
       Value->Error("Your answer is correct, but your score did not change since you used up the attempts allowed."); 
    } else {
       return 0;
    };
  }
));

}

#ANS( $answer[0]->cmp() );
#ANS( $answer[1]->cmp() );


COMMENT('Uses a custom answer checker to limit the number of attempts that count for credit to one.');

ENDDOCUMENT();
