##DESCRIPTION
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(06/08/2012)
## Institution(Cleveland State University)
## Author(Anna Ghazaryan, Fuli Zhang, Barbara Margolius)
## MLT(Exact_Solve)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Differential equations Theory Technique and Practice')
## AuthorText1('George F. Simmons')
## EditionText1('')
## Section1('1.5')
## Problem1('3')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();   

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#


Context("Numeric");
Context()->variables->are(
x=>"Real",
t=>"Real",y=>"Real",
c=>"Real",
);

$a=random(3,6,1);

##############################################################
#
#  Text
#
#

$answer = Compute("$a xy-x^{$a}+y^{$a}");
$dx = Formula("y-x^{$a-1}")->reduce;
$dy = Formula("x+y^{$a-1}")->reduce;


Context()->texStrings;
BEGIN_TEXT

Solve the following differential equation: 
$BR
$BR
\(\displaystyle   ($dx)dx+($dy)dy=0\).

$BR
$BR
\{ ans_rule(30) \} 
=
constant. 
\{ AnswerFormatHelp("formulas") \}





END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $context->variables->add('C1'=>'Parameter');
      my $c1 = Formula($context,'C1');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer + $c1");
      return $correct == $student;
    }
) );


;


ENDDOCUMENT();  
