##DESCRIPTION
##ENDDESCRIPTION


## Edition('7')
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(05/28/2012)
## Institution(Muohio and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Solve)
## Level(3)
## MO(1)
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## EditionText1('')
## Section1('2.4')
## Problem1('5')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();   

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#


Context("Numeric");
Context()->variables->are(
x=>"Real",
y=>"Real",
);


##############################################################
#
#  Text
#
#
$a=random(2,9,1);
$answer = Compute("$a*x^(2)y^(2)-3x+4y");




Context()->texStrings;
BEGIN_TEXT

Solve the following differential equation: 
$BR
$BR
\(\   ($a xy^{2}-3)dx+($a x^{2}y+4)dy=0\).

$BR
$BR
\{ ans_rule(30) \} 
=
constant. 
\{ AnswerFormatHelp("formulas") \}





END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $context->variables->add('C1'=>'Parameter');
      my $c1 = Formula($context,'C1');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer + $c1");
      return $correct == $student;
    }
) );

;

ENDDOCUMENT();  
