##DESCRIPTION
##ENDDESCRIPTION


## Edition('7')
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(05/28/2012)
## Institution(Muohio and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Initial_Value)
## Level(3)
## MO(1)
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## EditionText1('')
## Section1('2.4')
## Problem1('25')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();   

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#


Context("Numeric");
Context()->variables->are(
x=>"Real",
y=>"Real",
);


##############################################################
#
#  Text
#
#
$a=random(1,4,1);
$answer = Compute("y^(2)sinx-x^(3)y-x^(2)+ylny-y-(e^($a)ln(e^($a))-e^($a))")->reduce;




Context()->texStrings;
BEGIN_TEXT

Solve the following initial value problem: 
$BR
$BR
\(\displaystyle   (y^{2}\cos(x)-3x^{2}y-2x)dx+(2y\sin(x)-x^{3}+\ln(y))dy=0,\ y(0)=e^{$a}\).

$BR
$BR
\{ ans_rule(30) \} 
=
0. 
\{ AnswerFormatHelp("formulas") \}


END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer");
      return $correct == $student;
    }
) );

;

ENDDOCUMENT();  
