##DESCRIPTION
##  Differential Equation
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(02/07/2012)
## Institution(Miami University and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Initial_Value)
## Level(3)
## MO(1)
## TitleText1('A First Course in Differential Equation')
## AuthorText1('Dennis G.Zill')
## EditionText1('9')
## Section1('2.4')
## Problem1('21')
## KEYWORDS('Differential equaltion')

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric")->variables->add(y=>"Real",c=>"Real");
$a = random(2,5,1);
$b = Formula("-(1+3*$a^{2})")->reduce;

##############################################################
#
#  Text
#
#
#$answer = Formula("-4+x^(3)+3x^(2)y+3xy^(2)-3y");
$answer = Formula("$b+x^(3)+3x^(2)y+3xy^(2)-3y");
$dx = Formula("(x+y)^{2}")->reduce;
$dy = Formula("2xy+x^{2}-1")->reduce;

Context()->texStrings;
BEGIN_TEXT

Solve the following initial-value problem: 
$BR
$BR
\(\displaystyle (x+y)^{2}dx+(2xy+x^{2}-1)dy=0, \ y(1)=$a\)

$BR
$BR
\{ ans_rule(30) \} 
=
0. 
\{ AnswerFormatHelp("formulas") \}


END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#


ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer");
      return $correct == $student;
    }
) );


;


ENDDOCUMENT();        
    
