##DESCRIPTION
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(06/08/2012)
## Institution(Cleveland State University)
## Author(Anna Ghazaryan, Fuli Zhang, Barbara Margolius)
## MLT(Exact_Solve)
## Level(3)
## MO(1)
## TitleText1('Differential equations Theory Technique and Practice')
## AuthorText1('George F. Simmons')
## EditionText1('')
## Section1('1.5')
## Problem1('7')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();   

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
$choice = random(0,3,1);

Context("Numeric");
Context()->variables->are(
x=>"Real",
t=>"Real",y=>"Real",
c=>"Real",
);


##############################################################
#
#  Text
#
#

@answerArray = (Compute("xe^y+sin x cos y"),Compute("xe^y+sin x sin y"),Compute("xe^y+cos x cos y"),Compute("xe^y+cos x sin y"));
@lhsArray = (Formula("sin x sin y-xe^{y}")->reduce,Formula("-sin x cos y-xe^{y}")->reduce,Formula("cos x sin y-xe^{y}")->reduce,Formula("-cos x cos y-xe^{y}")->reduce);
@rhsArray = (Formula("e^{y}+cos x cos y")->reduce,Formula("e^{y}+cos x sin y")->reduce,Formula("e^{y}-sin x cos y")->reduce,Formula("e^{y}-sin x sin y")->reduce);
$answer = $answerArray[$choice];
$lhs = $lhsArray[$choice];
$rhs = $rhsArray[$choice];


Context()->texStrings;
BEGIN_TEXT

Solve the following differential equation: 
$BR
$BR
\(\displaystyle   ($lhs)dy=($rhs)dx\).

$BR
$BR
\{ ans_rule(30) \} 
=
constant. 
\{ AnswerFormatHelp("formulas") \}





END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer");
      return $correct == $student;
    }
) );





;




ENDDOCUMENT();  
