##DESCRIPTION
##ENDDESCRIPTION


## Edition('7')
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(05/28/2012)
## Institution(Muohio)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Verification)
## Level(2)
## Static(1)
## MO(1)
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## EditionText1('')
## Section1('2.4')
## Problem1('2 4 6')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
#Context("Numeric")->variables->add(t=>"Real",u=>"Real",r=>"Real",I0=>"Real",k=>"Real");


##############################################################
#
#  Text
#
#




$answer[0] = PopUp(["Choose","Exact","Not Exact"],"Not Exact");
$answer[1] = PopUp(["Choose","Exact","Not Exact"],"Exact");
$answer[2] = PopUp(["Choose","Exact","Not Exact"],"Not Exact");

#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Are the following differential equations exact? (You have only one attempt!  Submit all answers at the same time)
$BR
$BR
(a) \{ $answer[0]->menu() \} \((2x+y)dx-(x+6y)dy=0\).  
$BR
$BR
(b) \{ $answer[1]->menu() \} \((\sin(y)-y\sin(x))dx+(\cos(x)+x\cos(y)-y)dy=0 \).
$BR
$BR
(c) \{ $answer[2]->menu() \} \(\displaystyle \left(2y-\frac{1}{x}+\cos(3x) \right)\frac{dy}{dx}+\frac{y}{x^{2}}-4x^{3}+3y\sin(3x)=0\).

END_TEXT
Context()->normalStrings;


############################
#  Answers

$showPartialCorrectAnswers = 0;

foreach my $i (0..2) {
ANS( $answer[$i]->cmp(
  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( ($numOfAttempts < 2) && ($correct == $student) ) { 
       return 1; 
    } elsif ( ($numOfAttempts > 1) && ($correct == $student) ) {
       Value->Error("Your answer is correct, but your score did not change since you used up the attempts allowed."); 
    } else {
       return 0;
    };
  }
));

}


#ANS( $answer[0]->cmp() );
#ANS( $answer[1]->cmp() );
#ANS( $answer[2]->cmp() );


COMMENT('Uses a custom answer checker to 
limit the number of attempts that count for credit to one.');

ENDDOCUMENT();
