##DESCRIPTION
##  Differential Equation
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(02/07/2012)
## Institution(Miami University and Hope College and Cleveland State University)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson, Barbara Margolius)
## MLT(Exact_Initial_Value)
## Level(3)
## MO(1)
## TitleText1('A First Course in Differential Equation')
## AuthorText1('Dennis G.Zill')
## EditionText1('9')
## Section1('2.4')
## Problem1('22')
## KEYWORDS('Differential equaltion')

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric")->variables->add(y=>"Real",c=>"Real");
$a = random(2,4,1);
$b = Compute("-(1+$a)");
##############################################################
#
#  Text
#
#
$answer = Formula("e^(x)+xy+$a y+ye^(y)-e^(y)+$b")->reduce;

Context()->texStrings;
BEGIN_TEXT

Solve the following initial-value problem: 
$BR
$BR
\(\displaystyle (e^{x}+y)dx+($a+x+ye^{y})dy=0, \  y(0)=1\)

$BR
$BR
\{ ans_rule(30) \} 
=
0. 
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer");
      return $correct == $student;
    }
) );


COMMENT('');




ENDDOCUMENT();        
    
