##DESCRIPTION
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(06/08/2012)
## Institution(Cleveland State University)
## Author(Anna Ghazaryan, Fuli Zhang, Barbara Margolius)
## Level(3)
## MO(1)
## Static(1)
## TitleText1('Differential equations Theory Technique and Practice')
## AuthorText1('George F. Simmons')
## EditionText1('')
## Section1('1.5')
## Problem1('11')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();   

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#


Context("Numeric");
Context()->variables->are(
x=>"Real",
t=>"Real",y=>"Real",
c=>"Real",
);
Context()->variables->set(
  x=>{limits=>[-.8,.8]},
  y=>{limits=>[-.8,.8]},
);

##############################################################
#
#  Text
#
#

$answer = Compute("ln(abs((1+xy)/(1-xy)))-2x");
$dx = Formula("y/(1-x^{2}y^{2})")->reduce;
$dy = Formula("x/(1-x^{2}y^{2})")->reduce;

#\frac{y}{1-x^{2}y^{2}}dx+\frac{x}{1-x^{2}y^{2}}dy

Context()->texStrings;
BEGIN_TEXT

Solve the following differential equation: 
$BR
$BR
\(\displaystyle   dx=$dx dx+$dy dy\).

$BR
$BR
\{ ans_rule(30) \} 
=
constant. 
\{ AnswerFormatHelp("formulas") \}





END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(  $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $context->variables->add('C1'=>'Parameter');
      my $c1 = Formula($context,'C1');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * $answer + $c1");
      return $correct == $student;
    }
) );



ENDDOCUMENT();  
