##DESCRIPTION
##ENDDESCRIPTION


## Edition('7')
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(05/28/2012)
## Institution(Muohio and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Initial_Value)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## EditionText1('')
## Section1('2.4')
## Problem1('24')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();   

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#


Context("Numeric");
Context()->variables->are(
t=>"Real",
y=>"Real",
);


##############################################################
#
#  Text
#
#
$a=random(1,3,1);
#$answer = Compute("t^(2)/(4y^(4))-3/(2y^(2))-(1/(4*($a)^(4))-3/2/($a)^(2))")->reduce;
$answer = Compute("$a*sqrt((3*$a**2+ sqrt(9*$a**4 - (6*$a**2-1)*t**2))/(6*$a**2 - 1))")->reduce;

Context()->texStrings;
BEGIN_TEXT

Solve the following initial value problem: 
$PAR
\(\displaystyle   \left( \frac{3y^{2}-t^{2}}{y^{5}} \right) 
\frac{dy}{dt}+\frac{t}{2y^{4}}=0,\ y(1)=$a\).
$PAR
\(y(t)\)=\{ ans_rule(60) \} 
\{ AnswerFormatHelp("formulas") \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
#ANS(  $answer->cmp(checker => sub {
#      my ( $correct, $student, $self ) = @_;
#      my $context = Context()->copy;
#      return 0 if $student == 0;
#      $context->flags->set(no_parameters=>0);
#      $context->variables->add('C0'=>'Parameter');
#      my $c0 = Formula($context,'C0');
#      $student = Formula($context,$student);
#      $correct = Formula($context,"$c0 * $answer");
#      return $correct == $student;
#    }
#) );
ANS($answer->cmp());


ENDDOCUMENT();  
