##DESCRIPTION
##ENDDESCRIPTION


## Edition('7')
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Exact)
## Date(05/28/2012)
## Institution(Muohio and Hope College)
## Author(Anna Ghazaryan, Fuli Zhang, Paul Pearson)
## MLT(Exact_Verification)
## MLTleader(1)
## Level(2)
## Static(1)
## MO(1)
## TitleText1('Differential equations with Boundary-value problem')
## AuthorText1('Dennis G. Zill')
## EditionText1('')
## Section1('2.4')
## Problem1('14 16 18')
## KEYWORDS('differential equations')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
#Context("Numeric")->variables->add(t=>"Real",u=>"Real",r=>"Real",I0=>"Real",k=>"Real");


##############################################################
#
#  Text
#
#




$answer[0] = PopUp(["Choose","Exact","Not Exact"],"Exact");
$answer[1] = PopUp(["Choose","Exact","Not Exact"],"Exact");
$answer[2] = PopUp(["Choose","Exact","Not Exact"],"Exact");

#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Are the following differential equations exact? (You have only one attempt! Submit all answers at the same time)
$BR
$BR
(a) \{ $answer[0]->menu() \} \(\displaystyle \left(1-\frac{3}{y}+x\right)\frac{dy}{dx}+y=\frac{3}{x}-1\).  
$BR
$BR
(b) \{ $answer[1]->menu() \} \((5y-2x)y'-2y=0 \).
$BR
$BR
(c) \{ $answer[2]->menu() \} \( \left( 2y\sin(x)\cos(x)-y+2y^{2}e^{xy^{2}} \right) dx=(x-\sin^{2}(x)-4xye^{xy^{2}})dy\).

END_TEXT
Context()->normalStrings;


############################
#  Answers

$showPartialCorrectAnswers = 0;

foreach my $i (0..2) {
ANS( $answer[$i]->cmp(
  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( ($numOfAttempts < 2) && ($correct == $student) ) { 
       return 1; 
    } elsif ( ($numOfAttempts > 1) && ($correct == $student) ) {
       Value->Error("Your answer is correct, but your score did not change since you used up the attempts allowed."); 
    } else {
       return 0;
    };
  }
));

}


#ANS( $answer[0]->cmp() );
#ANS( $answer[1]->cmp() );
#ANS( $answer[2]->cmp() );


COMMENT('Uses a custom answer checker to 
limit the number of attempts that count for credit to one.');


ENDDOCUMENT();
