##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('PDE', 'wave', 'equation')

## DBsubject('Calculus')
## DBchapter('Partial Differential Equations')
## DBsection('Wave Equation')
## Date('7/25/2011')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('12.3')
## Problem1('')
########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "contextFraction.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   "PGcourse.pl",      # Customization file for the course
   "Alfredmacros.pl",
   "answerHints.pl",
   "parserMultiAnswer.pl",
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;
#force the use of jsmath mode for the inline entry answer boxes
jsmathmode();
##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("X"=>"Real","X'"=>"Real","X''"=>"Real","T"=>"Real","T''"=>"Real","L"=>"Real","lambda"=>"Real","a"=>"Real","n"=>"Real");
Context()->variables->set(lambda => {TeX => '{\lambda}'});
Context()->functions->undefine("acos","acosh");
Context()->{error}{msg}{"Function 'acosh' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cosh function, you need to put the multiplication symbol between the a and the cosh";
Context()->{error}{msg}{"Function 'acos' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cos function, you need to put the multiplication symbol between the a and the cos";
$DEcontext = Context()->copy;

Context("Numeric");
Context()->variables->add("y"=>"Real","t"=>"Real","L"=>"Real","n"=>"Real","alpha"=>"Real","a"=>"Real","A"=>"Real","B"=>"Real","C"=>"Real","D"=>"Real","n"=>"Real");
Context()->variables->set(alpha => {TeX => '{\alpha}'});
Context()->functions->undefine("acos","acosh");
Context()->{error}{msg}{"Function 'acosh' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cosh function, you need to put the multiplication symbol between the a and the cosh";
Context()->{error}{msg}{"Function 'acos' is not allowed in this context"} 
= "WeBWorK thinks you meant the inverse cos function, you need to put the multiplication symbol between the a and the cos";

$cnt = Formula("cos(n*pi*a*t/L)");
$snt = Formula("sin(n*pi*a*t/L)");
$cnx = Formula("cos(n*pi*x/L)");

## Sturm-Louiville problem

$waveleft = Formula($DEcontext,"a^2*X''*T");
$waveright = Formula($DEcontext,"X*T''");

$Xdiff1 = Formula($DEcontext,"X''/X")->reduce;
$Tdiff1 = Formula($DEcontext,"T''/(a^2*T)")->reduce;

$Xdiff2 = Formula($DEcontext,"X''+lambda*X")->reduce;
$Tdiff2 = Formula($DEcontext,"T''+lambda*a^2*T")->reduce;

## DE Initial Conditions
$left = Formula($DEcontext,"0");
$right = Formula($DEcontext,"L");
$homogeneousleft = Formula($DEcontext,"0");
$homogeneousright = Formula($DEcontext,"0");

$multians = MultiAnswer($left,$homogeneousleft,$right,$homogeneousright)->with(
  singleResult => 0,
  checker => sub 
    {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2,$stu3,$stu4 ) = @{$student};
      my ( $c1, $c2, $c3, $c4 ) = @{$correct};
      if ( ($c1 == $stu1) && ($c3 == $stu3) ) 
         {if ( ($c2 == $stu2) && ($c4 == $stu4) )
            {return [1,1,1,1];}
          elsif ($c2 == $stu2) {return [1,1,1,0]}
          elsif ($c4 == $stu4) {return [1,0,1,1]}
          else {return [1,0,1,0]};
         }
      elsif ( ($c3 == $stu1) && ($c1 == $stu3) )
         {if ( ($c4 == $stu2) && ($c2 == $stu4) )
             {return [1,1,1,1];}
         elsif ($c4 == $stu2) {return [1,1,1,0]}
         elsif ($c2 == $stu4) {return [1,0,1,1]}
         else {return [1,0,1,0]};
         }
      else {return [0,0,0,0];}
  }
);


$sn = Formula("sin(n*pi*x/L)");
$cn = Formula("cos(n*pi*x/L)");

$Xequal = Formula("A+Bx");
$Xequalprime = $Xequal->D('x');

$Xless = Formula("A*cosh(alpha*x)+B*sinh(alpha*x)");
$Xlessprime = $Xless->D('x');
$Xlessleft = $Xlessprime->substitute(x=>"0");
$Xsimplified = Formula("A*cosh(alpha*x)");
$Xlessright = Formula("A*alpha*sinh(alpha*L)");

$Xgreater =  Formula("A*cos(alpha*x)+B*sin(alpha*x)");
$Xgreaterprime = $Xgreater->D('x');
$Xgreaterleft = $Xgreaterprime->substitute(x=>"0");
$Xcos = Formula("A*cos(alpha*x)");
$Xgreaterright = $Xgreaterprime->substitute(x=>"L");

## Find T now
$Tdiff3 = Formula($DEcontext,"T''+(n*pi*a/L)^2*T")->reduce;

$Tgreater = Formula("C*cos(n*pi*a*t/L)+D*sin(n*pi*a*t/L)");
$ugreater = $cn*$Tgreater;

## Putting X and T together
$a0 = Formula("L/2");
$an = Formula("2*L^2(-1+cos(n*pi))/(a*n^3*pi^3)");
$bn = Formula("0");

$u0 = $a0;
$un = Formula("($an*$cnt+$bn*$snt)*$cnx");
##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
For the general solution of the differential equation in X use A and B for your constants and list the functions in alphabetical order, for example \(y = A\cos (x) + B \sin(x)\). For the differential equation in T use the  C and D.For the variable \(\lambda\) type the word lambda and type alpha for \(\alpha\),otherwise treat them as you would any other variable.
$BR
$BR
Use the prime notation for derivatives, so the derivative of \(X\) is written as \(X^\prime\). Do NOT use \(X^\prime(x)\)
$BR
$BR
The longitudinal displacement u(x,t) of a vibrating elastic bar can be modeled by a wave equation with free-end conditions
\[
\begin{aligned}
&a^2\frac{\partial^2 u}{\partial x^2} = \frac{\partial^2 u}{\partial t^2}, 0< x < L, t > 0\\
&\frac{\partial u}{\partial x}\Big\vert_{x=0} = 0,\hskip 10pt \frac{\partial u}{\partial x}\Big\vert_{x=L} = 0, t > 0\\
&u(x,0) = x
\hskip 10pt \frac{\partial u}{\partial t}\Big\vert_{t=0}=$a,0<x<L\\
\end{aligned}
\]
$BR
$BR
$BBOLD PART 1: Separate $EBOLD
$BR
$BR
We look for solutions of the form \(u(x,t) = X(x)T(t)\). 
$BR
Assuming solutions of this form the PDE can be written as 
$BR
$BCENTER \{$waveleft->ans_rule\} = \{$waveright->ans_rule\} $ECENTER
$BR
Then we can separate the PDE,placing constants in the DE into T: 
$BR
\{$Xdiff1->ans_rule\} = \{$Tdiff1->ans_rule\} = \(-\lambda\) (\(\lambda\) is typed as lambda)
$BR
$BR
The boundary conditions in the PDE translate into initial conditions for the differential equation in X:
$BR

Since these differential equations are independent of each other, they can be separated 
$BR
DE in X: \{$Xdiff2->ans_rule\} \( = 0\)
$BR IC's in X: 
$BR
$BCENTER
\(X^\prime(\)\{$multians->ans_rule(1)\}\() = \)\{$multians->ans_rule(1)\}
$BR
\(X^\prime(\)\{$multians->ans_rule(1)\}\() = \)\{$multians->ans_rule(1)\}
$ECENTER
$BR
DE in T: \{$Tdiff2->ans_rule\} \( = 0\)
$BR
$BR
This leads us to a Sturm-Louiville problem in X. In each case the general solution in X is written with constants A and B.
$BR
$BR
$BBOLD PART 2: Solve for X $EBOLD
$BR
$BR
$BBOLD Case 1: $EBOLD \(\lambda = 0\)
$BR
\(X(x) = \) \{$Xequal->ans_rule(70)\}
$BR
Using this formula
$BR
$BCENTER
\(X^\prime(0) = \)\{ans_rule(1)\} \( = 0\)
$BR
\(X^\prime(L) = \)\{ans_rule(1)\} \( = 0\)
$ECENTER
$BR
Making X a constant function. We get the constant solution another way, so we will discard this case.
$BR
$BR
$BBOLD Case 2: $EBOLD \(\lambda = -\alpha^2\)
$BR
\(X(x) = \) \{$Xless->ans_rule(70)\}
$BR
\(X^\prime(0) = \)\{ans_rule\} \( = 0\)
$BR
That simplifies X down to \{ans_rule\}. Using the simplified X:
$BR
\(X^\prime(L) = \)\{ans_rule(20)\} \( = 0\)
$BR
$BR
Therefore \(A = \) \{ans_rule\} and \(X = \) \{ans_rule\}
$BR
$BR
$BBOLD Case 3: $EBOLD \(\lambda = \alpha^2 \)
$BR
\(X(x) = \) \{$Xgreater->ans_rule(70)\}
$BR
\(X^\prime(0) = \)\{ans_rule\} \( = 0\)
$BR
So X simplifies to \{$Xcos->ans_rule\}.
$BR
Using the simplified X:
$BR
\(X^\prime(L) = \)\{ans_rule\} \( = 0\)
$BR
Which lead use to the eigenvalues \(\lambda = \)\{ans_rule\}
$BR 
$BBOLD Note $EBOLD If n = 0 we get the constant case from \(\lambda = 0\), including it would have been redundant.
$BR
$BR
$BR
$BBOLD PART 3: Solve for T $EBOLD
$BR
$BR
$BR
So the only case we need consider is the \(\lambda = \alpha^2\) case. Therefore the differential equation in T is now 
$BR
\{$Tdiff3->ans_rule\} \( = 0\)
$BR
which has the general solution (using C and D as the unknown constants)
$BR
\(T = \) \{$Tgreater->ans_rule(40)\}
$BR
$BR
$BBOLD PART 4: Find the series solution $EBOLD
$BR
$BR
Since the constant from the \($cnx\) term would be redundant, functions of the form \(u(x,t) = X(x)T(t) = \)\{ans_rule(50)\}, for \(n\geq 0\), are solutions of the PDE
\[
\begin{aligned}
&a^2\frac{\partial^2 u}{\partial x^2} = \frac{\partial^2 u}{\partial t^2}, 0< x < L, t > 0\\
&\frac{\partial u}{\partial x}\Big\vert_{x=0} = 0,\hskip 10pt \frac{\partial u}{\partial x}\Big\vert_{x=L} = 0, t > 0\\
\end{aligned}
\]
To add in the initial conditions we will have to use a series solution. For the series we relabel the coefficients so that \(C = A_n\) and \(D=B_n\) 
$BR
\(u(x,t) = \displaystyle\sum\limits_{n=0}^\infty \left( A_n \hbox{\{$cnt->ans_rule(13)\}} + B_n \hbox{\{$snt->ans_rule(13)\}}\right)\hbox{\{$cnx->ans_rule(12)\}}\)
$BR
Using the initial conditions we can compute
$BR
\(
\begin{aligned}
A_0&=\displaystyle\frac{\mbox{\{ans_rule(1)\}}}{\mbox{\{ans_rule(3)\}}}\int_0^L\mbox{\{ans_rule(1)\}}\hskip 3pt dx
&\\
&=\mbox{\{ans_rule(5)\}}\\
&\\
A_n&=\displaystyle\frac{\mbox{\{ans_rule(1)\}}}{\mbox{\{ans_rule(3)\}}}\int_0^L\mbox{\{ans_rule(30)\}}\hskip 3pt dx
&\\
&=\mbox{\{ans_rule(35)\}}\\
&\\
B_n&=\displaystyle\frac{\mbox{\{ans_rule(1)\}}}{\mbox{\{ans_rule(3)\}}}\int_0^L\mbox{\{ans_rule(5)\}}\hskip 3pt dx
&\\
&=\mbox{\{ans_rule(5)\}}\\
\end{aligned}
\)
$BR
$BR
$BBOLD PART 5: Finally! $EBOLD
$BR
$BR \(u(x,t) = \)\{$u0->ans_rule(4)\}+\(\displaystyle\sum\limits_{n=1}^\infty\)\{$un->ans_rule(80)\} 
$BR

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS($waveleft->cmp() );
ANS($waveright->cmp() );

#####
# Translate to ODE
#
ANS($Xdiff1->cmp);
ANS($Tdiff1->cmp);
ANS($Xdiff2->cmp);
ANS( $multians->cmp() ); 
ANS($Tdiff2->cmp);
## Sturm-Louiville
ANS($Xequal->cmp);
ANS(Formula("B")->cmp);
ANS(Formula("B")->cmp);

ANS($Xless->cmp);
ANS($Xlessleft->cmp);
ANS($Xsimplified->cmp);
ANS($Xlessright->cmp);
ANS(Compute("0")->cmp);
ANS(Compute("0")->cmp);


ANS($Xgreater->cmp);
ANS($Xgreaterleft->cmp);
ANS($Xcos->cmp);
ANS(Compute("-A*alpha*sin(alpha*x)")->cmp);
ANS(Compute("n*pi/L")->cmp);

#Solve for T
ANS($Tdiff3->cmp() );
ANS($Tgreater->cmp);

#Series solution
ANS($ugreater->cmp());

ANS($cnt->cmp() );
ANS($snt->cmp() );
ANS($cnx->cmp() );
# A_0
ANS(Compute("1")->cmp() );
ANS(Compute("L")->cmp() );
ANS(Formula("x")->cmp() );
ANS($a0->cmp() );
#A_n
ANS(Compute("2")->cmp() );
ANS(Compute("L")->cmp() );
ANS(Formula("x*$cnx")->cmp() );
ANS($an->cmp() );
#B_n
ANS(Compute("2")->cmp() );
ANS(Compute("n*pi*a")->cmp() );
ANS(Formula("0")->cmp() );
ANS($bn->cmp() );

ANS($u0->cmp);
ANS($un->cmp);
ENDDOCUMENT();        
