# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('5.2')
## Problem1('1')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Vector2D");

SRAND( $psvn );

$v11 = random(1,3,1);
$v12 = random(-3,-1,1);
$v21 = random(2,4,1);
$v22 = random(2,4,1);

$v1 = ColumnVector( $v11, $v12 );
$v2 = ColumnVector( $v21, $v22 );
if ( $v1.$v2 == 0 ) {
    $v21--;
    $v2 = ColumnVector( $v21, $v22 );
}

$u1 = unit( $v1 );
$u2 = unit( $v2 - ($u1 . $v2)*$u1 );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider the basis \( \vec v_1 = $v1 \), \( \vec v_2 = $v2 \)
of \(\mathbb R^2\). Apply the algorithm in the text (the Gram-Schmidt process) to this basis to find an orthonormal pair of basis vectors
\(\vec u_1\) and \(\vec u_2\).
$BR
\( \vec u_1 = \) \{ $u1->ans_array() \} $BR
\( \vec u_2 = \) \{ $u2->ans_array() \} $BR

END_TEXT
Context()->normalStrings;

ANS( $u1->cmp( checker => sub {
    my ( $c, $s, $a ) = @_;
    return ($c->isParallel($s) && norm($s) == 1); } ) );
ANS( $u2->cmp( checker => sub {
    my ( $c, $s, $a ) = @_;
    return ($u1.$s == 0 && norm($s) == 1); } ) );

ENDDOCUMENT();
