# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('5.5')
## Problem1('2')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Matrix");

$amat = Matrix( [ [ random(1,3,1), 0 ], [ random(-3,-1,1), random(1,3,1) ] ] );
$bmat = Matrix( [ [ random(1,3,1), random(1,3,1) ], [ 0, random(1,3,1) ] ] );
$amb = $amat - $bmat;
$d = sqrt( ($amb->transpose() * $amb)->trace() );
$norma = sqrt( ($amat->transpose() * $amat)->trace() );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider the inner product \(\langle A, B\rangle = \mbox{trace}(A^T B)\) on
\(\mathbb R^{2\times 2}\).
$BR
\(\textbf{a)}\) Find the norm of \(M=$amat\) with respect to this inner product.
$BR
Answer: \(||M||=\) \{$norma->ans_rule()\}
$BR
\(\textbf{b)}\) To find the distance between
\(M = $amat\) and \(N = $bmat\), we find the norm of what matrix?
$BR
Answer: \{$amb->ans_array()\}

END_TEXT
Context()->normalStrings;

ANS( $norma->cmp() );
ANS( $amb->cmp( checker=>sub {
    my ( $c, $s, $a ) = @_;
    return $c == $s || -1*$c == $s; } ) );

ENDDOCUMENT();
