# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('5.3')
## Problem1('2')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Vector2D");

$v1 = random(1,3,1);
$v2 = random(1,7,2);
$nsq = $v1*$v1 + $v2*$v2;
$n = "\sqrt{ $nsq }";

$col = ColumnVector( -$v2/sqrt($nsq), $v1/sqrt($nsq) );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find a pair of values \(a_{12}\) and \(a_{22}\) for which
\[ A = \begin{bmatrix} $v1/$n & a_{12}\\ $v2/$n & a_{22}\end{bmatrix} \]
is an orthogonal matrix.
$PAR
\(\begin{bmatrix} a_{12}\\ a_{22}\end{bmatrix} = \)
\{ $col->ans_array \}

END_TEXT
Context()->normalStrings;

ANS( $col->cmp( checker=>sub { my ( $c, $s, $a ) = @_;
			       return ($c->isParallel( $s ) && norm($s) == 1);
			   } ) );

ENDDOCUMENT();
