## DESCRIPTION
## Simplify Complex Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");

$a = non_zero_random(2, 7);
$c = non_zero_random(2, 7);

#Question
$g1 = Formula("(1)/(x)")->reduce;
$g2 = Formula("($a)/(x^2)")->reduce;
#$g3 = Formula("(1)/($c)")->reduce;
$g4 = Formula("($a**2)/($c*x^2)")->reduce;
$lcd = Formula("$c*x^2");

$num1 = Formula("$c*x - $a*$c")->reduce;
$num2 = Formula("$c(x-$a)")->reduce;
$den1 = Formula("x^2 - $a**2")->reduce;
$den2 = Formula("(x-$a)(x+$a)")->reduce;

#Answers
Context("RationalFunction-Strict");
$f = Formula("($c)/(x + $a)");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the complex rational expression.

[``\frac{[$g1]-[$g2]}{\frac{1}{[$c]}-[$g4]} = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
Multiply every term by the least common denominator.  The LCD is [`[$lcd]`].

[``\frac{[$g1]\cdot [$lcd]-[$g2]\cdot [$lcd]}{\frac{1}{[$c]}\cdot [$lcd]-[$g4]\cdot [$lcd]} ``]

[``\frac{[$num1]}{[$den1]}``]

[``\frac{[$num2]}{[$den2]}``]

[``[$f]``]

END_PGML_SOLUTION

ENDDOCUMENT();