## DESCRIPTION
##  Write equation defining image of T:R^2-->R^3 (plane)
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Kernel and image)
## Date(22/03/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(3)
## Static(0)
## MO(1)




##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
);


TEXT(beginproblem());

COMMENT('un-parameterize a plane in R^3');

#####################################
#  Setup

Context("Vector");
Context()->lists->set(Vector=>{open=>"(",close=>")"});

do {
 $v1 = Vector(random(-3,3,1),       # generating vectors
              random(-3,3,1),
              random(-3,3,1));
 $v2 = Vector(random(-3,3,1),
              random(-3,3,1),       # should be non-parallel
	      random(-3,3,1));      # and non-zero
} until (!(areParallel $v1 $v2) && !($v1->isZero()) && !($v2->isZero()));  

$perp = $v1 x $v2;                  # perpendicular direction 

@c  = ($perp->extract(1),$perp->extract(2),$perp->extract(3));
@c1 = ($v1->extract(1),  $v1->extract(2),  $v1->extract(3));
@c2 = ($v2->extract(1),  $v2->extract(2),  $v2->extract(3));

Context("Numeric");
Context()->variables->are(x=>'Real',y=>'Real',z=>'Real');

$plane = Formula("$c[0]*x+$c[1]*y+$c[2]*z")->reduce;
$f1    = Formula("$c1[0]*x+$c2[0]*y")->reduce;
$f2    = Formula("$c1[1]*x+$c2[1]*y")->reduce;
$f3    = Formula("$c1[2]*x+$c2[2]*y")->reduce;


#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Let \({T}:\mathbb{R}^2\to \mathbb{R}^3\) be the linear transformation 
defined by $BR
\(\displaystyle \quad {T}(x,y) = \bigl($f1,\,$f2,\,$f3\bigr)\) 
$PAR
Write a linear equation defining the subspace 
\(\mathrm{Im}({T})\)
$BR
\(\quad\)\{$plane->ans_rule(10) \} \( = 0\)
$PAR
(Write your answer in the form \(ax+by+cz\). 
For example "\(2x+3y-4z\)")

END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 0;

ANS( $plane->cmp(checker=>sub{
  my ($correct,$student,$ansHash) = @_;
  my $c = Formula($correct->value);
  my $s = Formula($student->value);    

  return 0 if ($s == Formula("0"));

  my @c_vect = ($c->D('x'),$c->D('y'),$c->D('z'));
  my @s_vect = ($s->D('x'),$s->D('y'),$s->D('z'));

  my $mult = ($c_vect[0] != 0) ? $s_vect[0]/$c_vect[0] : 
             ($c_vect[1] != 0) ? $s_vect[1]/$c_vect[1] :
	                         $s_vect[2]/$c_vect[2];

  return (($mult*$c_vect[0] == $s_vect[0]) && 
          ($mult*$c_vect[1] == $s_vect[1]) &&
          ($mult*$c_vect[2] == $s_vect[2]))
}) );


ENDDOCUMENT(); 
