# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector fields)
## DBsection(Graphs, flows lines, and level surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## Static(1)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.3')
## Section2('17.3')
## Section3('17.3')
## Problem1('21')
## Problem2('22')
## Problem3('23')
## KEYWORDS('vector field', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Vector2D");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);
$showPartialCorrectAnswers = 1;

$ans = Vector( Formula(1), Formula(1) );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Give an example of a vector field \( \vec F (x, y)  \) in
2-space, where \( \vec F \) is constant.

$PAR
\( \vec F(x,y) = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($ans->cmp( checker=>sub {
		   my ( $c, $s, $an ) = @_;
		   my @sv = $s->value;
		   return Compute(0) == $sv[0]->D('x') &&
		       Compute(0) == $sv[0]->D('y') &&
		       Compute(0) == $sv[1]->D('x') &&
		       Compute(0) == $sv[1]->D('y');
	       } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR


\(\vec F(x, y) = a\vec i + b\vec j\) for any real numbers \(a\) and
\(b\) is a constant vector field.  For example,
\(\vec F(x, y) = 3\vec i -4\vec j\).



END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
