# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector fields)
## DBsection(Graphs, flows lines, and level surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(writefield)
## Level(3)
## Static(1)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.3')
## Section2('17.3')
## Section3('17.3')
## Problem1('25')
## Problem2('26')
## Problem3('27')
## KEYWORDS('vector field', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context('Vector2D');
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);

$showPartialCorrectAnswers = 1;

$a = 1;
$b = -1;
$c = -1;
$v = 'x';
$t1 = Compute("x + $a*y")->reduce();
$t2 = Compute("$b + $c*$v^2")->reduce();

$v = Vector( $t1, $t2 );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Give an example of a vector field \( \vec F (x, y)  \) in 2-space
with the properties that \( \vec F \) is perpendicular to
\( \vec G = $v \) at every point.

$PAR
\( \vec F(x,y) = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($v->cmp( checker=>sub { my ( $c, $s, $an ) = @_;
    return $c.$s == Compute(0) && norm($s) != 0; } ) );


Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

If \(\vec F(x, y) = f(x,y)($h)\) where \(f(x, y)\) is any function,
then \(\vec F\cdot \vec G=0\), which shows that \(\vec F\) is
perpendicular to \(\vec G\).
$PAR
For example \(\vec F(x, y) = (y+\cos x  )($h)\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
