# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector fields)
## DBsection(Graphs, flows lines, and level surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(writefield)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.3')
## Section2('17.3')
## Section3('17.3')
## Problem1('29')
## Problem2('30')
## Problem3('32')
## KEYWORDS('vector field', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Vector2D");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);

$showPartialCorrectAnswers = 1;

$var = list_random('x','y');
$ij = ( $var eq 'x' ) ? 0 : 1;
$ji = ($ij + 1)%2;

$orient = list_random('vertical', 'horizontal');
$xy = ( $orient eq 'vertical' ) ? 'x' : 'y';
$yx = ( $xy eq 'x' ) ? 'y' : 'x';

@comp = ( Compute(0), Compute(0) );
$comp[$ij] = Compute( $xy );
$vec = Vector( @comp );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Write a formula for a two-dimensional vector field which has
all vectors parallel to the \( $var \)-axis and
all vectors on a $orient line having the same magnitude.

$PAR
\(\vec V = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($vec->cmp( checker=>sub {
		  my ( $c, $s, @ans ) = @_;
                  my $i = Vector(i);  
                  my $j = Vector(j);  
		  @sv = ($s . $i, $s . $j);  
		  return $sv[$ji] == 0 && $sv[$ij]->D($yx) == 0; 
	      } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

One possible solution is \(\vec{V}(x, y) = $vec\)

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
