# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector fields)
## DBsection(Graphs, flows lines, and level surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(writefield)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.3')
## Section2('17.3')
## Section3('17.3')
## Problem1('31')
## Problem2('32')
## Problem3('34')
## KEYWORDS('vector field', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context('Vector2D');
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);

$showPartialCorrectAnswers = 1;

$len = random(1,8,1);
$vec = Vector( Compute("$len*y/sqrt(x^2 + y^2)"),
	       Compute("-$len*x/sqrt(x^2 + y^2)") );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Write a formula for a two-dimensional vector field which has
all vectors of length $len and perpendicular to the position
vector at that point.

$PAR
\(\vec V = \) \{ ans_rule(45) \}

END_TEXT
Context()->normalStrings;

ANS($vec->cmp( checker=>sub {
		    my ( $c, $s, $ans ) = @_;
		    return $c == $s || -1*$c == $s;
		} ) );

$ld = ($len == 1) ? '' : $len;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The position vector at each point is
\(\vec{r}=x\,\vec i+y\,\vec j\). Any vector \(F\) perpendicular
to this must be \(\vec{F}(x,y)=A\,\vec i+B\,\vec j\) such that
\(\vec{F} \cdot \vec{r} = A\,x+B\,y=0\).  One possible answer
is to let \(A=y\) and \(B=-x\), so that \(\vec F = y\,\vec i - x\,\vec j\).
Then we want \(\vec V\) to have length $len, so we can take
\(\vec V = $len\,\vec F/|\!\!\!|\vec F|\!\!\!|\), so that
\[\vec V = \frac{$ld\,y\,\vec i - $ld x\,\vec j}{\sqrt{x^2+y^2}}.\]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
