# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/11/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector fields)
## DBsection(Graphs, flows lines, and level surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.3')
## Section2('17.3')
## Section3('17.3')
## Problem1('33')
## Problem2('33')
## Problem3('35')
## KEYWORDS('vector field', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context("Vector");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;


$fvec = Vector( Compute("x"),
		Compute("x + y"),
		Compute("x - y + z") );

$mx = non_zero_random(-5,5,1);
$my = non_zero_random(-5,5,1);
$mz = non_zero_random(-5,5,1);

$linevec = Vector( Compute($mx), Compute($my), Compute($mz) );

$vecContext = Context()->copy();

Context('Numeric');
Context()->variables->are( t=>'Real' );

$x0 = non_zero_random(-5,5,1);
$y0 = non_zero_random(-5,5,1);
$z0 = non_zero_random(-5,5,1);

$mx = non_zero_random(-5,5,1);
$my = non_zero_random(-5,5,1);
$mz = non_zero_random(-5,5,1);

$linex = Compute( "$x0 + $mx*t" )->reduce();
$liney = Compute( "$y0 + $my*t" )->reduce();
$linez = Compute( "$z0 + $mz*t" )->reduce();

$lineContext = Context()->copy();

Context('Point');

$parpoint = Point(Compute($mx), Compute($my - $mx),
		  Compute($mz - 2*$mx + $my));
$perpoint = Point(Compute(-1*$mz), Compute(0), Compute($mx+$my+$mz));

$ptContext = Context()->copy();

Context('ImplicitEquation');
Context()->variables->add( z=>'Real' );
$plane = ImplicitEquation( "$mx*x + $my*(x+y) + $mz*(x-y+z) = 0" );

$eqnContext = Context()->copy();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Let \( \vec F= $fvec \).  Let the line \(l\) be
\( x = $linex \),
\( y = $liney \),
\( z = $linez \).
$PAR
${BBOLD}(a)$EBOLD
Find a point \(P = (x_0, y_0, z_0)\) where
\( \vec F \) is parallel to \( l \).
$BR
\( P = \) \{ ans_rule(35) \}

$PAR
Find a point \(Q = (x_1,y_1,z_1)\) at which \( \vec F \) and \( l \)
are perpendicular.
$BR
\( Q = \) \{ ans_rule(35) \}

$PAR
Give an equation for the set of all points at
which \( \vec F \) and \( l \) are perpendicular.
$BR
equation: \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

Context($ptContext);
ANS( $parpoint->cmp( checker=>sub {
			 my ( $c, $s, $ans ) = @_;
			 my @cv = $c->value;
			 my @sv = $s->value;
			 return 0 if ( Compute(0) == $sv[0] );
			 my $c0 = $cv[0]/$sv[0];
			 return $c0*$cv[1] == $sv[1] && $c0*$cv[2] == $sv[2];
		     } ) );
ANS( $perpoint->cmp( checker=>sub {
			 my ( $c, $s, $ans ) = @_;
			 my @sv = $s->value;
			 return Compute($mx)*$sv[0] +
			     Compute($my)*($sv[0] + $sv[1]) +
			     Compute($mz)*($sv[0] - $sv[1] + $sv[2]) == 0;
		     }, showCoordinateHints=>0 ) );

Context( $eqnContext );
ANS( $plane->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The line \(l\) is parallel to the vector
\(\vec v = $linevec \).
The vector field \(\vec F\) is parallel to the line when \(\vec F\)
is a multiple of \(\vec v\).  Taking the multiple to be 1 and solving
for \(x,y,z\), we find the point \(P = $parpoint\).  Other answers
are possible.

$PAR
${BBOLD}(b)$EBOLD
The line and vector field are perpendicular if \(\vec F\cdot\vec v=0\),
that is
\[ (x)($mx) + (x + y)($my) + (x - y + z)($mz) = 0. \]
Solving, we can find any number of possibilities.  One of these is
\( Q = $perpoint\).

$PAR
$PAR
${BBOLD}(c)$EBOLD
The set of all points that is perpendicular to the line is the equation
shown in (b), \($plane\).  This is a plane through  the origin.


END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
