## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear combinations)
## Institution(TCNJ)
## Level(2)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('2.2')
## Problem1('')
## KEYWORDS('vectors' 'span')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
"PGstandard.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGgraphmacros.pl",
"PGnumericalmacros.pl",
"PGauxiliaryFunctions.pl",
"PGmatrixmacros.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$v01 = non_zero_random(-5,5,1);
$v02 = random(-5,5,1);
$v03 = non_zero_random(-5,5,1);

$mult = non_zero_random(-2,2,1);
$d1 = non_zero_random(-3,3,1);
$d2 = non_zero_random(-3,3,1);

$v11 = $mult*$v01;
$v12 = $mult*$v02+$d1;
$v13 = $mult*$v03+$d2;

$spot = 0;
@ans = (0,0,0,0);
Context("Vector");

$vecu = Matrix->new_column_matrix([$v01, $v02, $v03]);
$vecv = Matrix->new_column_matrix([$v11, $v12, $v13]);
$vu = Vector([$v01, $v02, $v03]);
$vv = Vector([$v11, $v12, $v13]);
$an1 = $vu+$vv;
$an2 = 2*$vu - $vv;

#$mp = MultiAnswer(1,1, $an1, 2, -1, $an2)->with(
#  singleResult=>1,
#  checker=> sub {
#    my ($correct,$student,$self) = @_;
	#my ($a1,$b1,$v1,$a2,$b2,$v2) = @{$student};
	#if(Real($a1)->isZero() or Real($b1)->isZero() or Real($a2)->isZero() or Real($b2)->isZero()) {
	#  $mp->setMessage(1,"The coefficients cannot be zero");
	#  return 0;
    #}
#    return 1;
#  }
#);

$mp=MultiAnswer(Real(1),Real(1),$an1,Real(2),Real(-1),$an2)->with(
  checker=> sub {
    my ($correct,$student,$self) = @_;
	my ($a1,$b1,$v1,$a2,$b2,$v2) = @{$student};
	my $res = [0,0,0,0,0,0];
	$self->setMessage(1,"The coefficients cannot be zero") if Real($a1)->isZero;
	$self->setMessage(2,"The coefficients cannot be zero") if Real($b1)->isZero;
	$self->setMessage(4,"The coefficients cannot be zero") if Real($a2)->isZero;
	$self->setMessage(5,"The coefficients cannot be zero") if Real($b2)->isZero;
	if(not (Real($a1)->isZero or Real($b1)->isZero)) {
		if($a1*$vu+$b1*$vv == Compute($v1) ) {
		  $res->[0] = 1;
		  $res->[1] = 1;
		  $res->[2] = 1;
		}
	}
	return res if(Real($a1)->isZero or Real($b1)->isZero);

	if(Real($a1) == Real($a2) and Real($b1) == Real($b2)) {
	  $self->setMessage(4, "The two linear combinations must be different");
	  return $res;
	}
	if($a2*$vu+$b2*$vv == Compute($v2) ) {
	  $res->[3] = 1;
	  $res->[4] = 1;
	  $res->[5] = 1;
	}
    return $res;
  }
);

$NO_SPACE = '@{}';

BEGIN_TEXT

\{ mbox( 'Let \( u = \)', display_matrix($vecu), ' and \( v = \)', display_matrix($vecv) ) \}
$BR
Find two different vectors in \( \textrm{span}(\left\{"\{"\} u, v\right\}) \) 
that are not multiples of \(u\) or \(v\) and show the weights on \( u \) and 
\( v \) used to generate them.
$BR
$BR
 \{$mp->ans_rule(5)\} \(u+\) \{$mp->ans_rule(5)\} \(v\)
\(=\) \{$mp->ans_rule(15)\}
$BR
$BR
 \{$mp->ans_rule(5)\} \(u+\) \{$mp->ans_rule(5)\} \(v\)
\(=\) \{$mp->ans_rule(15)\}
$BR$BR
$BBOLD Note:$EBOLD enter vectors using WeBWorK's 
\{helpLink('vector')\} notation.
END_TEXT

ANS($mp->cmp);


ENDDOCUMENT();       # This should be the last executable line in the problem.
