## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear combinations)
## Institution(TCNJ)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('2.2')
## Problem1('')
## KEYWORDS('vectors' 'span')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Vector");

$showPartialCorrectAnswers = 1;

$v01 = non_zero_random(-5,5,1);
$v02 = random(-5,5,1);
$v03 = non_zero_random(-5,5,1);

$mult = non_zero_random(-2,2,1);
$d1 = non_zero_random(-3,3,1);
$d2 = non_zero_random(-3,3,1);

$vecu = ColumnVector([$v01, $v02, $v03]);
$vecv = $mult*$vecu + (ColumnVector([0,$d1,$d2]));

$mp=MultiAnswer(1,1,$vecu + $vecv,1,-1,$vecu - $vecv)->with(
  checker=> sub {
    my ($correct,$student,$self) = @_;
    my ($a1,$b1,$v1,$a2,$b2,$v2) = @{$student};
    my $res = [1,1,1,1,1,1];  # initialize, optimistically assuming all answers are right
    my $badSumQ = 0;
    my $zerosQ = 0;  #zeroes in coefficients 
    
    # check the sums.  
    if($a1*$vecu+$b1*$vecv != $v1 ) {
      $badSumQ = 1;
      $res->[0] = 0;  
      $res->[1] = 0;  
      $res->[2] = 0;  
      $self->setMessage(3,"Check your vector arithmetic."); 
    }
    if($a2*$vecu+$b2*$vecv != $v2 ) {
      $badSumQ = 1;
      $res->[3] = 0;  
      $res->[4] = 0;  
      $res->[5] = 0;  
      $self->setMessage(6,"Check your vector arithmetic."); 
    }
    if (($res->[2])==0 or ($res->[5])==0){ return $res; }   #Quit if a sum is bad.

    #Check the coefficients.  Quit if one is zero.
    if ( ($a1->isZero) or ($b1->isZero)){
      $res->[0] = 0;  
      $res->[1] = 0;  
      $res->[2] = 0;  
      $self->setMessage(3,"Is multiple of u or v. Coefficients can't be zero.");
    }
    if ( ($a2->isZero) or ($b2->isZero)){
      $res->[3] = 0;  
      $res->[4] = 0;  
      $res->[5] = 0;  
      $self->setMessage(6,"Is multiple of u or v. Coefficients can't be zero.");
    }
    if (($res->[2])==0 or ($res->[5])==0){ return $res; }   #Quit if a coefficient is zero. 
      
    # if we get this far all values of $res are 1
    #check that the vectors are different
    if ($v1 == $v2){ 
      $res->[2] = 0;
      $res->[3] = 0;  
      $res->[4] = 0;  
      $res->[5] = 0;  
      $self->setMessage(3,"Different vectors? The two linear combinations must be different.");
      $self->setMessage(6,"Different vectors? The two linear combinations must be different.");
    }
    return $res;
  }
);


TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Let \( u = $vecu\) and \( v = $vecv\). 
$BR
Find two different vectors in span\(($LBRACE u, v $RBRACE) \) 
that are not multiples of \(u\) or \(v\) and show the weights on \( u \) and 
\( v \) used to generate them.
$PAR
 \{$mp->ans_rule(5)\} \(u+\) \{$mp->ans_rule(5)\} \(v\)
\(=\) \{$mp->ans_rule(15)\}
$PAR
 \{$mp->ans_rule(5)\} \(u+\) \{$mp->ans_rule(5)\} \(v\)
\(=\) \{$mp->ans_rule(15)\}
$PAR
$BBOLD Note:$EBOLD enter vectors using WeBWorK's 
\{helpLink('vector')\} notation.
END_TEXT

Context()->normalStrings;
ANS($mp->cmp);


ENDDOCUMENT();       # This should be the last executable line in the problem.
