## DESCRIPTION
## Parametric equations: parametric curve in space
## ENDDESCRIPTION

## DBsubject(WeBWorK)
## DBchapter(WeBWorK tutorial)
## DBsection(PGML tutorial 2015)
## Date(06/01/2015)
## Institution(Hope College)
## Author(Paul Pearson)
## MO(1)
## KEYWORDS('parametric', 'curve in space')


#################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "parserMultiAnswer.pl",
  "PGML.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


#################################
#  Setup

Context("Vector2D");
#Context("Vector"); # for 3D vectors
Context()->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[0,5]});
Context()->flags->set( ijk=>0 );

$a = random(2,5,1);
$Q = Point($a,$a**2);


$multians = MultiAnswer(Vector("<t,t**2>"),0,$a)->with(
  singleResult => 1,

  checker => sub {

  my ($correct,$student,$self) = @_;  # get the parameters
  my ($f,$x1,$x2) = @{$student};      # extract student answers
  if ( 
       ( ($f . i)**2 == ($f . j)  )
       && ($f->eval(t=>$x1) == Vector("<0,0>")) 
       && ($f->eval(t=>$x2) == Vector("<$a,$a**2>")) 
     ) { 
       return 1; 
     } elsif (
       ( ($f . i)**2 == ($f . j)  )
       && ($f->eval(t=>$x1) == Vector("<0,0>")) 
     ) {
       $self->setMessage(3,"Your right endpoint is not correct."); 
       return 0; 
     } elsif (
       ( ($f . i)**2 == ($f . j)  )
       && ($f->eval(t=>$x2) == Vector("<$a,$a**2>")) 
     ) {
       $self->setMessage(2,"Your left endpoint is not correct."); 
       return 0; 
     } elsif (
       ( ($f . i)**2 == ($f . j)  )
     ) {
       $self->setMessage(2,"Your left endpoint is not correct."); 
       $self->setMessage(3,"Your right endpoint is not correct."); 
       return 0;
     } else { return 0; }


}
);




#####################################
#  Main text

BEGIN_PGML
Find a vector parametric equation for the parabola
[` y = x^2 `] from the origin to the point 
[` [$Q] `] using [` t `] as a parameter.

[` \vec{r}(t) = `]
[_____________]{$multians} for
[___]{$multians}
[` \leq t \leq `] 
[___]{$multians}.
END_PGML


############################
#  Solution

BEGIN_PGML_SOLUTION
Solution explanation goes here.
END_PGML_SOLUTION

COMMENT('Uses PGML.');

ENDDOCUMENT(); 
