## DESCRIPTION
## Compute the eigenvalues of a 3x3 matrix.
##   All coefficients and eigenvalues will be small real integers. There will be at least 2 distinct eigenvalues, and if an eigenvalue is repeated (algebraic multiplicity > 1) then the student is expected to list it multiple times.
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(March 22, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('3.3')
## Problem1('')
## KEYWORDS('linear algebra','matrix','eigenvalue','algebraic multiplicity')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

$flg = random(1,5);
if($flg == 1) {
    $P = Matrix([[2,5,-2],[0,3,-1],[-1,-2,1]]);
    $PI = Matrix([[1,-1,1],[1,0,2],[3,-1,6]]);
} elsif($flg == 2) {
    $P = Matrix([[3,0,1],[1,1,1],[2,0,1]]);
    $PI = Matrix([[1,0,-1],[1,1,-2],[-2,0,3]]);
} elsif($flg == 3) {
    $P = Matrix([[1,-2,2],[-2,5,-4],[-1,2,-1]]);
    $PI = Matrix([[3,2,-2],[2,1,0],[1,0,1]]);
} elsif($flg == 4) {
    $P = Matrix([[1,1,-1],[0,1,-1],[2,0,1]]);
    $PI = Matrix([[1,-1,0],[-2,3,1],[-2,2,1]]);
} else {
    $P = Matrix([[2,1,1],[3,2,1],[1,1,1]]);
    $PI = Matrix([[1,0,-1],[-2,1,1],[1,-1,1]]);
}

do {
    $eig1 = random(-3,4);
    $eig2 = random(-3,4);
    $eig3 = random(-3,4);
    
    $D = Matrix([[$eig1,0,0],[0,$eig2,0],[0,0,$eig3]]);
    
    $A = $P * $D * $PI;
} until ($A->norm_one < 25 && ($eig1 != $eig2 || $eig1 != $eig3));

$multians = MultiAnswer($eig1,$eig2,$eig3)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1, $s2 ) = @{$student};
        if(($s0 == $eig1 && $s1 == $eig2 && $s2 == $eig3) || ($s0 == $eig1 && $s2 == $eig2 && $s1 == $eig3) || ($s1 == $eig1 && $s0 == $eig2 && $s2 == $eig3) || ($s1 == $eig1 && $s2 == $eig2 && $s0 == $eig3) || ($s2 == $eig1 && $s1 == $eig2 && $s0 == $eig3) || ($s2 == $eig1 && $s0 == $eig2 && $s1 == $eig3)) {
            return 1;
        } else {
            return 0;
        }
    }
);

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Compute all \( 3 \) eigenvalues of the matrix \( \displaystyle $A \). If an eigenvalue has algebraic multiplicity greater than \( 1 \), list it multiple times.
$BR
$BR
Eigenvalues: \{$multians->ans_rule(3)\}, \{$multians->ans_rule(3)\}, \{$multians->ans_rule(3)\}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
