## DESCRIPTION
## Compute an eigenvector corresponding to a given eigenvalue of a 3x3 matrix.
##    All matrix and vector entries will be small integers, as will the eigenvalue.
##    All eigenspaces will be 1-dimensional, so eigenvectors are only determined up to scalar multiple.
##    All eigenvectors corresponding to the given eigenvalue are accepted. The scaling of the eigenvector does not matter (as long as it is non-zero).
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(March 13, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('3.3')
## Problem1('')
## KEYWORDS('linear algebra','matrix','eigenvalue','eigenvector')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context('Matrix');

$flg = random(1,5);
if($flg == 1) {
    $P = Matrix([[2,5,-2],[0,3,-1],[-1,-2,1]]);
    $PI = Matrix([[1,-1,1],[1,0,2],[3,-1,6]]);
} elsif($flg == 2) {
    $P = Matrix([[3,0,1],[1,1,1],[2,0,1]]);
    $PI = Matrix([[1,0,-1],[1,1,-2],[-2,0,3]]);
} elsif($flg == 3) {
    $P = Matrix([[1,-2,2],[-2,5,-4],[-1,2,-1]]);
    $PI = Matrix([[3,2,-2],[2,1,0],[1,0,1]]);
} elsif($flg == 4) {
    $P = Matrix([[1,1,-1],[0,1,-1],[2,0,1]]);
    $PI = Matrix([[1,-1,0],[-2,3,1],[-2,2,1]]);
} else {
    $P = Matrix([[2,1,1],[3,2,1],[1,1,1]]);
    $PI = Matrix([[1,0,-1],[-2,1,1],[1,-1,1]]);
}

do {
    $eig1 = random(-5,5);
    $eig2 = random(-5,5);
    $eig3 = random(-5,5);
    
    $D = Matrix([[$eig1,0,0],[0,$eig2,0],[0,0,$eig3]]);
    
    $A = $P * $D * $PI;
} until ($A->norm_one < 25 && $eig1 != $eig2 && $eig1 != $eig3 && $eig2 != $eig3);


Context("Vector");
Context()->lists->set(Vector => {open => "(", close => ")"});
Context()->parens->set("(" => {type => "Vector", close => ")"});

$flg = random(1,3);
if($flg == 1) {
    $eg = $eig1;
} elsif($flg == 2) {
    $eg = $eig2;
} else {
    $eg = $eig3;
}
$ans = Vector($P->column($flg));

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Compute an eigenvector of the matrix \( \displaystyle $A \) corresponding to the eigenvalue \( $eg \).
$BR
$BR
Eigenvector: \{ $ans->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $ans->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    
    if( norm($student) > 0.0001 && $student->isParallel($ans) == 1 )
        { return 1; } else { return 0; }
}));

ENDDOCUMENT();
