##DESCRIPTION
## Word problem
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Linear equations)
## Institution(The College of Idaho)
## Author(RA Cruz)
## Level(2)
## MO(1)
## MLT(MultipleVarsSolveForOne)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('1.5')
## Problem1('')
## KEYWORDS('mathematical models','problem solving','formulas','algebra')
## Date: 2007/09

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
######################################
#  Setup
@vars = ("x","y","z","p","q","k","u","v","h","s");
@caps = ("C","N","M","R","L","J","F");
$c1 = random(0,scalar(@caps)-1,1);
$c2 = random(0,scalar(@vars)-1,1);
do {$c3 = random(0,scalar(@vars)-1,1)} until ($c3 != $c2);
do {$c4 = random(0,scalar(@vars)-1,1)} until ($c4 != $c2 && $c4 != $c3);
$var1 = $caps[$c1];
$var2 = $vars[$c2];
$var3 = $vars[$c3];
$var4 = $vars[$c4];
Context()->variables->are(lc($var1)=>"Real", lc($var2)=>"Real", lc($var3)=>"Real", lc($var4)=>"Real", uc($var1)=>"Real", uc($var2)=>"Real", uc($var3)=>"Real", uc($var4)=>"Real", );

for ($i = 0; $i < random(10,50); $i++) {
$n = random(0,2,1);
$a = non_zero_random(-15,15);
do {$b = random(2,15)} until (gcd($a,$b) == 1);
}
Context()->texStrings;
$dispeqn = "$a $var1 = $b $var2 $var3 + $b $var2 $var4 + $b $var3 $var4";
@nums = ("$a $var1 - $b $var3 $var4", "$a $var1 - $b $var2 $var4", "$a $var1 - $b $var2 $var3");
@dens = ("$b $var3 + $b $var4", "$b $var2 + $b $var4", "$b $var2 + $b $var3");
$num = Formula($nums[$n]);
$den = Formula($dens[$n]);
$dispans = "\displaystyle \frac{$num}{$den}";
@var = ("$var2", "$var3", "$var4");
$frac = $var[$n];



$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
#  $displayfrac = "\( $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \)";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Solve the equation for \($frac \):
$PAR
\[ $dispeqn \]
$PAR
$displayfrac
END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp(studentsMustReduceFractions=>1));

Context()-texStrings;
    BEGIN_TEXT
    $PAR
    $BBOLD Note: $EBOLD
    Write your answer as a single fraction, and type the numerator and denominator separately in the answer blanks provided.
    $PAR
    END_TEXT

Context()->normalStrings;

ENDDOCUMENT();
