##DESCRIPTION
## Word problem
## 
##ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Linear equations and functions)
## DBsection(Linear equations)
## Institution(The College of Idaho)
## Author(RA Cruz)
## Level(2)
## MO(1)
## MLT(MultipleVarsSolveForOne)
## TitleText1('Essentials of Intermediate Algebra')
## AuthorText1('Blitzer')
## EditionText1('1')
## Section1('1.5')
## Problem1('')
## KEYWORDS('mathematical models','problem solving','formulas','algebra')
## Date: 2007/09

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
######################################
#  Setup

Context()->variables->are(A=>'Real',
                          b=>'Real',
                          h=>'Real', a=>'Real', B=>'Real', H=>'Real');

$n = random(0,1,1);
@vars = ("b", "h");
@answer = (Formula("2A/h"), Formula("2A/b"));


$var = $vars[$n];
$a = random(2,10);
$b = random(2,10);
@nums = ("$a (A - $b h)", "$a A");
$e1 = $a * $b;
@dens = ("h", "b + $e1");
$num = Formula($nums[$n])->reduce;
$den = Formula($dens[$n])->reduce;
Context()->texStrings;
$dispeqn = "A = \frac{1}{$a} b h + $b h";
$dispans = "$var = \frac{$num}{$den}";

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $var = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $var = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

######################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Solve the equation for \( $var\): 
$PAR
\[ 
$dispeqn 
\]
$PAR
$displayfrac
END_TEXT

######################################
#  Answers

ANS($multians->cmp);

$showPartialCorrectAnswers = 1;

ENDDOCUMENT();
