# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('vectors', 'dot product')
## Tagged by glr 06/12/09


## DBsubject('Calculus')
## DBchapter('Vectors and the Geometry of Space')
## DBsection('The Dot Product')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('13.3')
## Problem1('13')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')
## HHSection1('The Dot Product')

Context("Vector");
Context()->flags->set( ijk=>1 );

Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(2,8,1);
$c = random(2,8,1);
$d = random(2,8,1);

$n = Vector( (-$a, $c, 1) );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a normal vector \(\vec n\) to the plane
\( z - $a (x - $b) = $c($d - y) \)

$PAR
\(\vec n = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( $n->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student->isParallel($correct);
  } ) );

$ab = $a * $b;
$cd = $c * $d;
$cdmab = $cd - $ab;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The equation can be rewritten as
\[ z - $a x + $ab = $cd - $c y, \]
or
\[ -$a x + $c y + z = $cdmab. \]
so \(\vec n=-$a\vec i + $c\vec j + \vec k\) (or any
vector parallel to that, of course).

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
