##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Algebra')
## DBchapter('Exponential and logarithmic expressions and functions')
## DBsection('Exponential and logarithmic equations')
## Date('7/15/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')
## TitleText1('Precalculus')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "bizarroArithmetic.pl",
   "parserMultiAnswer.pl",
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->variables->are(b=>Real, x=>Real);
Context()->flags->set(reduceConstants=>0, reduceConstantFunctions=>0,tolType=>'absolute');
Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
'*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
'/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
'//' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);

$coeff = random(3,9,1);
$base = random(2,8,1);
do{ $exp = random(12,58,1); } until ( gcf($coeff,$exp)==1 );
$result = Formula("x");

$ans = MultiAnswer(Formula("$base^($exp/$coeff)"),Formula("$result"))->with(
   singleResult => 1,
   separator => " = ",
   tex_separator => " = ",
   checker => sub {
      my ($correct,$student,$self) = @_;  # get the parameters
      my ($ansHash1, $ansHash2) = @{$self->{ans}};
      my $lhs = $ansHash1->{student_formula};
      my $rhs = $ansHash2->{student_formula};
      my ($exp,$res) = @{$correct};
      $lhs = Formula("$lhs"); $rhs = Formula("$rhs"); $exp = Formula("$exp");
      Context()->flags->set(bizarroDiv=>0, bizarroMul=>0, bizarroAdd=>0, bizarroSub=>0, bizarroPow=>0 );
      delete $exp->{test_values}, $lhs->{test_values}, $rhs->{test_values};
      return 0 unless (($lhs == $exp) && ($rhs == $res)) || (($rhs == $exp) && ($lhs == $res));
      Context()->flags->set(bizarroDiv=>1, bizarroMul=>1, bizarroAdd=>1, bizarroSub=>1, bizarroPow=>1 );
      delete $exp->{test_values}, $lhs->{test_values}, $rhs->{test_values};
      Value::Error("Your equation is not written in exponential form.") unless ((($lhs == $exp || $exp == $lhs) && ($rhs == $res || $res == $rhs)) || (($lhs == $res || $res == $lhs) && ($rhs == $exp || $exp == $rhs)));
      Context()->flags->set(bizarroDiv=>0, bizarroMul=>0, bizarroAdd=>0, bizarroSub=>0, bizarroPow=>0 );
      delete $exp->{test_values}, $lhs->{test_values}, $rhs->{test_values};
      return 1;
       },
  );


##############################################################
#
#  Text
#
#

BEGIN_PGML

>> ### Logarithmic Equations ### <<

In general, the logarithmic equation [`\color{blue}{\log_{B}(C)=A}`] is equivalent to the  exponential equation [`\color{green}{C=B^A}`].


How do we handle rewriting logarithmic equations as exponential ones if we have a coefficient in front of the logarithm?

For example, what about [`\color{red}{2}\color{blue}{\log_{3}(x)=31}`]?

Since the given equation is not yet in the form [`\color{blue}{\log_{B}(C)=A}`] because of the coefficient [`\color{red}{2}`], we first need to ''remove'' the coefficient. 

So, we divide by [`\color{red}{2}`] on _both_ sides of the equation to get [`\color{blue}{\log_{3}(x)=\dfrac{31}{2}}`].

Now, [`\color{blue}{\log_{3}(x)=\dfrac{31}{2}}`] is exactly in the form [`\color{blue}{\log_{B}(C)=A}`] and thus can be rewritten as [`\color{green}{x=2^{31/2}}`]

>> ### Practice ### <<

Rewrite [` [$coeff] \log_{[$base]}([$result]) = [$exp] `] as an exponential equation:

[_____________]{$ans} = [_____________]{$ans}

END_PGML

##############################################################
#
#  Answers
#
#

##############################################################
#
#  Hints
#
#



BEGIN_PGML_HINT

* Does your equation contain an expression that looks like [`\color{brown}{k}\log_{\color{red}{b}}\color{blue}{c}`]? If [`\color{brown}{k}`] is not [`1`], isolate [`\log_{\color{red}{b}}\color{blue}{c}`] by "removing" [`\color{brown}{k}`].

* [`\log_{\color{red}{b}}\color{blue}{c}=\color{green}{a}`] means that [`\color{red}{b}^{\color{green}{a}}=\color{blue}{c}`].

END_PGML_HINT

##############################################################
#
#  Solution
#
#



BEGIN_PGML_SOLUTION

Dividing both sides of [` \color{brown}{[$coeff]} \log_{\color{red}{[$base]}}(\color{blue}{[$result]}) = \color{green}{[$exp]} `] by [` \color{brown}{[$coeff]}`] yields that [`\log_{\color{red}{[$base]}}(\color{blue}{[$result]}) = \dfrac{\color{green}{[$exp]}}{\color{brown}{[$coeff]}}`]. This means that
[`\color{blue}{[$result]} = \color{red}{[$base]}^{\color{green}{[$exp]}/\color{brown}{[$coeff]}} `].

END_PGML_SOLUTION




ENDDOCUMENT();        
