# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus', 'integral', 'antiderivatives', 'indefinite integrals', 'partial fractions')
## Tagged by glr 02/08/09

## DBsubject('Calculus')
## DBchapter('Techniques of Integration')
## DBsection('Trigonometric Substitution')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('7.4')
## Section2('7.4')
## Problem1('4')
## Problem2('4')

## Textbook tags
## HHChapter1('Integration')
## HHChapter2('Integration')
## HHSection1('Algebraic Identities and Trigonometric Substitutions')
## HHSection2('Algebraic Identities and Trigonometric Substitutions')



DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
"parserMultiAnswer.pl",
"parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add(s=>"Real");

SRAND($psvn);

$a = random(2,8,1);
$fac1 = Formula("1/s");
$fac2 = Formula("-1/(s + $a)");
$ap1 = $a + 1;
$func = Formula("$a(1 + s)/(s(s^2 + $ap1 s + $a))");

$mp = MultiAnswer($fac1, $fac2)->with(
    singleResult => 0,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	my ( $f1stu, $f2stu ) = @{$student};
	my ( $f1, $f2 ) = @{$correct};
	if ( ( $f1 == $f1stu && $f2 == $f2stu ) ||
	     ( $f1 == $f2stu && $f2 == $f1stu ) ) {
	    return [1,1];
	} else {
	    if ( $f1 == $f1stu || $f2 == $f1stu ) {
		return [1,0];
	    } elsif ( $f1 == $f1stu || $f2 == $f2stu ) {
		return [0,1];
	    } else {
		return [0,0];
	    }
	}
    }
);
 
TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Split into partial fractions: 
$BR
\( \displaystyle $func = \) \{ $mp->ans_rule(15) \} + \{ $mp->ans_rule(15) \}

END_TEXT
Context()->normalStrings;

ANS($mp->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Since \( s^2 + $ap1 s+ $a =(s+$a)(s+1)\), we have
\[
\frac{$a(1+s)}{s(s^2+$ap1 s+$a)}=
\frac{$a(1+s)}{s(s+$a)(s+1)}=\frac{$a}{s(s+$a)},
\]
so we take
\[ \frac{$a}{s(s+$a)}= \frac{A}{s}  +\frac{B}{s+$a}.\]
Thus,
\[
$a = A(s+$a)+Bs,
\]
or
\[
$a = (A+B)s +$a A.
\]
This gives
\(A+B = 0\) and \($a A = $a\), so \(A = 1\) and \(B = -1\), and
\[
$func = $fac1 + $fac2.
\]

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
