## DESCRIPTION
## College Algebra
## ENDDESCRIPTION
 
 
## DBsubject(Algebra)
## DBchapter(Exponential and logarithmic expressions and functions)
## DBsection(Exponential and logarithmic equations)
## Date(04/21/2019)
## Institution(Community College of Denver)
## Author(Xian Feng Liang)
## MO(1)
## KEYWORDS('college algebra')
 
### Currently this problem DOES NOT require the fraction to be simplified
###########################
#  Initialization
 
DOCUMENT();
 
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"contextLimitedPoint.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
);
 
TEXT(beginproblem());

 
###########################
#  Setup
Context("Numeric");
install_problem_grader(~~&std_problem_grader);
 
# Array of 11 letters to be grabbed
@letter=qw/a b g h k r s t v w z/;


# Two random integer to grab letters
$n1 = random(0,10,1);
do{$n2 = random(0,10,1)} until ($n2 != $n1);


$root = $letter[$n1];
$power = $letter[$n2];


Context()->variables->add($root => "Real");
Context()->variables->add($power => "Real");
Context()->variables->add(y => "Real");


# The two sides of the correct answer
$ans1=Formula("log(y)");
$ans2=Formula("$power/$root");

# Allows for the correct equation to be given in either order

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2 ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $stu1 && $f2 == $stu2) ||
           ($f1 == $stu2 && $f2 == $stu1) ) {
          return [1,1];
      } else {
          if ($f1 == $stu1 || $f2 == $stu1) {
              return [1,0];
          } elsif ($f1 == $stu2 || $f2 == $stu2) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
 
 
###########################
#  Main text
 
BEGIN_PGML


Rewrite the following equation in logarithmic form. 

[`\sqrt[^[$root]]{10^[$power]}= y`] 

 
 [____]{$multians} [`=`] [____]{$multians} [@ AnswerFormatHelp("logarithms") @]*
 
 
END_PGML
 
 
 
############################
#  Solution
 
#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION
 
COMMENT('MathObject version. Uses PGML.');
 
ENDDOCUMENT();