# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Douglas Quinney, <d.a.quinney@maths.keele.ac.uk>
# revised by Gavin LaRose, <glarose@umich.edu> 11/2008

## Tagged by glr 11/19/08

## Author('Gavin LaRose');
## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Conceptual understanding of derivatives)
## Institution(University of Michigan)
## MLT(AppsEstimateDeriv)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('2.1')
## Section2('2.1')
## Section3('2.1')
## Section4('2.1')
## Problem1('17')
## Problem2('23')
## Problem3('22')
## Problem4('25')
## KEYWORDS('calculus','derivative','rate of change','velocity','average velocity','instantaneous velocity')

## Textbook tags
## HHChapter1('Key Concept: The Derivative')
## HHChapter2('Key Concept: The Derivative')

DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserNumberWithUnits.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

Context("Numeric");

## Do NOT show partial correct answers
$showPartialCorrectAnswers = 0;
TEXT(beginproblem());

$a1 = random(0.3,0.5,0.05);
$a2 = random(1.3,1.9,0.1);
$b  = ($a1-0)/0.2;
$c  = ($a2-$a1)/0.2;
$v  = ($b+$c)/2;
if ($c>$b){
  $max = $c;
  $min = $b;
} else {
  $max = $b;
  $min = $c;
}

Context()->texStrings;
BEGIN_TEXT

Consider a car whose position, \( s\), is given by the table
$PAR
$BCENTER
\{begintable(7) \}
\{row("\(t\) (s)" ,0,0.2,0.4,0.6,0.8,1.0) \}
\{row("\(s\) (ft)",0,$a1,$a2,3.8,6.5,9.6) \}
\{endtable()\}
$ECENTER

$PAR
Find the average velocity over the interval \(0 \le t \le 0.2\).
$BR
average velocity = \{ ans_rule(8) \}
\{ AnswerFormatHelp("units") \}
$PAR
Estimate the velocity at \(t=0.2\).
$BR
velocity = \{ ans_rule(8) \}
\{ AnswerFormatHelp("units") \}

END_TEXT
Context()->normalStrings;

ANS(NumberWithUnits( $b, "ft/s" )->cmp() );

#$bval = NumberWithUnits( $b, "ft/s" );
#$cval = NumberWithUnits( $c, "ft/s" );
#$vval = NumberWithUnits( $v, "ft/s" );

ANS(NumberWithUnits( $v, "ft/s" )->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    $stu_real = Real($student);
    if ( $min <= $stu_real && $stu_real <= $max ){   
#    if ( $student == $correct || $student == $cval || $student == $bval ) {
	return 1;
    } else {
	return 0;
    } } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
The average velocity over the interval \(0 \le t \le 0.2\) is given by

$PAR
\(\frac{s(0.2)-s(0)}{0.2}=\frac{$a1}{0.2}=$b \) ft/s.

$PAR
The average velocity over the interval \(0.2 \le t \le 0.4\) is given by

$PAR
\(\frac{s(0.4)-s(0.2)}{0.2}=\frac{$a2-$a1}{0.2}=$c \) ft/s.


$PAR

A reasonable estimate for the velocity at \(t=0.2\) is the average
\(\frac{1}{2}\cdot ($b+$c) = $v\) ft/s.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();

