## DESCRIPTION
## Eigenvalues and eigenvectors of 3x3 matrix (distinct, real)
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(11/06/2015)
## Institution(METU-NCC)
## Author(Benjamin Walter)
## MLT(find3x3eigen)
## MLTleader(1)
## Level(4)
## Static(0)
## MO(1)


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

#####################################
#  Setup

Context("Matrix");

# 1) choose eigenvalues a,bc for matrix of form
#    [ 0 -ac  ]     [ 0   ac  ]
#    [ b a+bc ]  or [ -b a+bc ]
#
# 2) phase shift to get
#    [ p   -ac   ]
#    [ b  a+bc+p ]
#
#  eigenvalues (a+p) and (bc+p)
#  eigenvectors [c -1] and [a -b]
#
# 3) place elements on corners of 3x3 matrix
#
#    [ p     n    -ac  ]
#    [  0    L     0   ]
#    [ b     0  a+bc+p ]
#
#  eigenvectors change to 
#    [ c  0  -1 ]     and 
#    [ a  0  -b ]
#  other eigenvector will be 
#  [ ... ugly ... ]   with eigenval L
#

my $a, $c, $b, $neg, $min, $max, $phase, $n, $lambda1, $lambda2;
my $v11, $v21, $v31, $v12, $v22, $v32, $v13, $v23, $v33;

$a = non_zero_random(-5,5,1);
$c = random(1,3,1);
do { $b = non_zero_random(-5,5,1); } until (($b*$c != $a) and (abs($b*$c) < 10));

$neg = non_zero_random(-1,1,1);

$min = max(-9,-9-$a,-9-$b*$c);  $max = min(9,9-$a,9-$b*$c);
do { $phase = non_zero_random($min,$max,1); } until ((abs($a+$b*$c+$phase) < 12) and ($a != -$phase) and ($b*$c != -$phase));

$n = non_zero_random(-4,4,1);

# eigenvalues

$lambda1 =  Real("$a+$phase");
$lambda2 = Real("$b*$c+$phase");
do { $lambda3 = non_zero_random(-5,5,1); } until ($lambda3 != $lambda1 && $lambda3 != $lambda2);
$lambda3 = Real($lambda3);

# matrix

$A11 = $phase;   $A12 = $n;        $A13 =  -$neg*$a*$c;
$A21 = 0;        $A22 = $lambda3;  $A23 = 0;
$A31 = $neg*$b;  $A32 = 0;         $A33 =  $a+$b*$c+$phase;

$A = Matrix([[$A11,$A12,$A13],
             [$A21,$A22,$A23],
             [$A31,$A32,$A33]]);

# eigenvectors

$v1 = Matrix([[$c],[0],[-$neg]]);

$v2 = Matrix([[$a],[0],[-$neg*$b]]);

$v3 = Matrix([[$A12*($A33-$A22)                 ],
              [$A13*$A31-($A11-$A22)*($A33-$A22)],
              [-$A12*$A31                       ]]);

$multians1 = MultiAnswer($lambda1,$v1, $lambda2,$v2, $lambda3,$v3)->with(
    singleResult => 1,
    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};
        my ($lambda1_stu,$v1_stu,$lambda2_stu,$v2_stu,$lambda3_stu,$v3_stu) = @{$student};

        my $flag1 = 0;
        my $flag2 = 0;
        my $flag3 = 0;
        my $flag  = 0;

        Value::Error("The correct answer has three different eigenvalues") if
          ($lambda1_stu == $lambda2_stu or $lambda1_stu == $lambda3_stu or $lambda2_stu == $lambda3_stu);
        
        Value::Error("Eigenvectors can't be zero") if ($v1->isZero or $v2->isZero or $v3->isZero);

        # check the eigenvalues and eigenvectors

        if ($lambda1_stu * $v1_stu == $A * $v1_stu){ $flag1 = 1;}
        if ($lambda2_stu * $v2_stu == $A * $v2_stu){ $flag2 = 1;}
        if ($lambda3_stu * $v3_stu == $A * $v3_stu){ $flag3 = 1;}
        

        if($flag1==1 && $flag2 ==1 && $flag3 ==1){ # then all answers are correct
          return 1;  
        }
        else {  # at least one answer is wrong, so create an error message
           if ($flag1 == 1 || $flag2 ==1 || $flag3 ==1) {  # so at least one answer is correct
	     if ($flag1 == 1) { $val1 = "correct"; } else { $val1 = "wrong";}
	     if ($flag2 == 1) { $val2 = "correct"; } else { $val2 = "wrong";}
             if ($flag3 == 1) { $val3 = "correct"; } else { $val3 = "wrong";}
             $errorMsg = "eigenpair \((\lambda_1, \boldsymbol{\vec{v}_1})\) is $val1 $BR" .
			 "eigenpair \((\lambda_2, \boldsymbol{\vec{v}_2})\) is $val2 $BR" .
			 "eigenpair \((\lambda_3, \boldsymbol{\vec{v}_3})\) is $val3";
             Value::Error($errorMsg);
           }
           else { Value::Error("None of your eigenpairs are correct"); }
         }
    }
);



#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Find the eigenvalues and eigenvectors of the matrix \($A \).
$PAR
\(\lambda_1 = \)\{$multians1->ans_rule(5)\}, $SPACE \(\mathbf{v}_1 = \)\{$multians1->ans_array()\} $PAR
\(\lambda_2 = \)\{$multians1->ans_rule(5)\}, $SPACE \(\mathbf{v}_2 = \)\{$multians1->ans_array()\} $PAR
\(\lambda_3 = \)\{$multians1->ans_rule(5)\}, $SPACE \(\mathbf{v}_3 = \)\{$multians1->ans_array()\} 
END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

COMMENT('MathObject version.  Randomized 3x3 matrix with 0 in WSE positions.  Distinct nonzero real integer eigenvalues. Two eigenvectors are [X 0 X].');

ENDDOCUMENT(); 
