# DESCRIPTION
# Factor an expression of the form A^2 - c^2, where A = x + b or A = x + b y.
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring: special forms)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(sub_diff_squares)
## Level(3)
## KEYWORDS('factoring', 'factor', 'factorization', 'quadratic', 'substitution')

###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPolynomial.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("LimitedPolynomial-Strict");
$case = random(0,1,1);
$a = non_zero_random(-6,6);
$b = random(1,9);
$b2 = ($b)**2;
@vars = ("x","y","z","a","b","c");
$ch1 = random(0,3,4);
$ch2 = $ch1+random(1,2,1);
$v1 = $vars[$ch1];
$v2 = $vars[$ch2];
Context()->variables->are($v1=>"Real", $v2=>"Real");

if ($case == 0) {
  $disp = "($v1 - $a)^2 - $b2";
  $e1 = -$a - $b;
  $e2 = -$a + $b;
  $fac1 = Formula("($v1 + $e1)")->reduce;
  $fac2 = Formula("($v2 + $e2)")->reduce;
} else {
  $disp = "($v1 - $a $v2)^2 - $b2";
  $fac1 = Formula("($v1 - $a $v2 - $b)")->reduce;
  $fac2 = Formula("($v1 - $a $v2 + $b)")->reduce;
}

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1$v1^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01$v1^{500}") if ($s2->type eq 'String'); #bogus answer
    if (($c1 == $s1 && $c1 == $s2) || ($c2 == $s1 && $c2 == $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c1 == -$s2) || ($c2 == -$s1 && $c2 == -$s2)) {
      return [1,0];
    } elsif (($c1 == $s1 && $c2 == $s2) || ($c1 == $s2 && $c2 == $s1)) {
      return [1,1];
    } elsif (($c1 == -$s1 && $c2 == -$s2) || ($c1 == -$s2 && $c2 == -$s1)) {
      return [1,1];
    } elsif (($c1 == $s1 && $c2 != $s2) || ($c2 == $s1 && $c1 != $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c2 == $s2) || ($c2 == -$s1 && $c1 == $s2)) {
      return [0,1];
    } elsif (($c1 == $s2 && $c2 != $s1) || ($c2 == $s2 && $c1 != $s1)) {
      return [0,1];
    }
    return [0,0];
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your 
answer as much as possible:
$PAR
${BCENTER}
\($disp = \) 
\( \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \)
${ECENTER}
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());




COMMENT('Desc: Factor an expression of the form A^2 - c^2, where A = x + b or A = x + b y.');
ENDDOCUMENT();
