##DESCRIPTION
##  College Algebra Review: Substitution
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring trinomials)
## Date(Spring 2012)
## Institution(University of Missouri)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.4')
## Problem1('')
## KEYWORDS('algebra', 'factoring', 'quadratic')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(x=>"Real",t=>"Real");

$m = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6,7,8,9,10);
do {$n = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6,7,8,9,10)} until ($n != $m && $n != -$m);
do {$d = list_random(-10,-9,-8,-7,-6,-5,-4,-3,-2,-1,1,2,3,4,5,6,7,8,9,10)} until ($d + $m != 0 && $d + $n != 0);
do {$a = list_random(2,3,4,5,6,7,8,9,10)} until (gcd($a, $d + $m) == 1 && gcd($a,$d + $n) == 1);
$e2 = $m + $n;
$e3 = $m*$n;
$u = $d + $m;
$v = $d + $n;


$fac1 = Formula(" $a t + $u ");
$fac2 = Formula(" $a t + $v ")->reduce;

$expression = "\(  ($a t + $d)^2 + $e2 ($a t + $d)  + $e3 \)";

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text

BEGIN_TEXT

Factor the expression $expression.  Simplify your 
answer as much as possible.
$BR
$BR
\( \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \big( \)
\{ $multians->ans_rule(10) \}
\( \big) \)
\{ AnswerFormatHelp("formula") \}
$PAR
$BR
END_TEXT


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());


ENDDOCUMENT();
