## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Date(6/3/2002)
## Level(3)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.5')
## Problem1('43 44')
## KEYWORDS('algebra', 'powers')
## MLT(Exponents_Division_Mix2015)

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl','PGcourse.pl');

$showPartialCorrectAnswers = 1;

Context("Numeric");
Context()->variables->add(y => 'Real');

$k=non_zero_random(-9,9,1);
$p=random(6,15,3);
$q=random(2,4,1);
$r=$p/3;
$r2 = 2*$r;
$s=$p-$q;

$ans1 = Formula("$k*(y**($r2) - y**{$r}*x**{$r} + x**{$r2})");
$ans2 = Formula("x**{$p} * y**{$s}");

# Allow answers n*$ans1/n*$ans2 where n is a nonzero real number

$mans = MultiAnswer($ans1,$ans2)->with(
    singleResult => 0,
    checker      => sub {
        my ($correct, $student, $self) = @_;
        my ($stu1, $stu2) = @{$student};
        my ($cor1, $cor2) = @{$correct};
        my $s1 = $stu1->eval(x=>1, y=>1);
        my $s2 = $stu2->eval(x=>1, y=>1);
        my $c1 = $cor1->eval(x=>1, y=>1);
        my $c2 = $cor2->eval(x=>1, y=>1);
        @score = (0,0);
        return (1,1) if ( 
             Formula(0) != $s1*$s2 && 
             ($s1*$cor1) == ($c1*$stu1) &&  
             ($s2*$cor2) == ($c2*$stu2) &&
             ($s1*$c2) == ($s2*$c1) );
        $score[0]=1 if ($cor1==$stu1);
        $score[1]=1 if ($cor2==$stu2); 
        return @score;
    }
);


BEGIN_PGML
Reduce the fraction to its simplest form, as a quotient of polynomials (no negative exponents).  The numerator and denominator should have no common factors.   

    [`` \frac{[$k] x^{-[$p]} + [$k] y^{-[$p]}}{(x^{[$r]}+y^{[$r]})y^{-[$q]}} = \frac{A}{B} ``]
    
where 

    A = [_]{$mans}{10} and 
    
    B = [_]{$mans}{10}
END_PGML

BEGIN_PGML_SOLUTION
Multiply the numerator and denominator by [``1=x^{[$p]}y^{[$p]}/(x^{[$p]}y^{[$p]})``].  Next factor the numerator using the formula [`P^3+Q^3=(P+Q)(P^2-PQ+Q^2)`] with [`P=x^{[$r]}`] and [`Q=y^{[$r]}`],  then cancel the common term from the numerator and denominator.

    [`` \begin{gathered} 
    \left(\frac{[$k] x^{-[$p]} + [$k] y^{-[$p]}}{(x^{[$r]}+y^{[$r]})y^{-[$q]}}\right)\left(\frac{x^{[$p]}y^{[$p]}}{x^{[$p]}y^{[$p]}}\right)\\ 
    = \frac{[$k]( y^{[$p]}+x^{[$p]} )}{(x^{[$r]}+y^{[$r]})x^{[$p]}y^{[$s]} } 
    = \frac{[$k]\Big( (y^{[$r]})^3+(x^{[$r]})^3 \Big)}{(x^{[$r]}+y^{[$r]})x^{[$p]}y^{[$s]} } \\
    = \frac{[$k](y^{[$r]}+x^{[$r]})\Big((y^{[$r]})^2 - y^{[$r]}x^{[$r]}+(x^{[$r]})^2\Big)}{(x^{[$r]}+y^{[$r]})x^{[$p]}y^{[$s]} } \\ 
    = \frac{[$ans1]}{[$ans2]}
    \end{gathered} ``]
END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.