## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(6/3/2002)
## MLT(RatExp8)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.5')
## Problem1('16 29')
## KEYWORDS('algebra', 'powers')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros('PGstandard.pl', 'PGML.pl', 'parserMultiAnswer.pl', 'PGcourse.pl');

$showPartialCorrectAnswers = 1;

Context("Numeric");

$A = non_zero_random(1,9,1);
$B = non_zero_random(1,9,1);
$C = non_zero_random(-9,9,1);
$D = non_zero_random(-9,9,1);
$E = non_zero_random(-9,9,1);
$F = non_zero_random(-9,9,1);
$G = non_zero_random(-9,9,1);

$ans1 = Formula("(1-$E)*x**2 - ($C+$G)*x + ($F-$D)")->reduce;
$ans2 = Formula("x**2 + $A x +$B")->reduce;

$mans = MultiAnswer($ans1,$ans2)->with(
    singleResult => 0,
    checker      => sub {
        my ($correct, $student, $self) = @_;
        my ($stu1, $stu2) = @{$student};
        my ($cor1, $cor2) = @{$correct};
        @score = (0,0);
        return (0,0) if ( Formula(0) == $stu2 );
        return (1,1) if ( Formula(0) == ($cor2*$stu1-$cor1*$stu2) );
        $score[0]=1 if ($cor1==$stu1);
        $score[1]=1 if ($cor2==$stu2); 
        return @score;
    }
);

BEGIN_PGML;
    [`` \frac{x^2 - ([$C] x + [$D])}{x^2+[$A] x +[$B]} + \frac{[$E] x^2 - [$F] + [$G] x}{-[$B] - [$A] x -x^2}=\frac{A}{B}``]
    
where 

    A = [_]{$mans}{10}
  
    B = [_]{$mans}{10}
END_PGML

BEGIN_PGML_SOLUTION
    [``\begin{gathered} \frac{x^2 - ([$C] x + [$D])}{x^2+[$A] x +[$B]} + \frac{[$E] x^2 - [$F] + [$G] x}{-[$B] - [$A] x -x^2} \\
    = \frac{x^2 - [$C] x - [$D]}{x^2+[$A] x +[$B]} + \frac{[$E] x^2 + [$G] x- [$F]}{(-1)(x^2+[$A] x +[$B])} \\
    = \frac{(x^2 - [$C] x - [$D])-([$E] x^2 + [$G] x- [$F] )}{x^2+[$A] x +[$B]} \\
    = \frac{[$ans1]}{[$ans2]} \end{gathered} ``]
END_PGML_SOLUTION

ENDDOCUMENT(); 