##DESCRIPTION
##  Algebra problem: Simplify and then add or subtract expressions involving radicals
##ENDDESCRIPTION

##KEYWORDS('algebra', 'radical', 'radicals', 'radical expression')

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Level(2)
## Author(K. Andrew Parker)
## Institution(CUNY City Tech)
## Language(en)

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "contextLimitedPowers.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

    ###########################
    #
    #  Subclass the numeric functions
    #
    package my::Function::numeric;
    our @ISA = ('Parser::Function::numeric');

    #
    #  Override sqrt() to return a special value (usually 1) when evaluated
    #  effectively eliminating it from the product.
    #
    sub sqrt {
      my $self = shift;
      my $value = $self->context->flag("setSqrt");
      return $value+1 if $value && $_[0] == 1;  # force sqrt(1) to be incorrect
      return $value if $value;
      return $self->SUPER::sqrt(@_);
    }

    #
    #  end of subclass
    #
    package main;

    ###########################

##############################################################
#
#  Setup
#
#
Context("Numeric")->flags->set(reduceConstants=>0, reduceConstantFunctions=>0);
Context()->variables->are(p=>[Real, limits=>[0,2]], y=>[Real, limits=>[0,2]], z=>[Real, limits=>[0,2]]);
Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
LimitedPowers::OnlyPositiveIntegers();
Parser::Number::NoDecimals();

$rad = list_random(2,3,5,6);

$c1 = random(2,6,1);
do { $c2 = random(2,6,1); } until ( $c1 != $c2 );
# do { $c3 = random(2,6,1); } until ( $c1 != $c3 && $c2 != $c3 );

$out1 = random(2,5,1);
$out2 = random(2,5,1);
# $out3 = random(1,5,1);

$c = $out1 * $c1 + $out2 * $c2;

$rad1 = $rad*$c1**2;
$rad2 = $rad*$c2**2;
# $rad3 = $rad*$c3**2;

$problem = Formula("$out1 p sqrt( $rad1 p^2 ) + $out2 p^2 sqrt( $rad2 )");

$ans = Formula("$c p^2 sqrt( $rad )");

### For solutions
$c12 = $c1**2;
$c22 = $c2**2;
$cOut1 = $c1*$out1;
$cOut2 = $c2*$out2;

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

$PAR
Add or subtract the radical expression, if possible. 
$BR Do not use rational exponents in your answer.
$PAR
\[ \large{ $problem } \]
$BR
\{ans_rule(10)\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($ans->cmp(checker => sub {
      my ($correct,$student,$ans) = @_;
      return 0 if $ans->{isPreview} || $correct != $student;
      Context()->flags->set(setSqrt => 1);
      delete $correct->{test_values}, $student->{test_values};
      my $OK = ($correct == $student); # check if equal when sqrt's are replaced by 1
      Context()->flags->set(setSqrt => 0);
      Value::Error("You must simplify your answer further") unless $OK;
      return $OK;
    })->withPostFilter(sub { 
  my $ans = shift; 
  $ans->{ans_message} = "Use 'sqrt(...)' instead of rational exponents." 
    if $ans->{ans_message} eq "Powers must be positive integer constants";
  return $ans; 
}));

COMMENT("Funded by US DoE Title V: Opening Gateways grant.");

Context()->texStrings;
BEGIN_HINT

Try simplifying each radical separately.
$PAR
You might find that they have a common factor after all!

END_HINT
BEGIN_SOLUTION

$BBOLD First, reduce \( $out1 p \sqrt{$rad1 p^2} \): $EBOLD $BR
\( \displaystyle $out1 p \sqrt{$rad1 p^2} = $out 1 \sqrt{$c12 \cdot $rad \cdot p^2} \)
$BR
\( \displaystyle $out1 p \sqrt{$c12 \cdot $rad \cdot p^2} = $out1 \cdot p \cdot \sqrt{$c12} \cdot \sqrt{$rad} \cdot \sqrt{p^2} \)
$BR
\( \displaystyle $out1 \cdot p \cdot \sqrt{$c12} \cdot \sqrt{$rad} \cdot \sqrt{p^2} = $out1 \cdot p \cdot $c1 \cdot \sqrt{$rad} \cdot p = $cOut1 p^2 \sqrt{$rad} \)
$PAR
$BBOLD Then reduce \( $out2 p^2 \sqrt{$rad2} \): $EBOLD $BR
\( \displaystyle $out2 p^2 \sqrt{$rad2} = $out2 p^2 \sqrt{$c22 \cdot $rad} \)
$BR
\( \displaystyle $out2 p^2 \sqrt{$c22 \cdot $rad} = $out2 \cdot p^2 \cdot \sqrt{$c22} \cdot \sqrt{$rad} \)
$BR
\( \displaystyle $out2 \cdot p^2 \cdot \sqrt{$c22} \cdot \sqrt{$rad} = $out2 \cdot p^2 \cdot $c2 \sqrt{$rad} = $cOut2 p^2 \sqrt{$rad} \)
$PAR
$BBOLD Now return to the original problem with our simplified radicals: $EBOLD $BR
\( \displaystyle $problem = $cOut1 p^2 \sqrt{$rad} + $cOut2 p^2 \sqrt{$rad} \)
$PAR
Factor out the \( \sqrt{$rad} \): $BR
\( \displaystyle $cOut1 p^2 \sqrt{$rad} + $cOut2 p^2 \sqrt{$rad} = \sqrt{$rad} \left( $cOut1 p^2 + $cOut2 p^2 \right) \)
$PAR
Combine the like terms: $BR
\( \displaystyle \sqrt{$rad} \left( $cOut1 p^2 + $cOut2 p^2 \right) = \sqrt{$rad}\left( $c p^2 \right) = $ans \)
$BR

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();        
