##DESCRIPTION
##  Algebra problem: Add or subtract expressions involving radicals
##ENDDESCRIPTION

##KEYWORDS('algebra', 'radical', 'radicals', 'radical expression')

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Level(2)
## Author(K. Andrew Parker)
## Institution(CUNY City Tech)
## Language(en)


########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "contextLimitedPowers.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

    ###########################
    #
    #  Subclass the numeric functions
    #
    package my::Function::numeric;
    our @ISA = ('Parser::Function::numeric');

    #
    #  Override sqrt() to return a special value (usually 1) when evaluated
    #  effectively eliminating it from the product.
    #
    sub sqrt {
      my $self = shift;
      my $value = $self->context->flag("setSqrt");
      return $value+1 if $value && $_[0] == 1;  # force sqrt(1) to be incorrect
      return $value if $value;
      return $self->SUPER::sqrt(@_);
    }

    #
    #  end of subclass
    #
    package main;
    ###########################

##############################################################
#
#  Setup
#
#
Context("Numeric")->flags->set(reduceConstants=>0, reduceConstantFunctions=>0);
Parser::Number::NoDecimals();
Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
LimitedPowers::OnlyPositiveIntegers();

$rad = list_random(2,3,5,6,7,10,11);

$c1 = random(3,9,1);
do { $c2 = random(3,9,1); } until ( $c1 != $c2 );

$cAns = $c1 - $c2;

$ans = Formula("$cAns sqrt( $rad )");

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

$PAR
Add or subtract the radical expression, if possible. 
$BR Do not use rational exponents in your answer.
$PAR
\[ \large{$c1 \sqrt{$rad} - $c2 \sqrt{$rad} } \]
$BR
\{ans_rule(10)\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($ans->cmp(checker => sub {
      my ($correct,$student,$ans) = @_;
      my $studentF = $ans->{original_student_ans};
      return 0 if $ans->{isPreview} || $correct != $studentF;
      Context()->flags->set(setSqrt => 1);
      delete $correct->{test_values}, $student->{test_values};
      my $OK = ($correct == $studentF); # check if equal when sqrt's are replaced by 1
      Context()->flags->set(setSqrt => 0);
      Value::Error("You must simplify your answer further") unless $OK;
      return $OK;
    })->withPostFilter(sub { 
  my $ans = shift; 
  $ans->{ans_message} = "Use 'sqrt(...)' instead of rational exponents." 
    if $ans->{ans_message} eq "Powers must be positive integer constants";
  return $ans; 
}));

Context()->texStrings;
BEGIN_HINT

It looks like both terms have a common factor of \( \sqrt{$rad} \)...
$PAR
What happens if you factor it out?

END_HINT
BEGIN_SOLUTION

\( \displaystyle $c1 \sqrt{$rad} - $c2 \sqrt{$rad} \) has a common factor of \( \sqrt{$rad} \).
$PAR
If we factor out the \( \sqrt{$rad} \), we get:
$BR
\( \displaystyle $c1 \sqrt{$rad} - $c2 \sqrt{$rad} = \sqrt{$rad}\left($c1 - $c2\right) \)
$PAR
Now there's nothing stopping us from combining the \($c1\) and \($c2\).
$PAR
\( \displaystyle \sqrt{$rad}\left($c1 - $c2\right) = \sqrt{$rad}\left($cAns\right) = $ans \)

END_SOLUTION
Context()->normalStrings;

COMMENT("Funded by US DoE Title V: Opening Gateways grant.");

ENDDOCUMENT();        