## DESCRIPTION
## Simplify Complex Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Simplify rational expressions)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(-10, 10);
$b = non_zero_random(-10, 10);
$n = random(2, 6);
$m = random(2, 6);

#Question
$g1 = Formula("($n)/($var^2 + ($a+$b)*$var + $a*$b)")->reduce;
$g2 = Formula("($n*$m)/($var + $a)")->reduce;
$flip = Formula("($var+$a)/($n*$m)")->reduce;
$factor = Formula("($var +$a)($var + $b)")->reduce;
$gcf = Formula("$n($var +$a)")->reduce;

#Answers
$f = Formula("(1)/($m*($var + $b))");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the complex rational expression.

[``\frac{[$g1]}{[$g2]} = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
To divide by a fraction, we can multiply by the reciprocal.

[``[$g1] \div [$g2] = [$g1] \cdot [$flip]``]

Then factor:

[``\frac{[$n]}{[$factor]} \cdot [$flip]``]

There is a common factor of [`[$gcf]`].  We can reduce and multiply to get:

[``[$f]``]
END_PGML_SOLUTION

ENDDOCUMENT();