## DESCRIPTION
## Simplify Complex Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Simplify rational expressions)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
  "contextLimitedPolynomial.pl",
 
);

##############################################
Context("Numeric");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(2, 5);
$b = non_zero_random(1, 9);
do{
$b = non_zero_random(1, 9);
}until($a != $b);

#Question
$g1 = Formula("$var")->reduce;
$g2 = Formula("($a*$var)/($var + $b)")->reduce;
$g3 = Formula("($var)/($var + $b)")->reduce;
$g4 = Formula("($var)/($var -  $b)")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');
$last = -$b**2 + $a*$b;
$f = Formula("($var^2 - $a*$var + $last)/(2*$var)")->reduce;

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

#For Solution
$LCD = Formula("($var + $b)($var - $b)");
$bsq = $b**2;
$ab = $a*$b;

##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the complex rational expression.

[``\frac{[$g1]-[$g2]}{[$g3]+[$g4]} = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The LCD (Least common denominator) is [`[$LCD]`].

If we multiply every term in both the numerator and the denominator by the LCD, we will get:

[``\frac{[$LCD] \cdot [$g1] - [$LCD] \cdot [$g2]}{[$LCD]\cdot [$g3] + [$LCD] \cdot [$g4]} ``]

Canceling out common factors, this becomes:

[`` \frac{[$var]([$var]^2 - [$bsq]) - [$a][$var]([$var] + [$b]) }{[$var]([$var]-[$b]) + [$var]([$var] + [$b])} ``]

Now distribute:
[`` \frac{[$var]^3 - [$bsq][$var] - [$a][$var]^2 + [$ab][$var] }{[$var]^2 - [$b][$var] + [$var]^2 + [$b][$var]} ``]

And combine like terms:

[`` \frac{[$var]^3 - [$a][$var]^2 + [$last][$var]}{2[$var]^2} ``]

Finally we can factor a [$var] out the numberator.

[`` \frac{[$var]([$var]^2 - [$a][$var] + [$last])}{2[$var]^2} ``]

There is a common factor of [$var] and the expression reduces to:

[`` [$f] ``]

END_PGML_SOLUTION

ENDDOCUMENT();