## DESCRIPTION
## Simplify Complex Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
Context()->variables->are(m =>'Real', n=>'Real');

$a = non_zero_random(2, 7);
$b = non_zero_random(2, 7);
$n = non_zero_random(2, 7);

#Question
$g1 = Formula("(n)/(m^$n)")->reduce;
$g2 = Formula("($a)/(n)")->reduce;
$g3 = Formula("($b)/(n)")->reduce;
$g4 = Formula("(n)/(m^$n)")->reduce;
$lcd = Formula("n*m^$n");

#Answers
Context("RationalFunction-Strict");
Context()->variables->are(m =>'Real', n=>'Real');
$f = Formula("(n^2 + $a*m^$n)/($b*m^$n - n^2)");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the complex rational expression.

[``\frac{[$g1]+[$g2]}{[$g3]-[$g4]} = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The LCD is [`[$lcd]`].  We can multiply every term by the LCD to simplify:

[``\frac{[$g1] \cdot [$lcd] +[$g2]\cdot[$lcd]}{[$g3]\cdot[$lcd]-[$g4]\cdot[$lcd]}``]

[``[$f]``]
END_PGML_SOLUTION

ENDDOCUMENT();