## DESCRIPTION
## Polynomial and Rational Functions
## ENDDESCRIPTION

## KEYWORDS('polynomial functions','rational functions')

## DBsubject('Precalculus')
## DBchapter('Polynomial and Rational Functions')
## DBsection('')
## Date('01/01/10')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Functions Modeling Change')
## EditionText1('3')
## AuthorText1('Connally')
## Section1('9 Tools')
## Problem1('52')


####################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl", 
"MathObjects.pl", 
"PGunion.pl",
"parserMultiAnswer.pl",
"PGcourse.pl",
"answerHints.pl",
);

TEXT(beginproblem());

####################################
#  Setup

Context("Numeric")->variables->are(t=>"Real");
Context()->variables->set(t=>{limits=>[1,3]});

$a = random(2,9,1);
$b = random(2,9,1);
$c = random(2,9,1);
$d = random(2,9,1);

$n = random(3,9,1);

$fraction = "\frac{ $a t^2 + $b t - $c}{ $d t^{$n} }";

$ans1 = Compute("($a/$d) t^(2-$n)")->reduce;
$ans2 = Compute("($b/$d) t^(1-$n)")->reduce;
$ans3 = Compute("($c/$d) t^(-$n)")->reduce;

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Split into a sum or difference of reduced fractions.
$BR
$BR
\( \displaystyle $fraction = \) 
\{ $multians->ans_rule(15) \}
+ 
\{ $multians->ans_rule(15) \}
-
\{ ans_rule(15) \}
END_TEXT
Context()->normalStrings;


####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );

ANS($ans3->cmp()
->withPostFilter(AnswerHints( 
  -$ans3 => "There is already minus sign in front of the third answer box",
))
);


COMMENT('MathObject version');
ENDDOCUMENT();