## DESCRIPTION
## Polynomial and Rational Functions
## ENDDESCRIPTION

## KEYWORDS('polynomial functions','rational functions')

## DBsubject('Precalculus')
## DBchapter('Polynomial and Rational Functions')
## DBsection('')
## Date('01/01/10')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Functions Modeling Change')
## EditionText1('3')
## AuthorText1('Connally')
## Section1('9 Tools')
## Problem1('35')


####################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl", 
"MathObjects.pl", 
"PGunion.pl",
"parserMultiAnswer.pl",
"PGcourse.pl",
);

TEXT(beginproblem());

####################################
#  Setup

Context("Numeric");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(6,9,1);
$b = random(2,5,1);
$m = random(3,9,1);
$n = random(5,9,1);

while ($m >= $n) { $m = random(3,9,1); }

$bma = $a - $b;
$ab = $a*$b;
$nmm = $n - $m;

$fraction = "\frac{ \displaystyle \frac{x^{$m}}{x-$a} }{ \displaystyle \frac{x^{$n}}{x^2-$bma x-$ab} }";

$num = Formula("x+$b")->reduce;
$den = Formula("x^$nmm")->reduce;

$numbogus = Formula("x^($m+$n)");
$denbogus = Formula("(x-$a)(x^2-$bma x-$ab)");

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ($f1 == $f1stu && $f2 == $f2stu) ) {
          return [1,1];
      } elsif ($f1 == $f1stu) {
          return [1,0];
      } elsif ($f2 == $f2stu) {
          return [0,1];
      } elsif ($numbogus==$f1stu || $denbogus==$f2stu) {
          $ansHash->setMessage(1,"To divide by a fraction, multiply by its reciprocal");
          $ansHash->setMessage(2,"To divide by a fraction, multiply by its reciprocal");
          return [0,0];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
  $showfraction =
  "\[ $fraction = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $showfraction =
  ColumnTable(
  "\( \displaystyle $fraction = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}
Context()->normalStrings;



####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Perform the indicated operations.
Express your answer in reduced form.
$BR
$BR
$BCENTER
$showfraction
$ECENTER
END_TEXT
Context()->normalStrings;


####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );


COMMENT('MathObject version');
ENDDOCUMENT();

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ) {
          return [1,1];
      } elsif ( Formula("x^($m+$n)")==$f1stu ) {
          $ansHash->setMessage(1,"To divide by a fraction, multiply by its reciprocal");
          $ansHash->setMessage(2,"To divide by a fraction, multiply by its reciprocal");
          return [0,0];
      } elsif ( $f1/$f2==$f1stu/$f2stu ) {
          $ansHash->setMessage(1,"Simplify your answer further");
          $ansHash->setMessage(2,"Simplify your answer further");
          return [0,0];          
      }
  }
);
