## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring polynomials: general)
## Level(2)
## MO(1)
DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric");
#Context()->variables->add(y=>"Real");
#Context("Fraction-NoDecimals");

$a = list_random(2,3,5);
$b = $a + 1;
$apb = $a + $b;

$fac1 = Formula("($a * x - $b)");
$fac2 = Formula("(x - 1)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ||
           ($f1 == -($f1stu) && $f2 == -($f2stu)) ||
           ($f1 == -($f2stu) && $f2 == -($f1stu))
         ) {
          return 1;
      } else {
          return 0;
      }
  }
);

#Context()->operators->undefine("+","-","*","/","**","^");
#Context()->functions->disable("All");

##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Factor: \( $a x^2 - $apb x + $b = \)
\( \big( \) 
\{$multians->ans_rule(6)\} 
\( \big) \big( \) 
\{$multians->ans_rule(6)\} 
\( \big) \)

$BR
$BR
$HR
Note: Your answers should be completely simplified.
Unsimplified answers will not be accepted.  

END_TEXT
Context()->normalStrings;

##################################################
#  Answer evaluation

install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 0;

ANS($multians->cmp() );


;
ENDDOCUMENT();
