## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring trinomials)
## Level(3)
## MO(1)
DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric");
#Context()->variables->add(y=>"Real");
#Context("Fraction-NoDecimals");

$a = random(3,6,1);
$b = $a - 1;
$ab = $a * $b;

$fac1 = Formula("(x + $a)");
$fac2 = Formula("(x - $b)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ||
           ($f1 == -($f1stu) && $f2 == -($f2stu)) ||
           ($f1 == -($f2stu) && $f2 == -($f1stu))
         ) {
          return 1;
      } else {
          return 0;
      }
  }
);

#Context()->operators->undefine("+","-","*","/","**","^");
#Context()->functions->disable("All");

##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Factor: \( x^2 + x - $ab = \)
\( \big( \) 
\{$multians->ans_rule(6)\} 
\( \big) \big( \) 
\{$multians->ans_rule(6)\} 
\( \big) \)

$BR
$BR
$HR
Note: Your answers should be completely simplified.
Unsimplified answers will not be accepted.  

END_TEXT
Context()->normalStrings;

##################################################
#  Answer evaluation

install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 0;

ANS($multians->cmp() );


;
ENDDOCUMENT();
