## DESCRIPTION
## Modeling with linear functions and equations
## ENDDESCRIPTION


## BookTitle('Algebra: Form and Function')
## BookChapter('Linear functions, expressions, and equations')
## DBsubject(Operations research)
## DBchapter(Linear programming)
## DBsection(Constrained optimization - planar)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(GraphProblems)
## MLTleader(1)
## Level(5)
## MO(1)
## TitleText1('Algebra: Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('5.5')
## Problem1('10')
## KEYWORDS('linear modeling')

##############################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserImplicitPlane.pl",
  "parserMultiAnswer.pl",
  "PGgraphmacros.pl",
  "unionTables.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$refreshCachedImages=1;

##############################################
#  Setup

Context("ImplicitPlane")->variables->are(w=>"Real",d=>"Real");

$speedwalk  = random(2,5,1);
$speeddrive = random(45,75,15);

$answera = ImplicitPlane("$speedwalk * w + $speeddrive * d = 60");
$answerb = Compute("60/$speedwalk");
$answerc = Compute("60/$speeddrive");

$w = Real("1");
$d = Formula("(60 - $speedwalk)/$speeddrive");

$multians = MultiAnswer($w, $d)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $f1stu<=0 || $f2stu<=0 ) {
           $self->setMessage(1,'Both w and d must be positive.');
           return 0;
      } 
      if ( ($speedwalk * $f1stu + $speeddrive * $f2stu) == 60 ) {
          return 1;
      } else {
          return 0;
      }
  }
);

#####################
#  Graph

$xmin = -1;
$xmax =  9;
$ymin = -1;
$ymax =  9;

$graph = init_graph($xmin,$ymin,$xmax,$ymax,'axes'=>[0,0]);
$graph->lb('reset');

#  axes labels
$graph->lb(new Label($xmax-0.1,0.05,"d",'black','right','bottom'));
$graph->lb(new Label(0.05,$ymax-0.1,"w",'black','left','top'));

@labels = ("A","B","C","D");
@colors = ("blue","red","gray","green");

@perm = shuffle(4);

# Graph A, correct answer
$graph->moveTo(0,8);
$graph->lineTo(3,0,$colors[0],2);
$graph->lb(new Label(3,-0.5,$labels[$perm[0]],$colors[0],'left','bottom'));

# Graph B, incorrect answer
$graph->moveTo(0,3);
$graph->lineTo(8,0,$colors[1],2);
$graph->lb(new Label(-0.5,3,$labels[$perm[1]],$colors[1],'left','bottom'));

# Graph C, incorrect answer
$graph->moveTo(0,0);
$graph->lineTo(2,9,$colors[2],2);
$graph->lb(new Label(1.5,8.5,$labels[$perm[2]],$colors[2],'left','bottom'));

# Graph D, incorrect answer
$graph->moveTo(0,0);
$graph->lineTo(9,2,$colors[3],2);
$graph->lb(new Label(8.5,1.5,$labels[$perm[3]],$colors[3],'left','bottom'));




#############################################
#  Main text

BEGIN_TEXT
\{
ColumnTable(
"Suppose you walk for \( w \) hours at \( $speedwalk \)
miles per hour and then drive a car 
for \( d \) hours at \( $speeddrive \) miles per hour, 
covering a total distance of 60 miles.".
$BR.
$BR.
"(a) Write a constraint equation that relates the 
time spent walking and driving to the total distance
covered.".
$BR.
ans_rule(20).$SPACE.
AnswerFormatHelp("equations").
$BR.
$BR.
"(b) If you only walk and do not drive, how much time 
do you spend walking?".
$BR.
"\( w \) = ".ans_rule(10)." hours ".
AnswerFormatHelp("numbers").
$BR.
$BR.
"(c) If you only drive and do not walk, how much time 
do you spend driving?".
$BR.
"\( d \) = ".ans_rule(10)." hours ".
AnswerFormatHelp("numbers").
$BR.
$BR.
"(d) Find a solution to your equation in part (a) that
has spends some time driving and some time walking.".
$BR.
"\( w \) = ".$multians->ans_rule(10)." hours ".
AnswerFormatHelp("numbers").
$BR.
"\( d \) = ".$multians->ans_rule(10)." hours ".
AnswerFormatHelp("numbers").
$BR.
$BR.
"(e) On a piece of paper, sketch a graph your equation
in part (a) and mark your solutions from parts (b)-(d)
on your graph.".
$BR.
$BR.
"Which graph A-D in the figure best resembles your graph? ".
pop_up_list(["?","A","B","C","D"]),
image(insertGraph($graph), width => 400, height => 400, tex_size => 700).
$BR.$BCENTER."(Click on graph to enlarge)".$ECENTER,
indent => 0, separation => 30, valign => "TOP"
)
\}
END_TEXT


#############################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($answera->cmp() );
ANS($answerb->cmp() );
ANS($answerc->cmp() );
ANS($multians->cmp() );
ANS(str_cmp($labels[$perm[0]]) );


;
ENDDOCUMENT();
