## DESCRIPTION
## Laplace transforms: convolution
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Laplace transforms)
## DBsection(Convolutions)
## Institution(Fort Lewis College)
## MLT(convolution_02)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('5.6')
## Problem1('04')
## KEYWORDS('Laplace transform')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
s=>"Real",
t=>"Real",
w=>"Real"
);

$a = random(2,6,1);

$ft = Compute("e^t");
$gt = Compute("e^(-$a*t)");

$int1 = Compute("e^(t-w)*e^(-$a*w)");
$int2 = Compute("e^w e^(-$a(t-w))");

$answera = Compute("e^t/(-(1+$a))*(e^(-($a+1)*t) - 1)");
$Fs = Compute("1/(s-1)");
$Gs = Compute("1/(s+$a)");
$FsGs = $Fs*$Gs;
$answerb = Compute("1/(1+$a)*(e^t - e^(-$a*t))");


#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT

For the functions \( f(t) = $ft \) and \( g(t) = $gt \), defined on \( 0\leq t < \infty \), compute \( f \ast g \) in two different ways:

\{ BeginList('OL', type=>'a') \}

$ITEM By directly evaluating the integral in the definition of \( f \ast g \).
$BR
$BR
\( \displaystyle (f \ast g)(t) = \int_0^t \) \{ ans_rule(20) \} \( dw \ = \ \) \{ ans_rule(40) \}

$ITEMSEP
$ITEM  By computing \( {\mathcal L}^{-1} \left\lbrace F(s)G(s) \right\rbrace \) where \( F(s) = {\mathcal L} \left\lbrace f(t) \right\rbrace \) and \( G(s) = {\mathcal L} \left\lbrace g(t) \right\rbrace \).
$BR
$BR
\( \displaystyle (f \ast g)(t) = {\mathcal L}^{-1} \left\lbrace F(s)G(s) \right\rbrace = {\mathcal L}^{-1} \big\lbrace \)
\{ ans_rule(20) \} 
\( \big\rbrace \)
$BR
$BR
\( \displaystyle \phantom{(f \ast g)(t)} = \) \{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}

\{ EndList('OL', type=>'a') \}
END_TEXT
Context()->normalStrings;


##############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $int1->cmp(
  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    return $int1 == $student || $int2 == $student;
  }
));
ANS( $answera->cmp() );
ANS( $FsGs->cmp() );
ANS( $answerb->cmp() );

ENDDOCUMENT();
