# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Aaron Wooton, <wootton@up.edu>
# revised by Gavin LaRose, <glarose@umich.edu> 10/2008

## Tagged by glr 10/21/08

## Author('Gavin LaRose');
## DBsubject(Algebra)
## DBchapter(Exponential and logarithmic expressions and functions)
## DBsection(Exponential functions)
## Institution(University of Michigan)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('1.2')
## Section2('1.2')
## Section3('1.2')
## Section4('1.2')
## Problem1('21')
## Problem2('13')
## Problem3('13')
## Problem4('13')
## KEYWORDS('calculus','function', 'exponential function', 'exponentials')

## Textbook tags
## HHChapter1('A Library of Functions')
## HHChapter2('A Library of Functions')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Numeric");

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a=random(2,5,1);
$b=random(70,80,0.1);
$c=random(7,12,1);
$d=random(170,180,0.1);
$f=$d/$b;
$e=($c-$a);
$ans1 = Compute("($f)**(1/$e)");
$ans2 = Compute("$ans1-1");
$P = Compute("100*$ans2");

Context()->texStrings;
BEGIN_TEXT

Suppose that \(f(t)=Q_{0}a^{t}=Q_{0}(1+r)^{t}\) with \(f($a)=$b\) and
\(f($c)=$d\). Find the following:

$PAR
(a) \(a=\) \{ans_rule( 30) \}
$PAR
(b) \(r=\) \{ans_rule( 30) \}
$PAR
${BITALIC}(Give both answers to at least 5 decimal places.)$EITALIC

END_TEXT
Context()->normalStrings;

Context()->{format}{number} = "%0.6g";
$a1 = Compute($ans1)->with(tolType=>'absolute', tolerance=>'0.00001');
$a2 = Compute($ans2)->with(tolType=>'absolute', tolerance=>'0.00001');
# ensure that fuzzy checking doesn't result in the two answers getting
#    marked inconsistently
# ANS($a1->cmp() );
ANS( $a1->cmp( checker=>sub { my ( $c, $s, $ah ) = @_;
     return $a2 == $s - 1; } ) );
ANS($a2->cmp() );

$fr = round(100000*$f)/100000;
$a1r = round(100000*$ans1)/100000;
$a2r = $a1r - 1;
$Pr = 100*$a2r;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
(a) Let \(Q=Q_{0}a^{t}\). Then \(Q_{0} a^{$a}=$b\) and \(Q_{0} a^{$c}=$d\). So

$PAR
\(Q_{0} a^{$c}/Q_{0} a^{$a}=$d/$b \approx $fr = a^{$e}\).

$BR
So \(a \approx $a1r\).

$PAR
(b) Since \(a \approx $a1r\), the growth rate is \(r \approx $a2\).
This is often expressed as a percent, which would be a $Pr
percent growth rate.

END_SOLUTION
Context()->normalStrings;



;
ENDDOCUMENT();
