# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 12/09/08

## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Implicit differentiation)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(xy_powers_square_root)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('3.7')
## Section2('3.7')
## Section3('3.7')
## Section4('3.7')
## Problem1('5')
## Problem2('5')
## Problem3('8')
## Problem4('8')
## KEYWORDS('calculus','derivative','implicit functions','implicit differentiation')

## Textbook tags
## HHChapter1('Short-Cuts to Differentiation')
## HHChapter2('Short-Cuts to Differentiation')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
 
Context("Numeric");
Context()->variables->add( y=>'Real' );
Context()->flags->set( limits=>[0.1,10] );

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$r = random(2,5,1);
$root = ($r == 2) ? '\sqrt' : "\root $r \of";
$a = random(2,15,1);

Context()->texStrings;
BEGIN_TEXT

Find \( dy/dx \) in terms of \( x \) and \( y \) 
if
\( $root {x} = $a \sqrt{y} \).

$PAR
\( {dy\over dx} = \) \{ ans_rule(65) \}

END_TEXT
Context()->normalStrings;

$cor2 = Formula( "2*x^(2/$r - 1)/($a*$a*$r)" );

ANS(Compute( "2*x^(-1+1/$r)*sqrt(y)/($r*$a)" )->cmp( checker=>sub {
    my ($cor, $stu, $ahash) = @_;
    return $cor == $stu || $cor2 == $stu; } ) );

$rexp = ($r-1) . "/$r";
$fr = ( $a*$r/2 == int($a*$r/2) ) ? "\frac{1}{" . ($a*$r/2) . "}" : 
    "\frac{2}{" . ($a*$r) . "}";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

We differentiate implicitly both sides of the equation with respect to
\(x\), noting that \( $root {x} = x^{1/$r} \) and 
\( \sqrt{y} = y^{1/2} \).
\[ 
\frac{1}{$r}\, x^{-$rexp} = \frac{$a}{2}\, y^{-1/2}\, \frac{dy}{dx}.
\]
Solving for \(dy/dx\),
\[
\frac{dy}{dx} = \frac{1}{$r} x^{-$rexp}\cdot \frac{2}{$a} y^{1/2},
\]
or
\[
\frac{dy}{dx} = $fr\cdot \frac{y^{1/2}}{x^{$rexp}}.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
