# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 12/09/08

## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Implicit differentiation)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(xy_powers_eval_approximate)
## MLTleader(1)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('3.7')
## Problem1('31')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('3.7')
## Problem2('31')
## TitleText3('Calculus')
## AuthorText3('Hughes-Hallett')
## EditionText3('6')
## Section3('3.7')
## Problem3('35')
## KEYWORDS('calculus','derivative','implicit functions','implicit differentiation')

## Textbook tags
## HHChapter1('Short-Cuts to Differentiation')
## HHChapter2('Short-Cuts to Differentiation')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( y=>'Real' );
Context()->flags->set( tolerance=>0.0001 );

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

# set up an initial equation and state
# x^r + y^3 - a x y^2 = b
# note that y=x is a solution to this if 
# x^r + x^3 - a x^3 = b, in particular, if a = 2 and r = 3, b goes to
# zero and we end up with y=x being a solution.  we don't like this,
# so tweak the parameters to be sure to avoid it.

$numTries = 0;
$notDone = 1;

# we try up to three times to be sure that the Newton's method on
#   the solution converges; if it doesn't work by then, we bail
#   and deal with the carnage later
while ( $numTries < 3 && $notDone ) {
    $r = random(2,5,1);
    $a = random(1,3,1);
    $x0 = 1;
    $y0 = random(1,4,1);

    # the derivative is 
    #    y' = (a y0^2 - r)/(y0 (3 y0 - 2 a))
    # we have a vertical slope if $a == 3 and $y0 == 2
    $a = 2 if ( $a == 3 && $y0 == 2 );
    # and don't like the solution y = x
    $a = 4 if ( $r == 3 && $a == 2 );
    # and don't like zero slopes, either
    $y0++ if ( $a*$y0*$y0 == $r );

    # what if this made y0 = 2 while a = 3?
    $y0++ if ( $a == 3 && $y0 == 2 );

    $ae = ( $a == 1 ) ? '' : $a;

    # this guarantees (x0,y0) is on the curve
    $b = 1 + $y0*$y0*$y0 - $a*$y0*$y0;

    # then the derivative is
    $dydx = "($a*y^2 - $r*x^($r-1))/(3*y^2 - 2*$a*x*y)";
    # or, as a numerical slope at (x0,y0)
    $y0sq = $y0*$y0;

    $m = ($a*$y0sq - $r*($x0**($r-1)))/(3*$y0sq - 2*$a*$x0*$y0);

    # and then values around 1 are estimated to be 
    #
    @yest = ( "$y0 + $m*(0.96 - 1)",
	      "$y0 + $m*(0.98 - 1)",
	      "$y0 + $m*(1.02 - 1)",
	      "$y0 + $m*(1.04 - 1)" );

    # finally, we need y when x = 0.96.  we'll newtonize it
    $maxIterations = 10;
    $tol = 0.0001;
    $y = $y0;
    $yprev = $y0-1;
    # $it = "$y - ";
    sub newt {
	my $yn = shift();
	$f = ($yn*$yn*($yn - 0.96*$a) - $b + (0.96**$r));
	$fd = ($yn*(3*$yn - 1.92*$a));
	return( $yn - $f/$fd );
    }
    for ( my $i=0; $i<$maxIterations; $i++ ) {
	$yprev = $y;
	$y = newt( $y );
    #     $it .= "$y - ";
        if ( abs($y - $yprev) < $tol ) {
	    $notDone = 0;
	    last;
	}
    }
    # $it =~ s/- $//;

    # the difference between the estimate and this value is
    $diff = "abs($yest[0] - $y)";

    $numTries++; 
}
# if that didn't work, just assign some values that we know will work
if ( $notDone ) {
    $a = 1;
    $r = 3;
    $b = 5;
    $y0 = 4;

    # then the derivative is
    $dydx = "($a*y^2 - $r*x^($r-1))/(3*y^2 - 2*$a*x*y)";
    # or, as a numerical slope at (x0,y0)
    $y0sq = $y0*$y0;
    $m = ($a*$y0sq - $r*($x0**($r-1)))/(3*$y0sq - 2*$a*$x0*$y0);

    # and then values around 1 are estimated to be 
    #
    @yest = ( "$y0 + $m*(0.96 - 1)",
	      "$y0 + $m*(0.98 - 1)",
	      "$y0 + $m*(1.02 - 1)",
	      "$y0 + $m*(1.04 - 1)" );

    # and the exact value at x = 0.96 is
    $y = 1.9944990;
    $diff = "abs($yest[0] - $y)";
}

@yestMO = ( map { Compute($_)->with(tolerance=>0.01) } @yest );
$yMO = Compute( $y );

$diffMO = Compute($diff)->with(tolerance=>0.15);

$mp = MultiAnswer( @yestMO, $yMO, $diffMO )->with(
    singleResult => 0,
    allowBlankAnswers => 1,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	my @stu = @{$student};
	my @cor = @{$correct};
	my $sdiffT = Compute($stu[5])->with( tolerance=>0.15 );

	my $ret = [ $stu[0] == $cor[0],  ## y(0.96)
		    $stu[1] == $cor[1],  ## y(0.98)
		    $stu[2] == $cor[2],  ## y(1.02)
		    $stu[3] == $cor[3],  ## y(1.04)
		    $stu[4] == $cor[4] ];  ## y at 0.96

	if ( ref($stu[4]) eq ref($stu[0]) ) {
	    push( @{$ret}, $sdiffT == abs( $stu[4] - $stu[0] ) );
	} else {
	    push( @{$ret}, 0 );
	}
	$self->setMessage(6, "Your difference is correct, but one or both " .
			  "of the values for y and the estimate may be " .
			  "wrong.  Be sure to calculate the new, correct " .
			  "difference when you correct the other value(s).")
	    if ( $sdiffT != $cor[5] && $ret->[5] );

	return $ret;
    }
);

# this omits the parts about the tangent being horizontal and vertical
# should they be included?  put into another problem?

Context()->texStrings;
BEGIN_TEXT

If \( x^{$r} + y^3 - $ae x y^2 = $b \), find \( dy/dx \) in terms of
\( x \) and \( y \).
$BR
\( {dy\over dx} = \) \{ ans_rule(65) \}
$PAR
Using your answer for \( dy/dx \), fill in the following table of
approximate \( y \)-values of points on the curve near 
\( x=$x0, y=$y0 \).
$BR
\{ begintable(4) \}
\{ row( '0.96', '0.98', '1.02', '1.04' ) \}
\{ row( $mp->ans_rule(8), $mp->ans_rule(8), $mp->ans_rule(8), 
        $mp->ans_rule(8) ) \}
\{ endtable() \}

$PAR
Finally, find the \( y \)-value for \( x=0.96 \) by substituting 
\( x=0.96 \) in the original equation and solving for \( y \) using a
computer or calculator.
$BR
\( y(0.96) \approx \) \{ $mp->ans_rule(8) \}
$PAR

How large (in magnitude) is the difference between your estimate for 
\( y(0.96) \)
using \( dy/dx \) and your solution with a computer or calculator?
\{ $mp->ans_rule(8) \}

END_TEXT
Context()->normalStrings;

ANS( Compute($dydx)->cmp() );
ANS( $mp->cmp() );

$rexp = ( $r==2 ) ? '' : '^{' . ($r-1) . '}';
$twoa = 2*$a;
$sl = ($a*$y0*$y0 - $r)/(3*$y0*$y0 - $twoa*$y0);
$eql = ( 1000*$sl == int(1000*$sl) ) ? '=' : '\approx';
$sl = round(1000*$sl)/1000;
$sld = ( $sl == 1 ) ? '' : $sl;

@yestr = ( round(10000*($y0 + $m*(0.96 - 1)))/10000,
           round(10000*($y0 + $m*(0.98 - 1)))/10000,
           round(10000*($y0 + $m*(1.02 - 1)))/10000,
           round(10000*($y0 + $m*(1.04 - 1)))/10000 );
$yr = round(100000*$y)/100000;

$diffn = round(10000*abs($yr - $yestr[0]))/10000;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Implicitly differentiating the equation yields
\( $r x$rexp + 3 y^2 \frac{dy}{dx} - $ae y^2 - $twoa x y \frac{dy}{dx} = 0 \),
so that 
\( {dy\over dx} = \frac{$a y^2 - $r x$rexp}{3y^2 - $twoa x y} \).

$PAR

We can approximate the curve near \(x=1\),  \(y=$y0\) by its tangent line.  The
tangent line will have slope
\( \frac{$a ($y0)^2 - $r (1)$rexp}{3 ($y0)^2 - $twoa (1) ($y0)} $eql $sl \).
Thus its equation is
\( y $eql $y0 + $sld(x - 1) \).

$PAR

Using the \(y\)-values of the tangent line to approximate the \(y\)-values
of the curve, we get 
\( y(0.96) \approx $yestr[0] \),
\( y(0.98) \approx $yestr[1] \),
\( y(1.02) \approx $yestr[2] \), and
\( y(1.04) \approx $yestr[3] \).

$PAR

When \(x=0.96\), we get the equation 
\( (0.96)^{$r} + y^3 - $ae (0.96) y^2 = $b \), whose
solution by numerical methods (for example, using a calculator) is $yr.

$PAR

The difference between $yr and the estimate is just 
\( |$yr - $yestr[0]| = $diffn \).

END_SOLUTION
Context()->normalStrings;



ENDDOCUMENT();
