## DESCRIPTION
## Add Rational Expressions w Same Denominator
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a','b','c','d','m','n','p','r','s','t','w','y');
Context()->variables->are($var =>'Real');

#Problem is of the form:
#ax/(bx + c) + (dx + e)/(bx + c)
$a = non_zero_random(-8, 8);
$c = list_random(-8, -6, -4, -2, 2, 4, 6, 8);
$b = list_random(-7, -5, -1, 1, 5, 7);
$e = non_zero_random(-9, 9, 2);
$sum = list_random(2,4,8);
$d = $sum - $a;

$den = Formula("$b*$var + $c")->reduce;

#Question
$g1 = Formula("($a*$var)/($b*$var+$c)")->reduce;
$g2 = Formula("($d*$var+$e)/($b*$var + $c)")->reduce;

#answer
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');

$f = Formula("($sum*$var+$e)/($b*$var+$c)")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Add and simplify.

[``[$g1] + [$g2] = ``]  [________]{$cmp}

 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
The denominators are already the same--add the numerators and simplify if possible.

[``\frac{[$a][$var] + [$d][$var] + [$e]}{[$den]}``]

[``[$f]``]

There are no common factors (we can't factor the numerator or the denominator), so this answer cannot be simplified.
END_PGML_SOLUTION

ENDDOCUMENT();