## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(1, 10);
do{
$b = non_zero_random(-10, 10);
}until($b != $a);


#Question
$g1 = Formula("($var)/($var - $a)")->reduce;
$g2 = Formula("($var + $b)/($var + $a)")->reduce;

$lcd = Formula("($var + $a)($var - $a)")->reduce;
$num1 = Formula("$var^2 + $a*$var")->reduce;
$num2 = Formula("$var^2 + $b*$var - $a*$var - $a*$b")->reduce;
$num2b = Formula("-$var^2 -($b-$a)*$var + $a*$b")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');

$first = 2*$a - $b;
$last = $a*$b;

$f = Formula("($first*$var + $last)/(($var + $a)($var-$a))");
$ans = $f->reduce;

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] - [$g2] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The LCD must contain each factor from each denominator, so the LCD = [`[$lcd]`].

We need to rewrite each expression with the LCD:

[`` [$g1] \cdot \frac{[$var]+[$a]}{[$var]+[$a]} - [$g2] \cdot \frac{[$var]-[$a]}{[$var]-[$a]} ``]

[`` \frac{[$num1]}{[$lcd]} - \frac{[$num2]}{[$lcd]} ``]

Notice the subtraction applies to every term in the 2nd fraction:

[``\frac{[$num1]}{[$lcd]} + \frac{[$num2b]}{[$lcd]}``]

[``[$ans]``]
END_PGML_SOLUTION

ENDDOCUMENT();