## DESCRIPTION
## Find a LCD of Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
Context()->variables->add(y =>'Real');

#constants
$a = random(2, 5);
$b = list_random(2, 4);
do{
$c = list_random(3, 5, 7);
}until($c != $a);

#exponents
$n1 = random(3, 10);
do{
$n2 = random(2, 10);
}until($n1 > $n2);
$m = random(2, 10);

$lcd = Formula("$a*$b*$c*x^$n1*y^$m");

#Question
$g1 = Formula("(y)/($a*$b*x^$n1)")->reduce;
$g2 = Formula("(1)/($b*$c*x^$n2*y^$m)")->reduce;

#Answers
$f1 = Formula("($c*y^($m+1))/$lcd");
$f2 = Formula("($a*x^($n1 - $n2))/$lcd");

#Custom Answer Checkers
$cmp1 = $f1->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer doesn't have the LCD") unless $cnum == $snum;
  return 1;
});

$cmp2 = $f2->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer doesn't have the LCD") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

For the expressions,

[``[$g1] , [$g2] = ``]

(a)  Find the LCD of the expressions.

LCD = [_________]{$lcd}

(b)  Rewrite each expression as an equivalent rational expressions with the LCD.

[``[$g1] = ``][__________]{$cmp1}

[``[$g2] = ``][__________]{$cmp2}



 
END_PGML

##############################################



BEGIN_PGML_SOLUTION

END_PGML_SOLUTION

ENDDOCUMENT();