## DESCRIPTION
## Linear equations
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Rational equations and functions)
## DBsection(Rational equations)
## Date(01/01/10)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(Cross_multiplication)
## Level(3)
## MO(1)
## TitleText1('Algebra Form and Function')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## EditionText1('1')
## Section1('5.3')
## Problem1('16')
## KEYWORDS('linear equations')

##############################################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "contextFraction.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

##############################################
#  Setup

Context("Numeric");
#Context()->variables->are(t=>"Real",y=>"Real");

$a = random(2,9,1);
$b = random(2,9,1);
$c = random(2,9,1);

$num = $a + $b * $c;
$den = $a + $c;

$answer = Compute("$num / $den");

$fac1 = Compute("$a (1-x)");
$fac2 = Compute("$c (x-$b)");

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
If 
\( \displaystyle \frac{$a}{x-$b} - \frac{$c}{1-x} = 0 \),
then
\( \displaystyle \frac{$a}{x-$b} = \)
\{ ans_rule(30) \}
\{ AnswerFormatHelp("formulas") \}
$BR
$BR
By cross-multiplying in the last equation, we get an 
equation
$BR
$BR
\{ $multians->ans_rule(20) \} = \{ $multians->ans_rule(20) \},
$BR
$BR
which can be solved for 
\( x \) = \{ ans_rule(20) \}
\{ AnswerFormatHelp("fractions") \}
END_TEXT
Context()->normalStrings;

ANS( Compute("$c / (1-x)")->cmp() );

ANS( $multians->cmp() );

ANS( $answer->cmp() );

ENDDOCUMENT();  
