## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Applications)
## Date(03/19/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-A-springs_n_things)
## Level(5)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('3.6')
## Problem1('5')


##################################
#  Initialization


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGmatrixmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "unionLists.pl",
  "unionTables.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


##################################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");

$m2 = random(1,4,1);
$m1 = 2*$m2;
$k2 = $m1;

#  Differential equation in matrix form:
#
#  [ $m1  0   ] [ x_1'' ]     [ -$k2   $k2 ] [ x_1 ]
#  [   0  $m2 ] [ x_2'' ]  =  [  $k2  -$k2 ] [ x_2 ]
# 

# matrix A = M^(-1) K listed by rows (not columns)
@A = ();
$A[0] = Real("-1"); # A11
$A[1] = Real("1");  # A12
$A[2] = Real("2");  # A21
$A[3] = Real("-2"); # A22


$multians1 = MultiAnswer( $A[0], $A[1], $A[2], $A[3] )->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "<table border='0' cellspacing='5'>
    <tr><td> %s </td><td> %s </td></tr>
    <tr><td> %s </td><td> %s </td></tr>
    </table>",
    tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\[2ex] %s & %s \end{array}\right\rbrack" ,
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($stuA11, $stuA12, $stuA21, $stuA22) = @{$student};

        if ( ($stuA11 == $A[0]) && ($stuA12 == $A[1]) && ($stuA21 == $A[2]) && ($stuA22 == $A[3]) ) 
        { return 1; } else { Value::Error('At least one entry is incorrect'); }

    }
);


#  Initial conditions
$x0 = 0;
$y0 = 0;
$xprime0 = 3;
$yprime0 = list_random(-3,-9);
$yprime0abs = abs($yprime0);

$multians2 = MultiAnswer( Real("$x0"), Real("$y0"), Real("$xprime0"), Real("$yprime0") )->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "%s, %s, %s, %s",
    tex_format =>  "x_0 = %s, y_0 = %s, x'_0 = %s, y'_0 = %s" ,
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($stu1, $stu2, $stu3, $stu4) = @{$student};

        if ( ($stu1 == $x0) && ($stu2 == $y0) && ($stu3 == $xprime0) && ($stu4 == $yprime0) ) 
        { return 1; }
        elsif ( ($stu1 == $x0) && ($stu2 == $y0) && ($stu3 == $xprime0) && ($stu4 == $yprime0abs) ) { Value::Error('Hint: What direction does car 2 travel initially?'); } 
        else { Value::Error('At least one entry is incorrect'); }

    }
);




if ($yprime0 == -3) {

  $pop = PopUp(["?","toward the wall","away from the wall","they will remain stationary"],"toward the wall");
  $x = Compute("t + 2/sqrt(3) sin(sqrt(3) t)");
  $y = Compute("t - 4/sqrt(3) sin(sqrt(3) t)");

} else {

  $pop = PopUp(["?","toward the wall","away from the wall","they will remain stationary"],"away from the wall");
  $x = Compute("-t + 4/sqrt(3) sin(sqrt(3) t)");
  $y = Compute("-t - 8/sqrt(3) sin(sqrt(3) t)");


}



#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{
ColumnTable(
"Consider a system of two toy rail cars (i.e., frictionless masses).
Suppose that car 1 has mass \( $m1 \ \mathrm{kg} \) and is traveling 
at \( $xprime0 \ \mathrm{m/s} \) toward the other car.  Suppose car 2 
has mass \( $m2 \ \mathrm{kg} \) and is moving toward the other car 
at \( $yprime0abs \ \mathrm{m/s} \).  
There is a bumper on the second rail car that 
engages at the moment the cars hit and does not let go (it connects 
the two cars).  The bumper acts as a spring with spring constant
\( $k2 \ \mathrm{N/m} \).  The second car is \( 10 \ \mathrm{m} \) 
from a wall.
$BR
$BR
Let \( t = 0 \) be the time that the cars link up.  Let \( x_1 \) be
the displacement of the first car from its position at \( t = 0 \), 
and let \( x_2 \) be the displacement of the second car from its original 
position."
,
$BCENTER.
image("rail-cars-crashing-new.png", width=>434, height=>134, tex_size=>800).
$ECENTER
,
indent => 0, separation => 30, valign => "TOP"
);
\}

\{ BeginList('OL',type=>'a') \}

$ITEM
Set up a system of second-order differential equations that models this
situation.
\{ 
mbox(
[
'\( \boldsymbol{\vec{x}^{\, \prime\prime}} = \) ', 
display_matrix([[$multians1->ans_rule(5),$multians1->ans_rule(5)],[$multians1->ans_rule(5),$multians1->ans_rule(5)]]),
' \( \boldsymbol{\vec{x}} \) '
],
allowbreaks=>'yes'
) 
\}
\( x_1(0) = \) \{ $multians2->ans_rule(10) \} meters
$BR
\( x_2(0) = \) \{ $multians2->ans_rule(10) \} meters
$BR
\( x_1^{\,\prime}(0) = \) \{ $multians2->ans_rule(10) \} meters/second
$BR
\( x_2^{\,\prime}(0) = \) \{ $multians2->ans_rule(10) \} meters/second

$ITEMSEP
$ITEM
Find the solution to this system of differential equations.
$BR
\( x_1(t) = \) \{ ans_rule(40) \} meters
$BR
\( x_2(t) = \) \{ ans_rule(40) \} meters

$ITEMSEP
$ITEM
Will the cars be moving toward the wall, away from the wall,
or will they be nearly stationary?
\{ $pop->menu() \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );
#ANS( Real("$x0")->cmp() );
#ANS( Real("$y0")->cmp() );
#ANS( Real("$xprime0")->cmp() );
#ANS( Real("$yprime0")->cmp() );

ANS( $x->cmp() );
ANS( $y->cmp() );

ANS( $pop->cmp() );

;

ENDDOCUMENT();
