## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(-9, 9);
$m = random(2, 9);
$n = random(2, 9);
$c = random(2, 9);
do{
$d = non_zero_random(-9, 9);
}until( ($d != -$a*$n) and ($d != ($m*$a**2 - $a*$c)));

#Question
$g1 = Formula("($m*$var)/($var - $a)")->reduce;
$g2 = Formula("($n)/($var)")->reduce;
$g3 = Formula("($c*$var + $d)/($var^2 - $a*$var)")->reduce;

#FOR SOLUTION
$lcd = Formula("$var($var - $a)")->reduce;
$num1 = Formula("$m*$var^2");
$num2 = Formula("$n*$var - $a*$n")->reduce;
$num3 = Formula("$c*$var + $d")->reduce;
$first = $m;
$middle = $n - $c;
$last = -$a*$n - $d;
$ans = Formula("($first*$var^2 + $middle*$var + $last)/($var($var-$a))")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');

$first = $m;
$middle = $n - $c;
$last = -$a*$n - $d;

$f = Formula("($first*$var^2 + $middle*$var + $last)/($var($var-$a))");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] + [$g2] - [$g3] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The first denominator is:

[`[$var] - [$a]`]

The second denominator is:

[`[$var]`]

The third denominator is:

[`[$var]([$var]-[$a])`]

The LCD must contain all factors of each denominator, so the LCD is:

[`[$lcd]`]

We have to rewrite each expression over the LCD:

[``[$g1] \cdot \frac{[$var]}{[$var]} + [$g2] \cdot \frac{[$var]-[$a]}{[$var]-[$a]} - [$g3] ``]

[``\frac{[$num1]}{[$lcd]} + \frac{[$num2]}{[$lcd]} - \frac{[$num3]}{[$lcd]}``]

Finally combine like terms in the numerator.  Notice that the subtraction applies to all of the terms in the third fraction.

[``[$ans]``]

END_PGML_SOLUTION

ENDDOCUMENT();