## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = list_random(2, 4, 8);
$b = list_random(-7, -5, -3, 3, 5, 7);
$c = non_zero_random(2, 7);
$m = random(2, 6);
do{
$n = non_zero_random(-10, 10);
}until($n != ($m*$b)/$a);


#Question
$g1 = Formula("($m*$var+$n)/($a*$var + $b)")->reduce;
$lcd = Formula("$a*$var + $b")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');

$first = $m+$c*$a;
$last = $n+$c*$b;

$f = Formula("($first*$var + $last)/($a*$var + $b)");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] + [$c] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
We want to add:

[``[$g1] + \frac{[$c]}{1} ``]

The LCD = [`[$lcd]`]

We can rewrite the problem as:

[``[$g1] + \frac{[$c]}{1} \cdot \frac{[$lcd]}{[$lcd]} ``]

Distribute and combine like terms to get:

[``[$f]``]
END_PGML_SOLUTION

ENDDOCUMENT();