## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(1, 7);
$c = list_random(2, 4, 5, 8);
$m = list_random(1, 3, 7, 9);
$n = non_zero_random(1, 10);
#If all constants are positive, there won't be any common factors.
#To allow negative constants n != (-2am)/c

#Question
$g1 = Formula("($m)/($c*$var - $c*$a)")->reduce;
$g2 = Formula("($n)/($var^2 - $a**2)")->reduce;

#FOR SOLUTION
$lcd = Formula("$c($var - $a)($var + $a)");
$num1 = Formula("$m*$var + $m*$a")->reduce;
$num2 = $n*$c;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');
$first = $m;
$last = $m*$a + $n*$c;
$f = Formula("($first*$var +$last)/($c*($var - $a)*($var + $a))");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] + [$g2] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The first denominator can be factored as:

[`[$c]([$var] - [$a])`]

The second denominator is:

[`([$var] - [$a])([$var] + [$a])`]

That makes the LCD:

[`[$c]([$var] - [$a])([$var] + [$a])`]

We can rewrite the expressions with the LCD:

[``[$g1] \cdot \frac{[$var] + [$a]}{[$var]+[$a]} + [$g2] \cdot \frac{[$c]}{[$c]}``]

Distribute in the numerators to get:

[``\frac{[$num1]}{[$lcd]} + \frac{[$num2]}{[$lcd]}``]

Then add and combine like terms:

[``[$f]``]
END_PGML_SOLUTION

ENDDOCUMENT();