## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(-9, 9, 2);
$b = non_zero_random(-5, 5);
do{
$c = non_zero_random(-5, 5);
}until($b != $c);

$m = random(1, 8);
$n = random(1, 8);

#Question
$g1 = Formula("($m)/(2*$var^2 + ($a+2*$b)*$var + $a*$b)")->reduce;
$g2 = Formula("($n)/(2*$var^2 + ($a+2*$c)*$var + $a*$c)")->reduce;

#FOR SOLUTION
$den1 = Formula("(2*$var + $a)($var + $b)")->reduce;
$den2 = Formula("(2*$var + $a)($var + $c)")->reduce;
$lcd = Formula("(2*$var + $a)($var + $b)($var + $c)")->reduce;

$num1 = Formula("$m*$var + $m*$c")->reduce;
$num2 = Formula("$n*$var + $n*$b")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');

$first = $m + $n;
$last = $c*$m + $b*$n;

$f = Formula("($first*$var + $last)/((2*$var + $a)($var+$b)($var+$c))");
$ans = Formula("($first*$var + $last)/((2*$var + $a)($var+$b)($var+$c))")->reduce;


#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] + [$g2] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The first denominator can be factored as [`[$den1]`].

The second denominator can be factored as [`[$den2]`].

The LCD must contain every factor of each denominator, so 

LCD = [`[$lcd]`]

We have to rewrite the expressions with the LCD:

[``[$g1] \cdot \frac{[$var]+[$c]}{[$var]+[$c]} + [$g2] \cdot \frac{[$var]+[$b]}{[$var]+[$b]} ``]

[``\frac{[$num1]}{[$lcd]} + \frac{[$num2]}{[$lcd]}``]

Then combine like terms to get:

[``[$ans]``]
END_PGML_SOLUTION

ENDDOCUMENT();