## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(-6, 6);
do{
$b = non_zero_random(-10, 10);
}until($b != $a);
do{
$c = non_zero_random(-6, 6);
}until(($c != $a) and ($c != $b));

#no common factors in answer
$n = random(1, 5);
do{
$m = non_zero_random(-10, 10);
}until($m != ($n*($a - $b - $b*$a))/($b*$c));

#Question
$g1 = Formula("($m*$var)/($var^2 + ($a+$b)*$var +$a*$b)")->reduce;
$g2 = Formula("($n)/($var^2 + ($c+$b)*$var + $c*$b)")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var =>'Real');
$middle = $c*$m + $n;
$last = $a*$n;
$f = Formula("($m*$var^2 + $middle*$var +$last)/(($var + $a)($var+$b)($var+$c))");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] + [$g2] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION

END_PGML_SOLUTION

ENDDOCUMENT();