## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(-6, 6);
do{
$b = non_zero_random(-6, 6);
}until($b != $a);

$n = $a - $b;

#Question
$g1 = Formula("(-$var)/($var + $b)")->reduce;
$g2 = Formula("($n*$var)/($var^2 + ($a+$b)*$var + $a*$b)")->reduce;

#Answers
Context("RationalFunction-Strict");
Context()->variables->are($var=>'Real');
$f = Formula("(-$var)/($var + $a)");

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] + [$g2] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION

END_PGML_SOLUTION

ENDDOCUMENT();