## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a', 'b', 'c', 'r', 's', 'x', 't', 'z', 'p');
Context()->variables->are($var =>'Real');

$a = non_zero_random(-10, 10);
do{
$b = non_zero_random(1, 10);
}until($b != $a);

$m = random(2, 7);
do{
$n = random(2, 7);
}until($n != $m);

#Question
$g1 = Formula("($m)/($var^2 - ($a+$b)*$var + $a*$b)")->reduce;
$g2 = Formula("($n)/($var^2 - 2*$b*$var + $b**2 )")->reduce;

#FOR SOLUTION
$den1 = Formula("($var - $a)($var - $b)")->reduce;
$den2 = Formula("($var - $b)($var - $b)")->reduce;
$lcd = Formula("($var - $a)($var - $b)($var - $b)")->reduce;

$num1 = Formula("$m*$var - $m*$b")->reduce;
$num2 = Formula("$n*$var - $n*$a")->reduce;

#Answers
#Use the strict context if you want students to enter repeated factors as
# (x - a)^n
#Context("RationalFunction-Strict");
#Context()->variables->are($var =>'Real');

$first = $m - $n;
$last = $a*$n - $m*$b;

$f = Formula("($first*$var + $last)/(($var - $b)^2($var - $a))");

$ans = Formula("($first*$var + $last)/(($var - $b)^2($var - $a))")->reduce;

#Custom Answer Checkers
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] - [$g2] = ``][__________]{$cmp}

END_PGML

##############################################



BEGIN_PGML_SOLUTION
The first denominator can be factored as [`[$den1]`].

The second denominator can be factored as [`[$den2]`].

The LCD must contain every factor of each denominator, so 

LCD = [`[$lcd]`]

We have to rewrite the expressions with the LCD:

[``[$g1] \cdot \frac{[$var]-[$b]}{[$var]-[$b]} - [$g2] \cdot \frac{[$var]-[$a]}{[$var]-[$a]} ``]

[``\frac{[$num1]}{[$lcd]} - \frac{[$num2]}{[$lcd]}``]

Then combine like terms, making sure that the subtraction applies to both terms in the 2nd fraction to get:

[``[$ans]``]
END_PGML_SOLUTION

ENDDOCUMENT();