## DESCRIPTION
## Add or Subtract Rational Expressions
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Rational expressions: add, subtract)
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
);

##############################################
Context("RationalFunction");
#$var = list_random('a', 'b', 'c', 'd', 'm', 'n', 'x', 'y', 't');
Context()->variables->add(a =>'Real', b =>'Real');

$m = random(2, 8);
$c = non_zero_random(1, 8);
$d = non_zero_random(1, 8);


$lcd = Formula("a^2*(b + $m)*(b - $m)");

#Question
$g1 = Formula("($c*b)/(a^2*b - $m*a^2)")->reduce;
$g2 = Formula("($d*b)/(b^2 - $m**2)")->reduce;

#Answers

$f = Formula("($c*b^2 + $m*$c*b - $d*b*a^2)/$lcd");

#Custom Answer Checkers

$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});

##############################################

TEXT(beginproblem());

BEGIN_PGML

Add or subtract the rational expressions.

[``[$g1] - [$g2] = ``][__________]{$cmp}

END_PGML

##############################################

BEGIN_PGML_SOLUTION

END_PGML_SOLUTION

ENDDOCUMENT();