# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus', 'integral', 'differential equations', 'higher derivatives')
## Tagged by glr 02/09/09

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Linear Equations')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('11.11')
## Section2('11.11')
## Problem1('11')
## Problem2('11')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')
## HHSection1('Linear Second-Order Differential Equations')
## HHSection2('Linear Second-Order Differential Equations')


DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( t=>'Real', A=>'Real', B=>'Real', C=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,9,1);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find the (real-valued) general solution
to the differential equation.
\( z''+ $a z'=0 \)

$PAR
\( z(t) = \) \{ ans_rule(35) \}
$BR
${BITALIC}(Use constants A, B, etc., for any constants
in your solution formula.)$EITALIC

END_TEXT
Context()->normalStrings;

## this is only good if students don't tweak the constants; they
##     could enter y = 3 A e^(-$a t) cos(t) + ... and it should be
##     correct, which this won't catch.  But we'll assume that they
##     won't be that contrarian...
ANS( Compute("A + B*e^(-$a*t)")->cmp( checker=>sub {
    my ( $cor, $stu, $ansHash ) = @_;
    my $opt = $cor->substitute( A=>'C' );
    $opt = $opt->substitute( B=>'A' );
    $opt = $opt->substitute( C=>'B' );
    return $cor == $stu || $opt == $stu;
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

If we try a solution of the form \( z(t) = e^{rt} \) then
\[ r^2 + $a r=0\]
which has  solutions \(r=0\) and \(r=-$a\),
so that the general solution is
\[ y(t) = A + B e^{-$a t}.\]

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
