# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 07/14/10


## DBsubject(Calculus - multivariable)
## DBchapter(Fundamental theorems)
## DBsection(Line integrals)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('18.3')
## Section2('18.3')
## Section3('18.3')
## Problem1('25')
## Problem2('36')
## Problem3('41')
## KEYWORDS('line integral', 'fundamental theorem of calculus', 'gradient', 'vector field')

## Textbook tags
## HHChapter1('Line Integrals')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add(t=>'Real', y=>'Real');
$showPartialCorrectAnswers = 1;

$a = random(1,5,1);
$b = random(1,5,2);
$bd = ( $b == 1 ) ? '' : $b;
$ab = $a*$b;

$rx = Compute( "t" );
$ry = Compute( "pi*t/2" );

$which = list_random(0,1);
$which = 1;
if ( $which ) {
    $fx = Compute( "-2*$a*x*e^(-x^2)*sin($b*y)" )->reduce();
    $fy = Compute( "$a*$b*e^(-x^2)*cos($b*y)" )->reduce();
    $f = Compute( "$a*e^(-x^2)*sin($b*y)" )->reduce();

    $integr = Compute( "-2*$a*t*e^{-t^2}*sin($b*pi*t/2) + " .
		       "(pi/2)*$a*$b*e^(-t^2)*cos($b*pi*t/2)" )->reduce();

    $x1 = 1;
    $y1 = 0;
    $r1x = "t";  $r2x = "1";
    $r1y = "0";  $r2y = "\pi t/2";
    $integr1 = "0"; $val1 = "0";
    $integr2 = "\frac{\pi}{2}\,$ab\,e^{-1}\,\cos($b\frac{\pi}{2}\,t)";
    $s2 = ($b == 3) ? '-' : '';
    $val2 = "$s2$a e^{-1}";

    $int1 = Compute("0");
    $int2 = Compute("$a*e^(-1)*sin($b*pi/2)");
    $chg1 = Compute("$a*e^(-1)*sin($b*pi/2)");

    $f0 = Compute("0");
    $sgn = ($b == 3) ? '-1' : '1';
    $f1 = Compute("$sgn*$a*e^{-1}");
    $f0d = 0;
    $f1d = "$s2$a e^{-1}";
    $chg2 = Compute("$a*e^(-1)*sin($b*pi/2) - 0");
    $chg2d = $f1d;
} else {
    $fx = Compute( "2*$a*x*e^(x^2)*cos($b*y)" )->reduce();
    $fy = Compute( "-$a*$b*e^(x^2)*sin($b*y)" )->reduce();
    $f = Compute( "$a*e^(x^2)*cos($b*y)" )->reduce();

    $integr = Compute( "2*$a*t*e^{t^2}*cos($b*pi*t/2) - " .
		       "(pi/2)*$a*$b*e^(t^2)*sin($b*pi*t/2)")->reduce();

    $x1 = 0;
    $y1 = "\pi/2";
    $r1x = "0";  $r2x = "t";
    $r1y = "\pi t/2";  $r2y = "\pi/2";
    $integr1 = "-\frac{\pi}2\,$ab\sin($bd\frac{\pi}{2}\,t)";
    $val1 = "-$a";
    $integr2 = "0";  $val2 = "0";

    $int1 = Compute("-$a");
    $int2 = Compute("0");
    $chg1 = Compute("-$a");

    $f0 = Compute("$a");
    $f1 = Compute("0");
    $f0d = $a;
    $f1d = 0;
    $chg2 = Compute("0 - $a");
    $chg2d = "-$a";
}

$part_b = MultiAnswer($f,$f0,$f1,$chg2)->with(
  singleResult => 0,
  checker => sub{ 
    my ( $correct, $student, $self ) = @_;
    my ( $f_stu, $f0_stu, $f1_stu, $chg2_stu ) = @{$student};
    my @score = (0,0,0,0);
    if ( ($fx == $f_stu->D('x')) and ($fy == $f_stu->D('y')) ){
      $score[0]=1;
    }
    else{ 
      $self -> setMessage(1,"The gradient of your function must equal \(\nabla f\)"); 
      return @score;
    }
    if ( $f0_stu == $f_stu->eval(x=>0,y=>0) ){
      $score[1]=1;
    }
    if ( $f1_stu == $f_stu->eval(x=>1,y=>pi/2) ){
      $score[2]=1;
    }
    if ( $score[1] == 1 and $score[2] == 1 and 
         $chg2_stu == ($f1_stu - $f0_stu) ){
       $score[3]=1;
    }
    return @score; 
  }
);


Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Let \( {\nabla f} = $fx\,\vec i + $fy\, \vec j \).
Find the change in \( f \) between \( (0,0) \) and \( (1,\pi/2) \)
in two ways.
$PAR
${BBOLD}(a)$EBOLD
First, find the change by computing the line integral
\(\int_C \nabla f\cdot d\vec r\), where \(C\) is a curve
connecting \( (0,0) \) and \( (1,\pi/2) \).
$PAR
The simplest curve is the line segment joining these points.  Parameterize
it:
$BR
with \(0\le t\le 1\),
\(\vec r(t) = \) \{ ans_rule(15) \}\(\vec i + \)
    \{ ans_rule(15) \}\(\vec j\)
$BR
So that
\(\int_C\nabla f\cdot d\vec r = \int_0^1\)
\{ ans_rule(45) \} \(dt\)
$BR
Note that this isn't a very pleasant integral to evaluate by hand
(though we could easily find a numerical estimate for it).  It's easier
to find \(\int_C\nabla f\cdot d\vec r\) as the sum
\(\int_{C_1}\nabla f\cdot d\vec r + \int_{C_2}\nabla f\cdot d\vec r\),
where \(C_1\) is the line segment from \((0,0)\) to \(($x1,$y1)\) and
\(C_2\) is the line segment from \(($x1,$y1)\) to \((1,\pi/2)\).
Calculate these integrals to find the change in \(f\).
$BR
\(\int_{C_1}\nabla f\cdot d\vec r = \) \{ ans_rule(45) \}
$BR
\(\int_{C_2}\nabla f\cdot d\vec r = \) \{ ans_rule(45) \}
$BR
So that the change in
\(f = \int_C\nabla f\cdot\vec r =
\int_{C_1}\nabla f\cdot d\vec r + \int_{C_2}\nabla f\cdot d\vec r =\)
\{ ans_rule(35) \}

$PAR
${BBOLD}(b)$EBOLD
By computing values of \( f \).  To do this,
$BR
First find \( f(x,y) = \) \{ $part_b->ans_rule(45) \}
$BR
Thus
$BR
\( f(0,0) = \) \{ $part_b->ans_rule(25) \} and
\( f(1,\pi/2) = \) \{ $part_b->ans_rule(25) \},
$BR
and the change in \(f\) is \{ $part_b->ans_rule(35) \}.

END_TEXT
Context()->normalStrings;

#part a answers

ANS($rx->cmp() );
ANS($ry->cmp() );
ANS($integr->cmp() );
ANS($int1->cmp() );
ANS($int2->cmp() );
ANS($chg1->cmp() );

#part b answers

ANS($part_b->cmp()); 

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The parameterization is
\(\vec r(t) = (x(t),y(t)) = \langle t, \pi\,t/2\rangle\).  Then
\(\vec r\,'(t) = \langle 1, \pi/2\rangle\), so that
\[\int_C\nabla f\cdot d\vec r =
  \int_0^1 \langle $fx, $fy\rangle\cdot \langle 1,\pi/2\rangle\,dt =
\]
\[
  \int_0^1 \left($integr\right)\,dt. \]
As noted, it's simpler to integrate on the pair of line segments
\(C_1\) from \((0,0)\) to \(($x1,$y1)\) and \(C_2\) from \(($x1,$y1)\)
to \((1,\pi/2)\).  Parameterizing these in a manner similar to our
parameterization of \(C\), on \(C_1\) we have
\(\vec r = \langle $r1x, $r1y\rangle\) and on \(C_2\)
\(\vec r = \langle $r2x, $r2y\rangle\).  Then
\[ \int_{C_1}\nabla f\cdot d\vec r =
   \int_0^1 $integr1\,dt = $val1, \]
and
\[ \int_{C_2}\nabla f\cdot d\vec r =
   \int_0^1 $integr2\,dt = $val2, \]
so that the change in \(f\) is \(($val2) + ($val1) = $chg2d\).

$PAR
${BBOLD}(b)$EBOLD
The other way to find the change in \(f\) between these two points is to
first find \(f\). To do this,
observe that
\[ $fx\,\vec i + $fy\,\vec j =
\frac{\partial}{\partial x}\left($f \right) \vec i +
\frac{\partial}{\partial y}\left($f \right) \vec j =
\]
\[
\nabla \left($f\right).\]
So one possibility for \(f\) is \(f(x,y)=e^{x^2}\sin y\). Thus,
\[
\mbox{Change in } f\bigg|_{(0,0)}^{(1,\pi/2)} = $f
\bigg|_{(0,0)}^{(1,\pi/2)}= $f1d - $f0d = $chg2d.\]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
