## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION

## KEYWORDS ('linear algebra','matrix','eigenvalue','eigenvector','systems','systems of differential equations')

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Systems of Differential Equations')
## Date('01/31/2012')
## Author('Paul Pearson')
## Institution('Rochester and Fort Lewis College')
## TitleText1('Elementary Differential Equations')
## EditionText1('2006')
## AuthorText1('Kohler and Johnson')
## Section1('4.6')
## Problem1('13')


##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);


TEXT(beginproblem());


#####################################
#  Setup

Context("Numeric")->variables->are(t=>"Real");


$b = random(2,6,1);

$vtop = "-1-i"; # TeX
$vbot = "1"; # TeX

$y1top = Formula("-cos($b t)+sin($b t)");
$y1bot = Formula("cos($b t)");

$y2top = Formula("-cos($b t) - sin($b t)");
$y2bot = Formula("sin($b t)");


$multians2 = MultiAnswer($y1top, $y1bot, $y2top, $y2bot)->with(

    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  
    "<table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table>
     <br />
     <table border='0' cellspacing='5'><tr><td> %s </td></tr><tr><td> %s </td></tr></table>",
    tex_format =>  
    "\left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack \\ 
     \left\lbrack \begin{array}{r} %s \\ %s \end{array}\right\rbrack",

    checker => sub {
        my ($correct, $student, $answerHash ) = @_;
        my @c = @{$correct};
	my @s = @{$student};

        $flag1 = 0;
        $flag2 = 0;

        Value::Error("Your solutions are not linearly independent") if
           $c[0]->typeMatch($s[0])
        && $c[1]->typeMatch($s[1])
        && $c[2]->typeMatch($s[2])
        && $c[3]->typeMatch($s[3])
        && ($s[0] * $s[3] == $s[1] * $s[2]);

        if (   $c[0]->typeMatch($s[0]) 
            && $c[1]->typeMatch($s[1])
            && ($s[0] != 0 || $s[1] != 0) 
            && ($s[0] * $c[1] == $s[1] * $c[0])
           ) 
        {   $flag1 = 1; }

        if (   $c[2]->typeMatch($s[2]) 
            && $c[3]->typeMatch($s[3])
            && ($s[2] != 0 || $s[3] != 0) 
            && ($s[2] * $c[3] == $s[3] * $c[2])
           ) 
        {   $flag2 = 1; }

        if ($flag1==1 && $flag2 ==1) {
           return 1;
        } elsif ($flag1 == 1) {
           Value::Error("Solution y<sub>1</sub> is correct");
        } elsif ($flag2 == 1) {
           Value::Error("Solution y<sub>2</sub> is correct");
        }else{
           Value::Error("Neither solution is correct");
        }
       
    }
);



#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose \( A \) is a \( 2 \times 2 \) real matrix 
with an eigenvalue \( \lambda = $b i \) and
corresponding eigenvector
$PAR
\[
\boldsymbol{\vec{v}}
=
\left\lbrack
\begin{array}{c}
$vtop \\
$vbot
\end{array}
\right\rbrack.
\]
$BR
Determine a fundamental set (i.e., linearly 
independent set) of solutions for
\( \boldsymbol{\vec{y}^{\,\prime}} = A \boldsymbol{\vec{y}} \),
where the fundamental set consists entirely 
of ${BITALIC}real${EITALIC} solutions.
$BR
$BR
Enter your solutions below.  Use \( t \)
as the independent variable in your answers.
$BCENTER
\{ 
mbox(
[
'\( \boldsymbol{\vec{y}_1}(t) =  \) ', 
display_matrix([[$multians2->ans_rule(30)],[$multians2->ans_rule(30)]]),
],
allowbreaks=>'yes'
) 
\}
$BR
\{ 
mbox(
[
'\( \boldsymbol{\vec{y}_2}(t) =  \) ', 
display_matrix([[$multians2->ans_rule(30)],[$multians2->ans_rule(30)]]),
],
allowbreaks=>'yes'
) 
\}
$ECENTER

END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians2->cmp() );

COMMENT('MathObject version.');

ENDDOCUMENT(); 