##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

## Tagged by cmd6a 4/26/06

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Applications)
## Date(1/25/10)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(5)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('1.3')
## Problem1('')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

$m = Compute(random(20,40));
$h = Compute(random(5,10));
$r = Compute(random(2,4));
$g = Compute(32);

Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("t"=>"Real","y"=>"Real","y'"=>"Real","y''"=>"Real","y'''"=>"Real");
Context()->variables->set("y"=>{limits=>[$h-2,$h+2]});# If you change this change it below in the answer checker
##############################################################
#
#  Setup
#
#

$bf = Formula("62.4*($h-y)*$r^2*pi");
$diffeqn = Formula("$m*y''/$g+$m-62.4*($h-y)*$r^2*pi"); ## If you modify this you must modify the differential equation below.

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
A cylindrical tank with height \($h\), radius \($r\), and weight \($m\) lbs is vertically submerged in water, with y(t) the height of the tank above the water. Two forces act on the tank, the acceleration due to gravity, \($g\frac{ft}{s^2}\) acting in the opposite direction of y and the buoyancy force by Archimedes principle. 
$BR
If water weighs \(62.4\frac{lbs}{ft^3}\) then the buoyancy force is \{$bf->ans_rule(20)\}
$BR
Therefore the differential equation modeling the forces on the tank is
\{$diffeqn->ans_rule(25)\}\( = 0\)
$BR
$BR
Note that weight = force = ma
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS($bf->cmp());
ANS($diffeqn->cmp(checker=>sub {
                 my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('C0'=>'Parameter');
                $context->variables->set('C0'=>{limits=>[0,5]});
                $context->variables->set("y"=>{limits=>[$h-2,$h+2]});
                my $c0 = Formula($context,'C0');
                $student = Formula($context,$student);
          $correct = Formula($context,"$c0($m*y''/$g+$m-62.4*($h-y)*$r^2*pi)");
                $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}
));
#$m*y''/$g+$m-62.4*($h-y)*$r^2*pi
ENDDOCUMENT();        
