## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## Tagged by cmd6a 4/26/06

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Applications)
## Date(12/17/09)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(5)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('1.3')
## Problem1('')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "answerHints.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showHint = 5;
$showPartialCorrectAnswers = 1;
#########################################################
##
## Set up the context to allow ' in variable names
##
Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("t"=>"Real","y"=>"Real","y'"=>"Real","y''"=>"Real","y'''"=>"Real");

$r = Compute(random(0.1,0.7,0.05));
$percent = Compute("100*$r");
$m = Compute(random(10,30));
$g = Compute("9.81");
$iv = Compute(random(10,20));
$DE = Formula("$m*y''+$m*$g-$r*sqrt(y')"); ## If you modify this you must also change the answer checker

Context()->texStrings;
BEGIN_TEXT
An object with mass \($m\)kg is in free fall. It has two forces acting on it: gravity and air resistance. For this model we will take the ground to be at height 0, the height of the object to be y(t) meters from the ground, and the acceleration due to gravity to be \(-$g \frac{m}{s^2}\).
$BR
We want to set up a differential equation for the $BBOLD forces $EBOLD acting on the object. The force due to gravity is is given by mass times acceleration acting in the opposite direction of y. The force due to air resistance is proportional to the square root of the velocity of the object in the same direction is y. The constant of proportionality for the wind resistance has been measured at \($r\). Find the differential equation that models the displacement of the object y at time t. 
$BR
\{$DE->ans_rule()\} \(= 0\)
$BR
If the object is dropped from a height of \($iv\) the IVP is
$BR
\(y(0) = \) \{$iv->ans_rule\}
$BR
\(y^\prime (0)= \)\{ans_rule(10)\}
$BR
Note: use y,y', etc instead of y(t), y'(t) in your answers. After you set up your differential equation you will have to set it equal to zero so that WeBWorK will understand your answer, do this in a way so that the highest order derivative has a positive coefficient.
END_TEXT
ANS($DE->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('zz'=>'Parameter');
                my $c0 = Formula($context,'zz');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0($m*y''+$m*$g-$r*sqrt(y'))");
                 $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));
## Cannot get this to work with the custom answer checker
#->withPostFilter(AnswerHints(
#                Formula("y''+$m*$g-$r*sqrt(y')")
#=>"The second derivative is the acceleration, not the force",
#                Formula("$m*y''-$m*$g-$r*sqrt(y')")=>"Mind your signs",
#                Formula("$m*y''+$m*$g+$r*sqrt(y')")=>"Mind your signs",
#                Formula("$m*y''-$m*$g+$r*sqrt(y')")=>"Mind your signs",
#                Formula("y''-$m*$g-$r*sqrt(y')")=>"Mind your signs",
#                Formula("y''+$m*$g+$r*sqrt(y')")=>"Mind your signs",
#                Formula("y''-$m*$g+$r*sqrt(y')")=>"Mind your signs",
#        )));
ANS($iv->cmp);
ANS(Compute("0")->cmp);

BEGIN_HINT
Remember that the acceleration of the object is given by \(y^{\prime\prime}\), but we are modeling the force on the object \(m*y^{\prime\prime}\). The velocity of the object is \(y^{\prime}\).
END_HINT
Context()->normalStrings;
ENDDOCUMENT();        # This should be the last executable line in the problem.
