##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

## Tagged by cmd6a 4/26/06

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Applications)
## Date(1/25/10)
## Institution(Alfred University)
## Author(Darwyn Cook)
## MLT(circuits)
## Level(4)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('1.3')
## Problem1('')
## RESOURCES('RLC_series_circuit.png')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add("t"=>"Real","q"=>"Real","q'"=>"Real","q''"=>"Real","q'''"=>"Real","E"=>"Real");
Parser::BOP::equality->Allow;
##############################################################
#
#  Setup
#
#

$C = Real(random(5,10));
$R = Real(2*random(1,5));
$b = Real(4*(random(-2,-1))*$C);
$L = Real((($C*$R)**2-$b)/(4*$C));# This could have been done better
$diffeqn = Formula("$L*q''+$R*q'+q/$C-E");

$q0 = Real(random(-9,9));
$qprime0 = Real(random(-9,9));


$q = Formula("-1/2*e^(-2*$C*$R/(($C)^2*($R)^2-$b)*t)*sin(2*(-$b)^(1/2)/(($C)^2*($R)^2-$b)*t)*(-$qprime0*($C)^2*($R)^2+$qprime0*$b-2*$C*$R*$q0+2*($C)^2*($R)*(E))/(-$b)^(1/2)+e^(-2*$C*$R/(($C)^2*($R)^2-$b)*t)*cos(2*(-$b)^(1/2)/(($C)^2*($R)^2-$b)*t)*($q0-E*$C)+E*$C"); 

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
A radio uses an LRC loop series circuit consisting of an inductor, resistor, and capacitor to tune radio stations. For this problem \(L = $L, R = $R, C = $C\). 
$BR
\{image("RLC_series_circuit.png",width=>140, height=>210);\}
$BR
Write down the differential equation modeling the circuit using E for the impressed voltage.
$BR
\{$diffeqn->ans_rule(35)\} \(=0\)
$BR
If the voltage is constant and \(q(0)= $q0\) and \(q^\prime(0)=$qprime0\) then
$BR
\(q(t) = \) \{$q->ans_rule(50)\}
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($diffeqn->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0($L*q''+$R*q'+q/$C-E)");
                $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));
ANS($q->cmp);
ENDDOCUMENT();        
