##DESCRIPTION
##   
##ENDDESCRIPTION

## Tagged by cmd6a 4/26/06

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Applications - exponential growth & decay)
## Date(2/1/10)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(2)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('1.3')
## Problem1('')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->add("t"=>"Real","P"=>"Real","P'"=>"Real","k"=>"Real");

##############################################################
#
#  Setup
#
#
$diffeqn = Formula("P'-kP");

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
Using k as the growth constant and P for population write down the Malthusian model for population growth \{$diffeqn->ans_rule(25)\} \(=0\)
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($diffeqn->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"$c0(P'-kP)");
                 $zero = Formula($context,"0");
                return (($correct == $student)&&($student!=$zero));}));

ENDDOCUMENT();        
