##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

## Tagged by cmd6a 4/26/06

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Applications - other)
## Date(1/25/10)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(5)
## TitleText1('Differential Equations: with Boundary Value Problems')
## AuthorText1('Zill')
## EditionText1('6')
## Section1('1.3')
## Problem1('')
########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

Context()->variables->{namePattern} = qr/[a-z][a-z0-9_]*'*/i;
Context()->variables->are("t"=>"Real","g"=>"Real","h"=>"Real","A"=>"Real","V"=>"Real","V'"=>"Real");
Context()->variables->set(V=>{limits=>[0,1]},g=>{limits=>[1,2]},h=>{limits=>[1,2]});

##############################################################
#
#  Setup
#
#

$l = Compute(random(3,8));
$w = Compute(random(5,10));

$e = Compute(random(4,9));
$v = Formula("-g*t");
$y = Formula("-g*t^2/2+h");
$t = Formula("sqrt(2*h/g)");
$h = Formula("V/($l*$w)");
$diffeqn = Formula("V'-A*sqrt(2*V/($l*$w)*g)");

$context = Context()->copy;
$context->operators->remove('/');
$context->{error}{msg}{"Unexpected character '/'"} 
= "Your answer cannot have any division symbols in it.";

$vofh = Formula($context,"sqrt(2*h*g)");

$veofh = Formula ($context, "-sqrt(2*h*g)");

$dVdt = Formula($context,"A*sqrt(2*h*g)");


##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
According to Torriccelli's law the speed of water leaving the bottom of a tank is the speed that a drop of the water would have acquired if it were dropped from the height h of the water in the tank.
$BR
If the acceleration due to gravity is g, the (negative) velocity of a drop of water at time t is \{$v->ans_rule(15)\}.
$BR 
Integrating we get the height of a drop of water at time t dropped from a height h \{$y->ans_rule(20)\}.
$BR
Now find the time t when the drop of water dropped from height h hits the ground: \{$t->ans_rule(20)\}.
$BR 
Plugging this back into the equation for the velocity of the rain drop, and simplifying, we find that the speed of the drop of water dropped from height h when it hits the ground is \{$vofh->ans_rule(20)\}.
$BR
By Torriccelli's law, the velocity of water leaving a tank from the bottom, where the height of the water in the tank is h, is: \{$veofh->ans_rule(20)\}
$BR
Now we know the velocity of the water leaving the tank, if the area of the hole from which the water is leaving the tank is A, then the volume of water  leaving the tank every unit of time is \(V^\prime = \) \{$dVdt->ans_rule\}
$BR
Now suppose that we know the tank is rectangular with width \($w\) and length \($l\), then the height of the water in the tank with respect to the volume \(V\) of water in the tank is 
$BR
h = \{$h->ans_rule(20)\}
$BR
Plugging this formula for h into the formula above for \(V^\prime\), we arrive at a differential equation for the volume \(V\) of the water in the tank.
$BR
\{$diffeqn->ans_rule(25)\} = 0
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS($v->cmp());
ANS($y->cmp());
ANS($t->cmp());
ANS($vofh->cmp()->withPostFilter(AnswerHints( 
  Formula($context,"-sqrt(2*h*g)") => "Take the absolute value for the speed",))
);
ANS($veofh->cmp());
ANS($dVdt->cmp());
ANS($h->cmp());
ANS($diffeqn->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context1 = Context()->copy;
                $context1->flags->set(no_parameters=>0);
                $context1->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context1,'c0');
                $student = Formula($context1,$student);
                $correct = Formula($context1,"$c0(V'-A*sqrt(2*V/($l*$w)*g))");
                $zero = Formula($context1,"0");
                return (($correct == $student)&&($student!=$zero));}));

ENDDOCUMENT();        
