## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Institution('TCNJ')
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('6.1')
## Problem1('')
## KEYWORDS('eigenvalues' 'eigenvectors')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
           "PGbasicmacros.pl",
           "PGchoicemacros.pl",
           "PGanswermacros.pl",
           "PGgraphmacros.pl",
           "PGmatrixmacros.pl",
           "PGnumericalmacros.pl",
           "PGdiffeqmacros.pl",
           "parserMultiAnswer.pl",
	   "PGauxiliaryFunctions.pl",
);


TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

$a= new Matrix(3,3);

# create an invertible matrix with det either 1  or -1

$a11 = random(-1,1,1);
$a21 = random(-1,1,2);
$a31 = random(-1,1,2);

$b1 = random(-1,1,2);
$a12 = $b1 * $a11;
$m = random(-1,1,2);
$a22 = $b1 * $a21 + $m;
$a32 = $b1 * $a31;

$c = random(-1,1,1);
$d = random(-1,1,2);
$n = random(-1,1,2);
$a13 = $c * $a11 + $d * $a12 + $n;
$a23 = $c * $a21 + $d * $a22;
$a33 = $c * $a31 + $d * $a32;

$det = - $a31 * $m * $n;

# define matrix

	$a->assign(1,1, $a11 );
	$a->assign(1,2, $a12 );
	$a->assign(1,3, $a13 );
	$a->assign(2,1, $a21 );
	$a->assign(2,2, $a22 );
	$a->assign(2,3, $a23 );
	$a->assign(3,1, $a31 );
        $a->assign(3,2, $a32 );
        $a->assign(3,3, $a33 );
	$a_lr = $a->decompose_LR();
	$a_det = $a_lr->det_LR();

# define inverse matrix
	$b = $a_lr->invert_LR();

# define eigenvalues
    $e = new Matrix(3,3);
    $e->one();
    $e->assign(1,1, random(-5,-2,1));
    $e->assign(2,2, random(-1,1,1));
    $e->assign(3,3, random(2,5,1));
    $e1 = $e->element(1,1);
    $e2 = $e->element(2,2);
    $e3 = $e->element(3,3);

# define eigenvectors
	$v1 = new Matrix(3,1);
	$v1->assign(1,1, $a11);
        $v1->assign(2,1, $a21);
        $v1->assign(3,1, $a31);
        	$v2 = new Matrix(3,1);
        	$v2->assign(1,1, $a12);
        	$v2->assign(2,1, $a22);
        	$v2->assign(3,1, $a32);
        $v3 = new Matrix(3,1);
        $v3->assign(1,1, $a13);
        $v3->assign(2,1, $a23);
        $v3->assign(3,1, $a33);

# define final matrix
	$matrix = $a * $e *$b;
	$matrix_lr = $matrix->decompose_LR();
	$matrix_det = $matrix_lr->det_LR();

# ensure that display doesn't have odd rounding issues
	$matrix_disp = new Matrix(3,3);
	for ( my $i=1; $i<=3; $i++ ) {
		for ( my $j=1; $j<=3; $j++ ) {
		    $matrix_disp->assign($i,$j, round($matrix->element($i,$j)));
					 # sprintf("%1.0d",$matrix->element($i,$j)));
		}
	}

$ev = MultiAnswer($a11, $a21, $a31, $a12, $a22, $a32, $a13, $a23, $a33)->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "Eigenvectors = (%s, %s, %s), (%s, %s, %s), (%s, %s, %s)",
    tex_format =>  "\text{Eigenvectors }= (%s, %s, %s), (%s, %s, %s), (%s, %s, %s)" ,
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($stuv11, $stuv12, $stuv13)= @{$student}[0..2];
        my ($stuv21, $stuv22, $stuv23)= @{$student}[3..5];
        my ($stuv31, $stuv32, $stuv33)= @{$student}[6..8];
        my $flag = 0;
        my $flag1=0;
        my $flag2=0;
        my $flag3=0;
        my $flag4 =0;
        my $counter =0;
        if($stuv11 || $stuv12 || $stuv13){
           my $v1 = new Matrix(3, 1);
           if($stuv11 != 0 || $stuv12 !=0 || $stuv13){
             $v1-> assign(1,1, $stuv11);
             $v1-> assign(2,1, $stuv12);
             $v1-> assign(3,1, $stuv13);
             if( $matrix * $v1 == $e1 * $v1){
                $flag1=1;
             }
           }
        }

        if($stuv21 || $stuv22 || $stuv23){
           my $v2 = new Matrix(3, 1);
           if($stuv21 != 0 || $stuv22 !=0 || $stuv23){
             $v2-> assign(1,1, $stuv21);
             $v2-> assign(2,1, $stuv22);
             $v2-> assign(3,1, $stuv23);
             if( $matrix * $v2 == $e2 * $v2){
                $flag2=1;
             }
           }
        }

        if($stuv31 || $stuv32 || $stuv33){
           my $v3 = new Matrix(3, 1);
           if($stuv31 != 0 || $stuv32 !=0 || $stuv33){
             $v3-> assign(1,1, $stuv31);
             $v3-> assign(2,1, $stuv32);
             $v3-> assign(3,1, $stuv33);
             if( $matrix * $v3 == $e3 * $v3){
                $flag3=1;
             }
           }
        }

        if($flag1==1 && $flag2 ==1 && $flag3==1){
           $flag=1;
        }else{
           $flag4 = $flag1+ $flag2 + $flag3;
           if($flag4 == 1 ){
              Value::Error("One of the eigenvectors is correct");
           }else{
              if($flag4 == 2 ){
                 Value::Error("Two of the eigenvectors are correct");
              }else{
                 Value::Error("None of the eigenvectors is correct");
              }
           }
        }
        return $flag;

    }
);

BEGIN_TEXT
The matrix \(A = \{display_matrix_mm( $matrix_disp )\}\)
has eigenvalues \( $e1 \), \( $e2 \) and \( $e3 \).
$BR
Find its eigenvectors.
$PAR
The eigenvalue \( $e1 \) is associated with eigenvector ( \{ $ev-> ans_rule(5)\}, \{$ev-> ans_rule(5)\}, \{$ev-> ans_rule(5)\}).
$PAR
The eigenvalue \( $e2 \) is associated with eigenvector  ( \{ $ev-> ans_rule(5)\}, \{$ev-> ans_rule(5)\}, \{$ev-> ans_rule(5)\}).
$PAR
The eigenvalue \( $e3 \) is associated with eigenvector  ( \{ $ev-> ans_rule(5)\}, \{$ev-> ans_rule(5)\}, \{$ev-> ans_rule(5)\}).

END_TEXT

ANS($ev->cmp() );

ENDDOCUMENT() ;




