# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 07/31/09


## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Partial derivatives)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('14.2')
## Section2('14.2')
## Section3('14.2')
## Problem1('37')
## Problem2('42')
## Problem3('45')
## KEYWORDS('derivatives', 'partial', 'functions', 'multivariable')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

Context("Numeric");
Context()->variables->add( w=>'Real', z=>'Real' );
$showPartialCorrectAnswers = 1;

$w0 = random(2,5,1);
$z0 = random(2,5,1);
$h  = random(0.01, 0.02, 0.01);
$f  = Compute("e^(w*ln(z))");

$f0   = $f->eval( w=>$w0, z=>$z0 );
$f0wp = $f->eval( w=>($w0+$h), z=>$z0 );
$f0wm = $f->eval( w=>($w0-$h), z=>$z0 );
$f0zp = $f->eval( w=>$w0, z=>($z0+$h) );
$f0zm = $f->eval( w=>$w0, z=>($z0-$h) );

$dfw0 = Compute( "($f0wp - $f0)/$h" );
$dfw1 = Compute( "($f0w - $f0wm)/$h" );
$dfw2 = Compute( "($f0wp - $f0wm)/(2*$h)" );

$dfz0 = Compute( "($f0zp - $f0)/$h" );
$dfz1 = Compute( "($f0z - $f0zm)/$h" );
$dfz2 = Compute( "($f0zp - $f0zm)/(2*$h)" );

$fw  = $f->D('w');
$fw0 = $fw->eval( w=>$w0, z=>$z0 );

$fz  = $f->D('z');
$fz0 = $fz->eval( w=>$w0, z=>$z0 );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Let \[ f(w,z) = $f. \]

$PAR
${BBOLD}(a)$EBOLD
Use difference quotients with \( h=$h \) to approximate each of the
partial derivatives:
$BR
\( f_w($w0, $z0) \approx \) \{ ans_rule(35) \}
$BR
\( f_z($w0, $z0) \approx \) \{ ans_rule(35) \}.

$PAR
${BBOLD}(b)$EBOLD
Now evaluate the partial derivatives exactly:
$BR
\( f_w($w0, $z0) = \) \{ ans_rule(35) \}
$BR
\( f_z($w0, $z0) = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($dfw0->cmp( checker=>sub {
		    my ( $c, $s, $a ) = @_;
		   return $s == $c || $s == $dfw1 || $s == $dfw2;
		   } ) );
ANS($dfz0->cmp( checker=>sub {
		    my ( $c, $s, $a ) = @_;
		   return $s == $c || $s == $dfz1 || $s == $dfz2;
		   } ) );
ANS($fw0->cmp() );
ANS($fz0->cmp() );

$w0h = $w0 + $h;
$z0h = $z0 + $h;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The difference quotient for evaluating \(f_w($w0,$z0)\) is
\[ f_w($w0,$z0) \approx \frac{f($w0+$h,$z0)-f($w0,$z0)}{$h} =
    \frac{e^{($w0h)\ln $z0} - e^{$w0 \ln $z0}}{$h} \approx
    $dfw0. \]
The difference quotient for evaluating \(f_z($w0,$z0)\) is
\[ f_z($w0,$z0) \approx \frac{f($w0,$z0 + $h)-f($w0,$z0)}{$h} =
    \frac{e^{($w0)\ln $z0h} - e^{$w0 \ln $z0}}{$h} \approx
    $dfz0. \]
(We could, of course, use a central or left-handed difference quotient
instead.)
$PAR
${BBOLD}(b)$EBOLD
Using the derivative formulas we get
\[ f_w = \frac{\partial f}{\partial w} = $fw, \]
\[ f_z = \frac{\partial f}{\partial z} = $fz, \]
so
\[ f_w($w0,$z0) = \ln($z0) e^{$w0 \ln($z0)} \approx $fw0, \]
and
\[ f_z($w0,$z0) = \frac{$w0}{$z0} e^{$w0 \ln($z0)} = $fz0. \]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
