## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/29/06

## DBsubject(Linear algebra)
## DBchapter(Determinants)
## DBsection(Applications)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "PGmatrixmacros.pl",
  "PGnumericalmacros.pl",
  "PGmorematrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->are(x=>'Real');

$a = random(2,3,1) * random(-1,1,2);

$s = random(-1,1,2);
$d = random(1,3,1) * $s; 
if ($a == $d) {$d = $s;}
$e = random(2,3,1) * $s;

$f = random(-2,2,4);
$g = non_zero_random(-2,2,4);

$a12free = $g * (1 - $a); 
$a31free = $f * ($a - $d);
$a32free = $f * $g * ($d - $a);
$g2 = $g * $g;
$mfe = - $f * $e;
$feg = $f * $g * $e;

$det = Compute("$a * (1 + x^2)($d + $e * x^2)");



# the inverse matrix

$ans11 = Compute("(1 + x^2 + $g*x^3)/($a + $a*x^2)");
$ans12 = Compute("(- $a12free - $g*x^2 - $g2*x^3)/($a + $a*x^2)");
$ans13 = Compute("0");
$ans21 = Compute("x^3/($a + $a*x^2)");
$ans22 = Compute("($a - $g*x^3)/($a + $a*x^2)");
$ans23 = Compute("0");

$k = $f*($d - $a);
$l = $f*($d + $e - $a);
$fe = $f*$e;
$ans31 = Compute("($k + $l*x^2 + $fe*x^4)/(($a + $a*x^2)*($d + $e * x^2))");

$fg = $f*$g;
$ans32 = Compute("$fg*(1/($d + $e * x^2) - 1/$a)");
$ans33 = Compute("1/($d + $e * x^2)");

$multians = MultiAnswer($ans11, $ans12, $ans13, $ans21, $ans22, $ans23, $ans31, $ans32, $ans33)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rrr} %s & %s & %s \\ %s & %s & %s \\ %s & %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $i (0..8) {
        $score[$i] = 0;
        if ($c[$i]->typeMatch($s[$i]) && $c[$i] == $s[$i]) { $score[$i] = 1; }
      }
      return [ @score ];
  }
);



Context()->texStrings;
BEGIN_TEXT

Find the determinant of the matrix
$BR
\{ mbox( '\(M=\)', 
display_matrix([["\($a - $g x^3\)", "\($a12free + $g x^2 + $g2 x^3\)", "0"], 
		["\(- x^3\)",       "\(1 + x^2 + $g x^3\)",            "0"], 
		["\($a31free + $mfe x^2\)", "\($a32free + $feg x^2\)", "\($d + $e x^2\)"]]), ',' ) \} 
$BR 
and use the adjoint method to find \(M^{-1}\).
$BR
$BR
\( {\rm det}(M)=\) \{ans_rule(20)\}, 
$BR
$BR
\{ 
mbox(
  "\( M^{-1} = \)",
  display_matrix([
  [$multians->ans_rule(25),$multians->ans_rule(25),$multians->ans_rule(25)], 
  [$multians->ans_rule(25),$multians->ans_rule(25),$multians->ans_rule(25)], 
  [$multians->ans_rule(25),$multians->ans_rule(25),$multians->ans_rule(25)], 
  ], align=>'ccc')
);
\}

END_TEXT
Context()->normalStrings;

ANS( $det->cmp );
ANS( $multians->cmp );

;
ENDDOCUMENT();       # This should be the last executable line in the problem.

