## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/29/06

## DBsubject(Linear algebra)
## DBchapter(Determinants)
## DBsection(Applications)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.3')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "PGmatrixmacros.pl",
  "PGnumericalmacros.pl",
  "PGmorematrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->are(t=>'Real');

$a11 = random(3,5,2) * random(-1,1,2);
$a12 = random(2,6,1) * random(-1,1,2);
$a21 = random(2,6,2) * random(-1,1,2);
$a22 = random(3,5,2) * random(-1,1,2);

$m = random(2,4,1);
$n = random(2,4,1);
if ($m==$n) {$n=$n+1;}
$s = $m+$n;

$det_num = $a11 * $a22 - $a12 * $a21;

$Adet = Compute("$det_num * e^($s*t)"); 

#$A = Matrix([
#["$a11 e^{$m t}", "$a12 e^{$n t}"], 
#["$a21 e^{$m t}", "$a22 e^{$n t}"]
#]);

$Acof11 = Compute("$a22 * e^($n*t)");    $Acof12 = Compute("- $a21 * e^($m*t)"); 
$Acof21 = Compute("- $a12 * e^($n*t)");  $Acof22 = Compute("$a11 * e^($m*t)"); 

$multians2 = MultiAnswer($Acof11, $Acof12, $Acof21, $Acof22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }
);

$Aadj11 = Compute("$a22 * e^($n*t)"); $Aadj12 = Compute("- $a12 * e^($n*t)");
$Aadj21 = Compute("- $a21 * e^($m*t)");  $Aadj22 = Compute("$a11 * e^($m*t)");

$multians3 = MultiAnswer($Aadj11, $Aadj12, $Aadj21, $Aadj22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }
);


$Ainv11 = Compute("$a22 * e^(- $m*t)/$det_num");
$Ainv12 = Compute("- $a12 * e^(- $m*t)/$det_num");
$Ainv21 = Compute("- $a21 * e^(- $n*t)/$det_num");  
$Ainv22 = Compute("$a11 * e^(- $n*t)/$det_num");

$multians4 = MultiAnswer($Ainv11, $Ainv12, $Ainv21, $Ainv22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'><tr><td> %s </td><td> %s </td></tr> <tr><td> %s </td><td> %s </td></tr></table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if ($c11->typeMatch($s11) && $c11 == $s11) { $score[0] = 1; }
      if ($c12->typeMatch($s12) && $c12 == $s12) { $score[1] = 1; }
      if ($c21->typeMatch($s21) && $c21 == $s21) { $score[2] = 1; }
      if ($c22->typeMatch($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }
);


Context()->texStrings;
BEGIN_TEXT

Let \(A=
\left\lbrack
\begin{array}{cc}
$a11 e^{$m t} & $a12 e^{$n t} \\
$a21 e^{$m t} & $a22 e^{$n t}
\end{array}
\right\rbrack
\).
$BR
$BR
(a) Find the determinant of \( A \).
$BR
\( \det (A) = \) \{ans_rule(20)\},  
$BR
$BR
(b) Find the matrix of cofactors of \( A \).
$BR
\{ 
mbox(
  "\( C = \)",
  display_matrix([
  [$multians2->ans_rule(15),$multians2->ans_rule(15)], 
  [$multians2->ans_rule(15),$multians2->ans_rule(15)]
  ], align=>'cc')
);
\}
$BR
$BR
(c) Find the adjoint of \( A \).
$BR
\{ 
mbox(
  "\( {\rm adj}(A) = \)",
  display_matrix([
  [$multians3->ans_rule(15),$multians3->ans_rule(15)], 
  [$multians3->ans_rule(15),$multians3->ans_rule(15)]
  ], align=>'cc')
);
\}
$BR
$BR
(d) Find the inverse of \( A \).
$BR
\{ 
mbox(
  "\( A^{-1} = \)",
  display_matrix([
  [$multians4->ans_rule(15),$multians4->ans_rule(15)], 
  [$multians4->ans_rule(15),$multians4->ans_rule(15)]
  ], align=>'cc')
);
\}
END_TEXT
Context()->normalStrings;

ANS($Adet->cmp);
ANS($multians2->cmp);
ANS($multians3->cmp);
ANS($multians4->cmp);

;




ENDDOCUMENT();       # This should be the last executable line in the problem.
