# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Simplify the following expression
#
#   b * a x^n y ^m              d * c x^p y^q
#   --------------------   -    ----------------
#      b x^(n-2) y^(m-1)        d x^(p-2) y^(q-2)
#
#
# This means that the problem simplifies to
#
#      ax^2y - cx^2y = (a-c) x^2y
#
# Obviously we choose a!=c
#
# Last edited: Hughes 9/6/13, Carl Yao 6/28/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions')
## KEYWORDS('exponent','simplify','fraction','divide')
## DBCCSS('8.EE.1')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextLimitedPolynomial.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});


$var1 = 'a';
$var2 = 'b';
Context()->variables->are($var1=>'Real',$var2=>'Real');

#   b * a x^n y ^m              d * c x^p y^q
#   --------------------   -    ----------------
#      b x^(n-2) y^(m-1)        d x^(p-2) y^(q-2)

# first term
$a = (random(0.5,6.0,1)+0.5)*2*random(-1,1,2); #non_zero_random(-13,13,1);
$b = (random(1.5,5.5,1)+0.5)*random(-1,1,2);
$n = random(4,20,1);
$m = random(3,$n-1,1);

$term1 = Formula("$b*$a*$var1^$n*$var2^$m/($b*$var1^($n-2)*$var2^($m-1))")->reduce;

# second term
$c = $a;
$c = non_zero_random(1,13,1) while($c==$a);
$d = (random(1.5,5.5,1)+0.5);
$p = random(4,20,1);
$q = random(3,$n-1,1);

$term2 = Formula("$d*$c*$var1^$p*$var2^$q/($d*$var1^($p-2)*$var2^($q-1))")->reduce;

# reduction check
$ans = Formula("($a-$c)*$var1^2*$var2")->reduce->reduce;


Context("LimitedPolynomial-Strict");
Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1});

Context()->variables->are($var1=>'Real',$var2=>'Real');
$ans = Formula($ans);


##############################################
TEXT(beginproblem());

BEGIN_PGML

Simplify the following expression

    [`\displaystyle [$term1] - [$term2]= `][________________________________]

END_PGML


ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 



BEGIN_PGML_SOLUTION
We perform the division on each term first, and the subtract the result as follows

    [`
        \begin{aligned}
            \displaystyle [$term1] - [$term2] & = [$a][$var1]^2[$var2]-[$c][$var1]^2[$var2] \\
                                             & = [$ans]
        \end{aligned}
`]

END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
