# WeBWorK problem written by Chris Hughes
# Portland Community College
#
# Template:
# Here is an expression with negative exponents.
#
#      1
#   ------
#   a x^-n
#
# Rewrite the expression using only positive exponents.
#
# a is integer on [2,20]. 
#
# n is integer on [2,20].
# 
# Last updated: Hughes 9/5/13, Carl Yao, 7/01/13
# ENDDESCRIPTION

## DBCCSS('8.EE.1')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## Level(2)
## MO(1)
## KEYWORDS('exponent','negative','fraction','simplify')
## MLT(Algebra_PropExp_PropExp)

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "contextLimitedPolynomial.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl"
);

##############################################

Context("LimitedPolynomial-Strict");
$var = RandomVariableName(type=>'variable');
Context()->variables->are($var=>'Real');

Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});

# the base
$a = random(2,50,1);

# the exponent
$n=$a;
$n = random(2,20,1) while($n==$a);

Context()->variables->set($var=>{limits=>[0.25,0.4]});
$ans = Formula("$var^$n/$a");

##############################################
TEXT(beginproblem());

BEGIN_PGML
Here is an expression with negative exponents.

    [`
        \displaystyle\frac{1}{[$a][$var]^{[$n*(-1)]}}=
     `][____________]

Rewrite the expression using only positive exponents.

END_PGML

ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
)); 


BEGIN_PGML_SOLUTION

We evaluate the expression by remembering that [`x^{-n}`] is the same thing as [`\frac{1}{x^n}`] for any non-zero,
real value of [`x`].

    [`
         \displaystyle\frac{1}{[$a][$var]^{[$n*(-1)]}} = \displaystyle\frac{[$var]^{[$n]}}{[$a]}
    `]

Notice that the [`[$a]`] in the denominator does _not_ go in the numerator because it is not raised to the power [`[$n*(-1)]`]; in other words, 
[`[$a][$var]^{[$n*(-1)]}\ne ([$a][$var])^{[$n*(-1)]}`].
END_PGML_SOLUTION
##############################################


ENDDOCUMENT();
