##DESCRIPTION
##  Number theory: Diophantine equation applied to groceries
##ENDDESCRIPTION


## DBsubject(Number Theory)
## DBchapter(Diophantine equations)
## DBsection(Solving Diophantine equations)
## Date(07/15/2018)
## Institution(University of Southern Mississippi)
## Author(John Perry)
## MO(1)
## KEYWORDS('number theory', 'diophantine', 'equation')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

@primes = ( 7, 11, 13, 17 );
$d = @primes[ random(0,3,1) ];
$a = $d * random(2,6);
$b = $a;
while ($a == $b) {
  $b = $d * random(2,6);
}
$x0 = random(101,300,1);
$y0 = random(50,300,1);
if ($x0 == 200) { $x0 -= 1; }
if ($y0 == 100 || $y0 == 200) { $y -= 1; }
$c = ($a * $x0 + $b * $y0)/100;
$ma = MultiAnswer($x0,$y0,$a,$b)->with(
  singleResult => 0,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $x1, $y1, $x2, $y2 ) = @{ $student };
    my ( $z1 ) = ($a * $x1 + $b * $y1) / 100;
    my ( $z2 ) = ($a * $x2 + $b * $y2) / 100;
    if ( ($a * $x1 + $b * $y1) / 100 == $c ) {
      if ( ($x1 != $x2) && ($a * $x2 + $b * $y2) / 100 == $c ) { return [1,1,1,1]; }
      else { return [1,1,0,0]; }
    } else {
      if ( ($a * $x2 + $b * $y2) / 100 == $c ) { return [0,0,1,1]; }
      else { return [0,0,0,0]; }
    }
  }
);
##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

A grocer orders apples and bananas at a total cost of $$c.
If the apples cost $a cents each and the bananas $b cents each,
how many of each type of fruit did he order? (There can be more than one answer.)
$PAR
<table>
<tr><td>apples:</td><td>\{ $ma->ans_rule(5) \}</td></tr>
<tr><td>bananas:</td><td>\{ $ma->ans_rule(5) \}</td></tr>
</table>
$PAR
Also give a second solution. (There can be more than one answer.)
$PAR
<table>
<tr><td>apples:</td><td>\{ $ma->ans_rule(5) \}</td></tr>
<tr><td>bananas:</td><td>\{ $ma->ans_rule(5) \}</td></tr>
</table>
END_TEXT
Context()->normalStrings;

ANS( $ma->cmp() );

ENDDOCUMENT();
