##DESCRIPTION
##  Number theory: simple linear congruence
##ENDDESCRIPTION


## DBsubject(Number Theory)
## DBchapter(Diophantine equations)
## DBsection(Solving linear congruence)
## Date(07/15/2018)
## Institution(University of Southern Mississippi)
## Author(John Perry)
## MO(1)
## KEYWORDS('number theory', 'linear', 'congruence')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "parserPopUp.pl",
   "scaffold.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

$has_solution = PopUp( [ "?", "yes", "no" ], "yes" );

@primes = ( 11, 13, 17, 19, 23, 29 );
$d1 = @primes[ random(0,2,1) ];
$d2 = @primes[ random(3,5,1) ];
$d = $d1 * $d2;
$a = $d * random(2,7,1);
$b = $a;
while ($a == $b || ($a % $b == 0) || ($b % $a == 0)) {
  $b = $d * random(2,7,1);
}

$x0 = random(-9,-6,1);
$y0 = random(6,9,1);
$c = $a * $x0 + $b * $y0;

$x_d = $b/$d;
$y_d = $a/$d;

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

<b>This is a three-part problem. Once you answer the first part correct, a second part will appear. Once you answer that correctly, a third part will appear.</b>
$PAR
Consider the linear congruence \[ $c \equiv $a x \pmod{$b} . \]

END_TEXT
Context()->normalStrings;

Scaffold::Begin();

Section::Begin("Part 1");

Context()->texStrings;
BEGIN_TEXT

Does this linear congruence have integer solutions?
\{ $has_solution->menu() \}

END_TEXT
Context()->normalStrings;

ANS( $has_solution->cmp() );

Section::End();

Section::Begin("Part 2");

Context()->texStrings;

BEGIN_TEXT

$PAR
Good! What is one possible solution?
$PAR
\( x\equiv \) \{ ans_rule(5) \}

END_TEXT
Context()->normalStrings;

ANS( Compute($x0)->cmp(
  checker => sub {
    my ( $correct, $student, $hash ) = @_;
    return ( ( $c - ( int( $student ) * $a ) ) % $b == 0);
  }
) );

Section::End();

Section::Begin("Part 3");

Context()->texStrings;

BEGIN_TEXT
$PAR
To find other solutions, we should add or subtract integer multiples of what number? Give the smallest positive integer possible.
$PAR
Add (or subtract) \{ ans_rule(5) \} to \( x \).

END_TEXT
Context()->normalStrings;

ANS( Compute($x_d)->cmp() );

Section::End();

Scaffold::End();

ENDDOCUMENT();
