##DESCRIPTION
#          
# File Created: 5/31/2000
# Last Modified: 5/31/2000
# Problem Author: Joseph Neisendorfer
# WeBWorK Entry: Victoria Sweetser
# Location: University of Rochester
#
##ENDDESCRIPTION   

  ## tsuda tagged and PAID on 2-20-2004

## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Lines with planes)
## Date(6/3/2002)
## MLT(TwoPlanes)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('6')
## Section1('12.5')
## Problem1('33')
## KEYWORDS('Vector','Multivariable','Geometry')
         
DOCUMENT();	# This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
              
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric");
Context()->variables->add(t=>'Real');
$a = random(1, 5, 1);
$b = random(1, 5, 1);
$c = random(1, 5, 1);

$ansX = Formula("t*$c/sqrt( ($a)^2 + ($c)^2 )");
$ansY = Formula("1/$b");
$ansZ = Formula("t*(-$a)/sqrt( ($a)^2 + ($c)^2 )");

$multians = MultiAnswer($ansX, $ansY,$ansZ)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stuX, $stuY,$stuZ ) = @{$student};
      my ( $ansX, $ansY, $ansZ ) = @{$correct};
      #Check if function lies on both planes:
      $FstFn = Formula("$a*$stuX + $b*$stuY + $c*$stuZ");
      $SndFn = Formula("$a*$stuX + $c*$stuZ")->with(tolType => 'absolute',
	            tolerance => 0.0001);

      if(($FstFn == Formula("1"))&&($SndFn == Formula("0"))) {
           return 1; }
      elsif(($FstFn == 1)||($SndFn == 0)) {
           Value->Error("Your function must be on $BBOLD both $EBOLD planes!");}
      else { return 0; }
  }
);

Context()->texStrings;
BEGIN_TEXT
Consider the planes 
\( $a\!x + $b\!y + $c\!z = 1 \) and 
\( $a\!x + $c\!z = 0. \)
$PAR
(A) Find the unique point P on the y-axis which is on both planes.
(\{ ans_rule(5) \}, \{ ans_rule(5) \}, \{ ans_rule(5) \})
$PAR

(B) Find a unit vector \( \bf{u} \) with positive first coordinate that
is parallel to both planes.$BR
\{ ans_rule(15) \} \( \mathbf{i} \) +
\{ ans_rule(15) \} \( \mathbf{j} \) +
\{ ans_rule(15) \} \( \mathbf{k} \)
$PAR
(C) Use parts (A) and (B) to find a vector equation for the line of
intersection of the two planes,\(\bf{r(t) =}\)$BR
\{ $multians->ans_rule(15) \} \( \bf{i} \) +
\{ $multians->ans_rule(15) \} \( \bf{j} \) +
\{ $multians->ans_rule(15) \} \( \bf{k} \) $BR $BR
$BBOLD NOTE: $EBOLD Part (C) is counted as a single question. $BBOLD All three $EBOLD components must be filled in to receive credit. $BR
There are $BITALIC many $EITALIC different correct answers to this question.

END_TEXT
Context()->normalStrings;

ANS(Compute("0")->cmp);   
ANS(Compute("1/$b")->cmp);
ANS(Compute("0")->cmp);

ANS(Compute("$c/sqrt( $a^2 + $c^2 )")->cmp);
ANS(Compute("0/sqrt( $a^2 + $c^2 )")->cmp);
ANS(Compute("(-$a)/sqrt( $a^2 + $c^2 )")->cmp);

ANS($multians->cmp);

ENDDOCUMENT();	# This should be the last executable line in the problem.
