# DESCRIPTION
# Factor an expression of the form x + b sqrt(x) + c.
# ENDDESCRIPTION

## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring trinomials)
## Institution(University of Missouri-Columbia)
## Level(3)
## KEYWORDS('factoring', 'factor', 'factorization', 'substitution')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "MUHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
$var = list_random('x','y','z','w','p','q','r','s','t','u','v');
Context()->variables->are($var=>"Real", uc($var)=>"Real");
LimitedPowers::OnlyPositiveIntegers();

do {
  $x = random(-1,1,2)*list_random(2,3,5,7,11);
  $y = random(-1,1,2)*list_random(2,3,5,7,11);
} until (gcd($x,$y) == 1);
$b = $x+$y;
$c = $x*$y;

$fac1 = "\sqrt{$var} + $x";    $afac1 = Formula("sqrt($var) + $x")->reduce;
$fac2 = "\sqrt{$var} + $y";    $afac2 = Formula("sqrt($var) + $y")->reduce;

$expn = nicestring([1,$b,$c], [$var, "\sqrt{$var}", ""]); #"$var +$b \sqrt{$var} + $c";
Context()->flags->set(limits=>[2,5]);
$answer = MultiAnswer($afac1, $afac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1$var^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01$var^{500}") if ($s2->type eq 'String'); #bogus answer
    if (($c1 == $s1 && $c1 == $s2) || ($c2 == $s1 && $c2 == $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c1 == -$s2) || ($c2 == -$s1 && $c2 == -$s2)) {
      return [1,0];
    } elsif (($c1 == $s1 && $c2 == $s2) || ($c1 == $s2 && $c2 == $s1)) {
      return [1,1];
    } elsif (($c1 == -$s1 && $c2 == -$s2) || ($c1 == -$s2 && $c2 == -$s1)) {
      return [1,1];
    } elsif (($c1 == $s1 && $c2 != $s2) || ($c2 == $s1 && $c1 != $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c2 == $s2) || ($c2 == -$s1 && $c1 == $s2)) {
      return [0,1];
    } elsif (($c1 == $s2 && $c2 != $s1) || ($c2 == $s2 && $c1 != $s1)) {
      return [0,1];
    }
    return [0,0];
  }
);
$texanswer = "\(($fac1)($fac2)\)";


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
Context()->texStrings;
  BEGIN_TEXT
  $BR
  ${BBOLD}Help:${EBOLD} \{ MUHelp("sqrt") \}
  $BR
  END_TEXT
Context()->normalStrings;

COMMENT('Desc: Factor an expression of the form x + b sqrt(x) + c.');
ENDDOCUMENT();
