# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION




##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextRationalFunction.pl",
  "bizarroArithmetic.pl",
  "PCCmacros.pl",
  "PGML.pl",
  
);

##############################################
TEXT(beginproblem());


Context("Numeric");
$index = random(0,4,1);
$var[0] = ('x','a','s','p','m')[$index];
$var[1] = ('y','b','t','q','n')[$index];
$var[2] = ('z','c','u','r','k')[$index];
Context()->variables->are($var[0]=>'Real',$var[1]=>'Real',$var[2]=>'Real');

$a = random(3,10,1);
$b = random(2,10,1);
$num = Formula("$var[0]*$var[1]^2/($a $var[2])");
$den = Formula("($var[0])/($b*$var[1]*$var[2])");

$as = $a/gcd($a,$b);
$bs = $b/gcd($a,$b);

Context("RationalFunction");
Context()->operators->set(
'*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
' *' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
'/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
);


Context()->variables->are($var[0]=>'Real',$var[1]=>'Real',$var[2]=>'Real');
$ans = Formula("($bs*$var[1]^3)/$as")->reduce;

$evaluator = $ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      my $OK = ($correct == $student);
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      Value::Error("Your answer is equivalent to the correct answer, but not in the expected form.
                    Maybe it needs to be simplified, factored, expanded, have its denominator rationalized, etc.") unless $OK;
      return $OK;
});

##############################################

BEGIN_PGML


Simplify this expression. [@KeyboardInstructions('Just enter the simplified expression. Do not try to enter its restricted domain.')@]**

[``\frac{[$num]}{[$den]}=``][__]{$evaluator}{width=>40}

END_PGML


##############################################


ENDDOCUMENT();
