## DESCRIPTION
## Multivariable Calculus -average value
## ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Applications of double integrals)
## Level(3)
## Date('7/21/2023')
## Author('Hunter Lehmann')
## Institution('Georgia Institute of Technology')
## TitleText1('Thomas' Calculus: Early Transcendentals')
## EditionText1('13')
## AuthorText1('Hass, Heil, Weir, Bogacki')
## Section1('15.3')
## Problem1('')
## KEYWORDS('calculus','average value','iterated integral')

########################################################################

DOCUMENT();   

loadMacros(
	'PGstandard.pl', 
	'PGML.pl', 
	'parserMultiAnswer.pl', 
	'PGcourse.pl'
);

Context()->variables->are(
    x  => 'Real',
    dx => 'Real',
    y  => 'Real',
    dy => 'Real'
);
Context()->flags->set(reduceConstants => 0);

#parameter setup

$a=random(2,5,1);
$b=random(1,6,1);
$c=random(1,15,1);

#bounds
$uby=Formula("$a-x/2");
$ubx=Formula("2*$a-2y");

# integrand and volume

$f = Formula("$b x^2-y+$c");
$area = Formula("$a^2");
$value = Formula("2*$a^2*$b/3-$a/3+$c");

# differentials and limits of integration
# Case 0, element 0 of each array below, is if the order of integration is dx dy
# Case 1, element 1 of each array below, is if the order of integration is dy dx
# 'id' and 'od' stand for inner and outer differential

@id = (Formula('dx'), Formula('dy'));    # (case 0, case 1)
@od = (Formula('dy'), Formula('dx'));    # (case 0, case 1)

# A = outer integral, lower limit
# B = outer integral, upper limit
# C = inner integral, lower limit
# D = inner integral, upper limit

@A = (Formula(0), Formula(0));    # (case 0, case 1)
@B = (Formula($a), Formula("2*$a"));    # (case 0, case 1)
@C = (Formula(0), Formula(0));    # (case 0, case 1)
@D = ($ubx, $uby);    # (case 0, case 1)

$multians = MultiAnswer($f, $id[0], $od[0], $A[0], $B[0], $C[0], $D[0])->with(
    singleResult => 1,
    checker      => sub {
        my ($correct, $student, $self) = @_;
        my ($fstu, $idstu, $odstu, $Astu, $Bstu, $Cstu, $Dstu) =
            @{$student};
        if (
            (
                $f == $fstu
                && $id[0] == $idstu
                && $od[0] == $odstu
                && $A[0] == $Astu
                && $B[0] == $Bstu
                && $C[0] == $Cstu
                && $D[0] == $Dstu
            )
            || ($f == $fstu
                && $id[1] == $idstu
                && $od[1] == $odstu
                && $A[1] == $Astu
                && $B[1] == $Bstu
                && $C[1] == $Cstu
                && $D[1] == $Dstu)
            )
        {
            return 1;
        } elsif (
            (
                $f == $fstu
                && $id[0] == $idstu
                && $od[0] == $odstu
                && ($A[0] != $Astu || $B[0] != $Bstu)
                && $C[0] == $Cstu
                && $D[0] == $Dstu
            )
            || ($f == $fstu
                && $id[1] == $idstu
                && $od[1] == $odstu
                && ($A[1] != $Astu || $B[1] != $Bstu)
                && $C[1] == $Cstu
                && $D[1] == $Dstu)
            || ($f == $fstu
                && $id[0] == $idstu
                && $od[0] == $odstu
                && $A[0] == $Astu
                && $B[0] == $Bstu
                && ($C[0] != $Cstu || $D[0] != $Dstu))
            || ($f == $fstu
                && $id[1] == $idstu
                && $od[1] == $odstu
                && $A[1] == $Astu
                && $B[1] == $Bstu
                && ($C[1] != $Cstu || $D[1] != $Dstu))
            )
        {
            $self->setMessage(1, 'Check your limits of integration.');
            return 0.94;
        } elsif (
            (
                $f == $fstu
                && $id[0] == $idstu
                && $od[0] == $odstu
                && ($A[0] != $Astu || $B[0] != $Bstu)
                && ($C[0] != $Cstu || $D[0] != $Dstu)
            )
            || ($f == $fstu
                && $id[1] == $idstu
                && $od[1] == $odstu
                && ($A[1] != $Astu || $B[1] != $Bstu)
                && ($C[1] != $Cstu || $D[1] != $Dstu))
            )
        {
            $self->setMessage(1,
                'Check your limits of integration and order of integration.'
            );
            return 0.47;
        } else {
            return 0;
        }
    }
);

BEGIN_PGML
Set up a double integral in rectangular coordinates for calculating the average value of the function [`f(x,y) = [$f]`] over the region bounded by the lines
[`x=0, y=0,`] and [`y=[$a]-\dfrac{x}{2}`].

_Instructions:_ Compute the area of the region and enter it in the first answer box. 
Enter the integrand in the second answer box . Depending
on the order of integration you choose, enter _dx_ and _dy_ in either order into
the third and fourth answer boxes with only one _dx_ or _dy_ in each box . Then,
enter the limits of integration and evaluate the integral to find the average value.

Area = [_]{$area}{10}

Average value = [``\frac{1}{\text{area}}\int_A^B \int_C^D``] [_]{$multians}{10} [_]{$multians}{5} [_]{$multians}{5}, where

A = [_]{$multians}{10}

B = [_]{$multians}{10}

C = [_]{$multians}{10}

D = [_]{$multians}{10}

Average value = [_]{$value}{10}
END_PGML

BEGIN_PGML_SOLUTION

We give the solution with the integration order [`dx\ dy.`]  The area of the triangular region of integration is [`\frac{1}{2}([$a])([@ $a*2@])=[@$a**2@]`].
Thus the average value is

>>[``\begin{aligned}
      f_{avg} &= \frac{1}{[@$a**2@]} \int_0^{[$a]} \int_0^{[@2*$a@]-2y} [$f]\ dx\ dy \\
       &=  \frac{1}{[@$a**2@]} \int_0^{[$a]} \frac{[$b]([@2*$a@]-2y)^3}{3}-([@2*$a@]-2y)(y+[$c])  \ dy \\
	   &= \frac{1}{[@$a**2@]}\left( -\frac{[$b]([@2*$a@]-2y)^4}{24}-[$a]y^2+\frac{2}{3}y^3-[@ 2*$a*$c@]y+[$c]y^2\right) |_0^[$a] \\
	   &= [$value]\\
	   &= \frac{[@ 2*$a**2*$b-$a+3*$c @]}{3}
      \end{aligned} ``]<<

END_PGML_SOLUTION

COMMENT('Allows integration in either order.');

ENDDOCUMENT();