#Use laws of exponents with rational exponents.
# ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Date(6/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)

 
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
  "contextFraction.pl",
);

######################################
TEXT(beginproblem());
Context("Fraction");
$var = list_random('a','b','c','m','n','p','y','z');
Context()->variables->add($var=>'Real');
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);
#Context()->flags->set(bizarroPow=>1);

do{$n1 = list_random(2, 4, 8);
$d1 = list_random(3, 5, 9);
$n2 = 1;
$d2 = $d1;

$e1 = Fraction($n1, $d1)->reduce;
$e2 = Fraction($n2, $d2)->reduce;
$e3 = Fraction($n1*$d2 + $n2*$d1, $d1*$d2)->reduce;
$e4 = Fraction($n1 + $d1, $d1)->reduce;}
until($e3 !=1);

#Answer
$ga = Formula("$var^($e3)+$var^($e4)" );

#####################################################

BEGIN_PGML
Multiply by applying the laws of exponents. Write your answers with positive rational exponents only.


[``{[$var]}^{[$e1]}({[$var]}^{[$e2]}+[$var]) =``][__________]


END_PGML
####################################################

ANS($ga -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

####################################
BEGIN_PGML_SOLUTION

[``{[$var]}^{[$e1]}({[$var]}^{[$e2]}+[$var])``]

[``[$var]^{[$e1]}\cdot [$var]^{[$e2]} + [$var]^{[$e1]}\cdot [$var]``]

[``[$var]^{[$e1]+[$e2]} + [$var]^{[$e1]+1}``]

[``[$ga]``]
END_PGML_SOLUTION

ENDDOCUMENT();