#Use laws of exponents with rational exponents.
# ENDDESCRIPTION
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Date(6/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)

 
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
  "contextFraction.pl",
);

######################################
TEXT(beginproblem());
Context("Fraction");
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);
#Context()->flags->set(bizarroPow=>1);

#Part A: Use the product rule.
$d1a = random(2, 8);
$n1a = 1;
$d2a = random(2, 8);
do{
$n2a = random(2, 8);
}until(($n1a/$d1a)<($n2a/$d2a));

#exponents: EXP 2 is bigger than EXP 1 and EXP 1 is negative
$e1a = Fraction($n1a, $d1a)->reduce;
$e2a = Fraction($n2a, $d2a)->reduce;
$e3a = Fraction(-$n1a*$d2a + $n2a*$d1a, $d1a*$d2a)->reduce;

#Answer
$ga = Formula("x^($e3a)" );

#Part B:  Use the power rule.
$d1b = random(2, 10, 2);
$n1b = random(1, 9, 2);
$d2b = random(2, 10);
$n2b = random(2, 10);
$e1b = Fraction($n1b,$d1b)->reduce;
$e2b = Fraction($n2b,$d2b)->reduce;
$e3b = Fraction($n1b*$n2b, $d1b*$d2b)->reduce;

#Answer
$gb = Formula("x^$e3b");

#Part C: Use the quotient rule
$d1c = list_random(3, 5, 7, 9);
$n1c = list_random(2, 4, 6, 8);
do{
$e2c = random(2, 5);
}until($e2c > ($n1c/$d1c));
$e1c = Fraction($n1c, $d1c);

$e3c = Fraction($e2c*$d1c - $n1c, $d1c)->reduce;

#Answer
$gc = Formula("x^($e3c)");

#####################################################

BEGIN_PGML
Simplify by applying the laws of exponents. Write your answers with positive rational exponents only.


[``{x}^{-[$e1a]}\cdot{x}^{[$e2a]} =``][__________]


[``({{x}^{-[$e1b]}})^{-[$e2b]} = ``][_________]


[``\frac{{x}^{[$e2c]}}{{x}^{[$e1c]}} = ``][_________]


END_PGML
####################################################

ANS($ga -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

ANS($gb -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

ANS($gc -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

################################################################
BEGIN_PGML_SOLUTION
To multiply expressions with exponents, we add the exponents.  You will need to find the LCD in order to do the addition.

[``{x}^{-[$e1a]}\cdot{x}^{[$e2a]}``]

[`` -[$e1a] + [$e2a] = [$e3a]``]

The answer is: [``x^[$e3a]``]

The power rule applies to the next problem--to raise an expression with an exponent to a power, we multiply the exponents.  You don't need an LCD to multiply fractions.

[``({{x}^{-[$e1b]}})^{-[$e2b]}``]

[``-[$e1b] \cdot -[$e2b] = [$e3b]``]

The answer is: [``x^{[$e3b]}``]

The quotient rule says we subtract exponents when dividing.  You will need an LCD to do the subtraction.

[``\frac{{x}^{[$e1c]}}{{x}^{[$e2c]}}``]

[``[$e2c] - [$e1c] = [$e3c]``]

The answer is;  [``x^{[$e3c]}``]
END_PGML_SOLUTION

ENDDOCUMENT();

