#Use laws of exponents with rational exponents.
# ENDDESCRIPTION
## DBsubject()
## DBchapter()
## DBsection()
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## Date(6/2019)
## KEYWORDS()

 
DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGcourse.pl",
  "bizarroArithmetic.pl",
  "contextFraction.pl",
);

######################################
TEXT(beginproblem());
Context("Fraction");
Context()->variables->add(y=>'Real', a =>'Real', b=>'Real');
Context()->flags->set(reduceConstants=>0);
Context()->flags->set(limits=>[1,10]);
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
'**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
'^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
);
#Context()->flags->set(bizarroPow=>1);

#Part A: Use the power rule
$a = random(2, 5);
$acube = $a**3;
$d1 = list_random(2, 4, 8, 10);
$n1 = list_random(1, 3, 6);
$d2 = 3;
do{
$n2 = list_random(1, 2, 4);
}until(($n2/$d2) != ($d1/$n1));

$e1 = Fraction($n1, $d1)->reduce;
$e2 = Fraction($n2, $d2)->reduce;
$e3 = Fraction($n1*$n2, $d1*$d2)->reduce;
$a3 = $a**$n2;

#Answer
$ga = Formula("$a3*y^($e3)" );

#Part B:  Use the power rule.
$db1 = random(2, 7);
$nb2 = $db1;
do{
$db2 = random(2, 7);
}until($db1 != $db2);

$eb1= Fraction(1, $db1)->reduce;
$eb2 = Fraction($db1 - 1, $db1)->reduce;
$eb3 = Fraction($nb2, $db2)->reduce;

$eb4 = Fraction(1, $db2)->reduce;
if ( ($db1 - 1) == $db2 ) {
$gb = Formula("a^($eb4)*b");
} else {
$eb5 = Fraction($db1 - 1, $db2)->reduce;
$gb = Formula("a^($eb4)*b^($eb5)");
}



#####################################################

BEGIN_PGML
Simplify by applying the laws of exponents. Write your answers with positive rational exponents only.


[``{([$acube]*{y}^{[$e1]})}^{[$e2]} =``][__________]


[``({{a}^{[$eb1]}{b}^{[$eb2]}})^{[$eb3]} = ``][_________]


END_PGML
####################################################

ANS($ga -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));

ANS($gb -> cmp(
checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroDiv=>1,bizarroPow=>1);
delete $correct->{test_values};
delete $student->{test_values};
my $OK = (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroDiv=>0,bizarroPow=>0);
Value::Error("Your answer is equivalent, but please use the proper format") unless $OK;
return $OK;
}));
###############################################################
BEGIN_PGML_SOLUTION
[``{([$acube]*{y}^{[$e1]})}^{[$e2]} = [$acube]^{[$e2]}y^{[$e1]\cdot [$e2]} = [$ga]``]


[``({{a}^{[$eb1]}{b}^{[$eb2]}})^{[$eb3]} = a^{[$eb1] \cdot [$eb3]} b^{[$eb2]\cdot[$eb3]}= [$gb]``]
END_PGML_SOLUTION

ENDDOCUMENT();
