## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Optimization - general)
## Date(5/29/2005)
## Institution(ASU)
## Author(Jeff Holt)
## MLT(Optimization_Box_SquareCut)
## MLTleader(1)
## Level(5)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('5')
## Section1('4.7')
## Problem1('')
## TitleText2('Calculus: Early Transcendentals')
## AuthorText2('Stewart')
## EditionText2('6')
## Section2('4.7')
## Problem2('')
## KEYWORDS('derivative' 'extrema' 'optimization','Optimization Problems','Optimization' 'Maximum' 'Minimum')



DOCUMENT();        # This should be the first executable line in the problem.


loadMacros(
                "PG.pl",
                "PGgraphmacros.pl",
                "freemanMacros.pl",
                "PGbasicmacros.pl",
                "PGchoicemacros.pl",
                "PGanswermacros.pl",
                "PGauxiliaryFunctions.pl",
                "MathObjects.pl",
                "PGunion.pl",            # Union College utilities
                "parserMultiAnswer.pl",
                "PGcourse.pl",           # Customization file for the course
                "PGstandard.pl",
                "scaffold.pl"
               );

## Show partial correct answers
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;



###########################################
#
#  The scaffold
#

Scaffold::Begin(
#      can_open => "when_previous_correct", 
#      is_open  => "always" ,
#      instructor_can_open => "always", 
#      after_AnswerDate_can_open =>"when_previous_correct", 
#      hardcopy_is_open  => "always",
    );

Context ("Numeric");
Context() -> variables -> add(y => 'Real');

# Setup

$a = random(20,30,1);
$b = random(1,8,1);
$c = random(5,25,1);
$v= Formula("x**2+y**2");
$B=($a)**2;
$t=Formula("(($a)**2-x**2)**(1/2)");
$u=(($a)**2-($c)**2)**(1/2);
$p=Formula("-x/y");
$n=-$c/$u*$b;
$A=Formula("x/2*(($a)**2-x**2)**(1/2)");
$g=Formula("1/2*(($a)**2-x**2)**(1/2)-x**2/(2*(($a)**2-x**2)**(1/2))");
$h=Formula("x/y");
$k=(($b)*($u)-($c)*($n))/($a)**2;


$multians1 = MultiAnswer($v, $B)->with(
  singleResult => 0,
  checkTypes => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ($correct,$student,$self) = @_;
      my ($s1, $s2) = @{$student};
      my ($c1, $c2) = @{$correct};

      return (1,1) if ($c1 == $s1 and $c2 == $s2) or
                      ($c1 == $s2 and $c2 == $s1);
      return (1,0) if $c1 == $s1 or $c2 == $s1;
      return (0,1) if $c1 == $s2 or $c2 == $s2; 
      return (0,0);
   }
);



TEXT(EV2(<<EOT));
A ladder $a feet long is leaning against the wall of a house. The base of the ladder is pulled away from the wall at a rate of $b feet per second.
$PAR
Use this image to find the relation between \(x\), \(y\) and ladder using the Pythagorean theorem.
$PAR \{ image( "Px.png", width=>400, height=>400) \} $BR
EOT








####################################################
Section::Begin("Part 1: State the Pythagorean theorem for this problem as function of x,y and the ladder $a feet.");

BEGIN_TEXT
$BR
The Pythagorean theorem for this problem is: \{$multians1->ans_rule(20)\}\(=\)\{$multians1->ans_rule(10)\}

END_TEXT

ANS($multians1 -> cmp);



Section::End();




####################################################
Section::Begin("Part 2: Find the value of \(y\) using Part 1.");

BEGIN_TEXT

$BR
 \(y =\) \{ans_rule(20)\}
END_TEXT

ANS($t->cmp);

Section::End();



####################################################

Section::Begin("Part 3: Find the rate of change of the top of the ladder, \(\displaystyle \frac{dy}{dt}\), using part Part 1.");

BEGIN_TEXT

$BR
 \(dy/dt =\) \{ans_rule(20)\}\(dx/dt\)
END_TEXT

ANS($p-> cmp );

Section::End();



####################################################

Section::Begin("Part 4: Find the velocity of the top of the ladder using the value of \(x\), \(y\) and \(dx/dt\) when \(x\) is $c feet.");

BEGIN_TEXT

$BR
 \(Velocity =\) \{ans_rule(20)\}
END_TEXT

ANS(num_cmp($n) );

Section::End();



####################################################

Section::Begin("Part 5: Find the area of the right triangle as a function of \(x\).");

BEGIN_TEXT

$BR
 \(A(x)=\) \{ans_rule(20)\}
END_TEXT

ANS($A-> cmp );

Section::End();



####################################################

Section::Begin("Part 6: Find the first derivative of the area with respect to time \(t\).");

BEGIN_TEXT

$BR
 \(dA/dt =\) \{ans_rule(50)\}\(dx/dt\)
END_TEXT

ANS($g-> cmp );

Section::End();



####################################################

Section::Begin("Part 7: Find \(\tan(\theta)\) as a function of \(x\) and \(y\).");

BEGIN_TEXT

$BR
 \(tan(\theta) =\) \{ans_rule(20)\}

$BR Use this image to find \( \tan(\theta) \) as a function of \(x\) and \(y\).
$PAR \{ image( "Px1.png", width=>400, height=>400) \} $BR
END_TEXT

ANS($h-> cmp );


Section::End();



####################################################
Section::Begin("Part 8: Find the rate of change of theta with respect to \(t\).");

BEGIN_TEXT

$BR
 \(d \theta /dt =\) \{ans_rule(50)\}
END_TEXT

ANS(num_cmp($k) );

Section::End();

Scaffold::End();



ENDDOCUMENT();  # This should be the last executable line in the problem.
