#DESCRIPTION
##Type of
#ENDDESCRIPTION

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "contextLeadingZero.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$ans1="7^4"; 
$ans2="2^2\times 3^5\times y^2)";
$ans3=8;
$ans4=32;
$ans5=64;
$ans6=64;
$ans7="2^6";
$ans8="add";
$a=random(3,7,2);
$b=random(4,6,2);
$c=random(3,6,1);
$d=$b+$c;
$e=random(3,5,1);
$f=random(2,4,1);
$g=random(2,6,1);
$h=random(3,7,1);
$ans9="$a^$d";
$i=$e+$g;
$j=$f+$h;
$ans10="($a^$i) ($b^$j)";
BEGIN_TEXT
$PAR  
The parts of this problem concern exponentiation.  You know
that \(3^5\)  is the product of five 3's, that is,
 \(3^5=3\times 3\times 3\times 3\times 3=243\).  $SPACE
 We will say that \(3^5\) is the 
${BITALIC}base-exponent form${EITALIC} of that number, i.e,
\( 3 \) is the base and \( 5 \) is the exponent. To avoid confusion,
when we have two or more numbers in base-exponent form multiplied
together we will separate them by multiplication signs ( \(\times\)).
$SPACE
For example if we want
to express \(5\times 5\times  
5\times 5\times 5\times 5\times x\times x\) we write \(5^6\times x^2\), but since you cannot make a \(\times\) symbol
 in the answer box, you should write  (5^6)( x^2) .
$BR
Some questions will have answers that involve exponents and so you
will need to know how to create your answers to these questions. 
Here we will use the ^ symbol (which is over the 6 on most keyboards).
For example, you should write 3^5 when you want to say that the 
answer is \(3^5\).

$PAR
To write \(7\times 7\times 7\times 7 \) as an answer in base-exponent 
form you should write \{ans_rule(6)\}
$PAR
To write \(2\times 2\times 3\times 3\times 3\times 3 \times 3\times y
\times y\) as an answer in base-exponent 
form you should write \{ans_rule(15)\} 
$PAR
Evaluate each of the following expressions.
$PAR
\(2^3\)=\{ans_rule(4)\}
$PAR
\(2^5\)=\{ans_rule(4)\}
$PAR
\(2^6\)=\{ans_rule(4)\}
$PAR
\(2^3\times 2^3\) = \{ans_rule(10)\}
$PAR
Write the following product in base-exponent form.
$PAR
\((2\times 2\times 2)(2\times 2\times 2)\)=  \{ans_rule(10)\}
$PAR
Looking at the evidence above we can see that when we multiply two 
numbers in base-exponent form, if the numbers have the same base our
answer has the same base but we must \{ans_rule(10)\} the exponents.
(answer with the word \(add\) or the word \(multiply\)).
$PAR
Multiply the following pairs of numbers and leave your answer in base-exponent form.
$PAR
\($a^$b\times $a^$c\)=\{ans_rule(2)\}
$PAR
This one includes one more trick
$BR
\((($a^$e)($b^$f))\times(($a^$g)($b^$h))\)= \{ans_rule(10)\} 
$PAR
END_TEXT
Context("Numeric");

ANS(Compute("7^4")->cmp(checker => sub {
my ($correct,$student,$ans) = @_;
return 0 unless $correct == $student;
my $tree = $student->{equation}{tree};
if ($tree->class eq 'BOP' && ($tree->{bop} eq '^' || $tree->{bop} eq '**') &&
$tree->{lop}->class eq 'Number' && $tree->{rop}->class eq 'Number') {
return 1 if $tree->{rop}{value} != 1;
Value->Error("Your exponent should not be 1") unless $ans->{isPreview};
return 0;
}
Value->Error("Your answer should be in the form of a power: a^b")
unless $ans->{isPreview};
return 0;
}));



ANS(str_cmp($ans2));
ANS(Real($ans3)->cmp);
ANS(Real($ans4)->cmp);
ANS(Real($ans5)->cmp);
ANS(Real($ans6)->cmp);
ANS(str_cmp($ans7));
ANS(str_cmp($ans8));
ANS(str_cmp($ans9));
ANS(str_cmp($ans10))     ;

BEGIN_SOLUTION
$PAR Solution $PAR
Just a few sample solutions$PAR
\(2\times 2\times 3\times 3\times 3\times 3 \times 3\times y \times y\).$PAR
Here you have two 2's, five 3's and two y's so \(2^2\times 3^5\times y^2\) is the base exponent form.$PAR
\((($a^$e)($b^$f))\times(($a^$g)($b^$h))\).$PAR
Here you may use the associative law to rewrite this as 
\((($a^$e)($a^$g))\times(($b^$f)($b^$h))\)$BR and so you can deal with the $a factors and the $b factors  separately to get the answer that you would write as  $ans10 .
END_SOLUTION	


;
ENDDOCUMENT();

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Level(2)
## MO(1)
## MLT(Algebra_PropExp_PropExp)



