# WeBWorK problem written by Alex Jordan 
# Portland Community College
#
# Template:
# Solve the following radical equation
#
#       f = 1/(2pi*sqrt(LC))     solve for C
#
#
# Last updated: Spoddeck 12/8/15
# ENDDESCRIPTION
 
## DBsubject('')
## DBchapter('')
## DBsection('')
## KEYWORDS('')
## DBCCSS('')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan', 'Heiko Spoddeck")
## Institution('PCC')


##############################################


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",  
  "bizarroArithmetic.pl",
  "PGML.pl",
  "PGcourse.pl",
);


##############################################


TEXT(beginproblem());

Context("Numeric");
Context()->variables->are(f=>'Real',L=>'Real',C=>'Real');
Context()->operators->set(
   '/' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
   ' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '//' => {class => 'bizarro::BOP::divide', isCommand => 1},
   '*' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   '* ' => {class => 'bizarro::BOP::multiply', isCommand => 1},
   ' *' => {class => 'bizarro::BOP::multiply', isCommand => 1});

$left = Formula("f");
$right = Formula("1/(2 pi sqrt(L C))");

$ans = Formula("1/(4 pi^2 L f^2)");


##############################################


BEGIN_PGML

In an electric circuit, resonance occurs when the frequency [`f`], inductance [`L`], and capacitance C fulfill the following equation:

>>[``[$left] = [$right]``]<<

Solve the equation for the capacitance  [`C`].

The frequency is measured in Hertz, the inductance in Henry, and the capacitance in Farad.

If you need to use a square root, type *sqrt(23)* for [`\sqrt{23}`] (for example). To enter [`\pi`], type *pi*.

    [`C = `] [____________________].

END_PGML


##############################################


ANS($ans -> cmp(
   checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      Context()->flags->set(limits=>[0.95,1.05]);
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1,bizarroMul=>1);
      delete $correct->{test_values}, $student->{test_values};
      Context()->flags->set(limits=>[0.25,0.4]);
      Value->Error("Your answer is correct, but please simplify it further") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0,bizarroMul=>0);
      return 1;
}
));


##############################################


BEGIN_PGML_SOLUTION


Since there is no addition or subtraction in this problem, we can square the equation at any point. We will do that again as our first step. So to get rid of the square root, we square the left side and the right side of the equation, using exponent rules to simplify.

    [``
\begin{aligned}
   [$left] &= [$right] \\
   \left(f\right)^2 &= \left(\frac{1}{2\pi\sqrt{LC}}\right)^2 \\
   f^2 &= \frac{1}{4\pi^2LC} \\
\end{aligned}
    ``]

In order to solve for [`C`], we need to get it out of the denominator by multiplying with [`C`] (which is also [`\frac{C}{1}`]) on both sides:

    [``
\begin{aligned}
   f^2 &= \frac{1}{4\pi^2LC} \\
   f^2\cdot C &= \frac{1}{4\pi^2LC}\cdot \frac{C}{1} \\
   f^2C &= \frac{1}{4\pi^2L} \\
\end{aligned}
    ``]

Since we want to solve for [`C`], we need to divide by [`f^2`]. Since we have a fraction on the right side, we will multiply by the reciprocal of [`f^2`] instead.

    [``
\begin{aligned}
   f^2C &= \frac{1}{4\pi^2L} \\
   \frac{f^2C}{f^2} &= \frac{1}{4\pi^2L}\cdot \frac{1}{f^2} \\
   C &= \frac{1}{4\pi^2Lf^2} \\
\end{aligned}
    ``]

So [`C=[$ans]`] is the solution when we solve [`[$left]=[$right]`] for [`C`].

END_PGML_SOLUTION


##############################################


ENDDOCUMENT();
