##DESCRIPTION
##  Integration of Trig integrals cos(mx)cos(nx)
##  Problems goes step by step through setting up an integral in terms of
##  product to sum formulas before integrating.
##ENDDESCRIPTION

## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Trigonometric integrals)
## Institution(Agnes Scott College)
## Author(Larry Riddle)
## Level(2)
## MO(1)
## TitleText1('APEX Calculus')
## AuthorText1('Hartman')
## EditionText1('4.0')
## Section1('6.3')
## Problem1('16')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",  
  "contextFraction.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);



TEXT(beginproblem());

Context("Fraction");
Context()->noreduce('(-x)-y','(-x)+y');
Context()->noreduce('(-x)-y');

$n = random(1,6);
do {$m = random(1,6);} until ($n != $m);
$f = Formula("cos($m*x)*cos($n*x)")->reduce;

$subt = Formula("1/2*cos(($m-$n)*x)")->reduce;
$add = Formula("1/2*cos(($m+$n)*x)")->reduce;
$ans2 = Formula("1/(2*abs($m-$n))*sin(abs($m-$n)*x)+1/(2*($m+$n))*sin(($m+$n)*x)")->reduce;

$multians = MultiAnswer($subt, $add)->with(
  singleResult => 1,
  separator => " + ",
  tex_separator => "\text{ + }",
  allowBlankAnswers => 1,
  checker => sub {
     my $correct= shift; my $student = shift;
     my ($ca,$cb) = @$correct;
     my ($sa,$sb) = @$student;
      if ( ($ca == $sa && $cb == $sb) ||
           ($ca == $sb && $cb == $sa) ) {
          return [1,1];
      } else {
          if ($ca == $sa || $cb == $sa) {
              Value->Error("The second term is incorrect");
              return [0,0];
          } elsif ($ca == $sb || $cb == $sb) {
              Value->Error("The first term is incorrect");
              return [0,0];
          } else {
              return [0,0];
          }
      }
  }
);

BEGIN_PGML

The goal of this problem is to evaluate [``\int [$f]\;dx``].

First convert this to a new integral with [`[$f]`] rewritten as the sum of two trig functions rather than the product of two trig functions, then integrate this new integral.

[``\int [$f]\;dx = \int \Big(``] [__]{$multians}{10} [`+`] [__]{$multians}{10} [`\Big)\;dx = `] [__]{$ans2}{25} [`\;+\; C`]



END_PGML

BEGIN_PGML_HINT
Use the product to sum trig identity:

[``\cos(m x)\cos(n x) = \frac{1}{2}\Big(\cos((m-n)x)+ \cos((m+n)x)\Big)``]

If necessary, you can also use the trig identities:

[`\sin(-x) = -\sin(x)`]  
[`\cos(-x) = \cos(x)`]

END_PGML_HINT

$a = $n-$m;
$a2 = 2*$a;
$b = $n+$m;
$c = $m-$n;
$c2 = 2*$c;
$a3 = ($a == 1) ? "" : $a;
if ($c == 1) {
  $c3 = "";
} elsif ($c == -1) {
  $c3 = "-";
} else {
  $c3 = $c;
}
if ($m-$n < 0) {
   $s1 = "&= \int \frac{1}{2}\Big(\cos($a3 x) + \cos($b x)\Big)\;dx \\ \\";
   $s2 = "\frac{1}{$a2}\sin($a3 x)";
} else {
   $s1 = "";
   $s2 = "\frac{1}{$c2}\sin($c3 x)";
}

BEGIN_PGML_SOLUTION
[``\begin{aligned}
\int [$f]\;dx &= \int \frac{1}{2}\Big(\cos(([$m]-[$n])x)+ \cos(([$m]+[$n])x)\Big)\;dx \\ \\
&= \int \frac{1}{2}\Big(\cos([$c3]x) + \cos([$m+$n]x)\Big)\;dx \\ \\
[$s1]
&= [$s2] + \frac{1}{[$m*2+$n*2]}\sin([$m+$n]x) + C
\end{aligned}
``]

END_PGML_SOLUTION

COMMENT('Randomization provides 30 different possible versions of this question.');

ENDDOCUMENT();