# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus','derivative','limits','continuity','piecewise functions')
## Tagged by glr 11/19/2008

## DBsubject('Calculus')
## DBchapter('Limits and Derivatives')
## DBsection('The Limit of a Function')
## Date('')
## Author('Gavin LaRose');
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('4')
## AuthorText1('Hughes-Hallett')
## Section1('1.8')
## Problem1('15')

## Textbook tags
## HHChapter('A Library of Functions')
## HHSection('Limits')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
"parserPopUp.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->strings->add( any=>{}, none=>{} );
$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(1,8,1);
$c = random(1,8,1);
$k = Compute("ln($a)");

$others = PopUp( ["?","yes","no"], "yes" );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Find a value of the constant \( k \) such that
the indicated limit exists.
$BR
\[ \lim_{x\to\infty}\,\frac{$a^x - $b}{e^{kx} + $c}. \]
$PAR
\(k = \) \{ ans_rule(25) \}
$BR
${BITALIC}(If any value of k will work, enter ${BBOLD}any$EBOLD;
if none will work, enter ${BBOLD}none$EBOLD.)$EITALIC
$PAR
Are there any other values of \(k\) that will work?
\{ $others->menu() \}

END_TEXT
Context()->normalStrings;

ANS( $k->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return $student >= $correct;
  } ) );
ANS( $others->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Dividing numerator and denominator by \($a^x\) gives
\[ \lim_{x\to\infty}\,\frac{$a^x - $b}{e^{kx} + $c} =
   \lim_{x\to\infty}\,\frac{1 - \frac{$b}{$a^x}}{\frac{e^{kx}}{$a^x} + \frac{$c}{$a^x}}. \]
Then, as \(x\to\infty\), we have \(\frac{$b}{$a^x}\to 0\) and
similarly \(\frac{$c}{$a^x}\to 0\).  Thus for the limit to exist, we must
require that the limit
\[ \lim_{x\to\infty}\,\frac{$a^x}{e^{kx}}  \]
is well-defined.  This will only occur if the two exponentials are actually
the same, that is, if \(e^{kx} = $a^x\), or if \(e^{kx} > $a^x\) for large
\(x\).  Thus, we must have \(k\ge \ln($a)\), and no matter what value of
\(k\) we pick for the first answer there must be others that will also
work.

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
