## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/30/06

## DBsubject(Linear algebra)
## DBchapter(Matrix factorizations)
## DBsection(Diagonalization)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## MLT(diag2x2)
## MLTleader(1)
## Level(4)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('6.4')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');

$r_sq = non_zero_random(1,6,1);
$a = non_zero_random(-9,9,1);
$d = non_zero_random(-9,9,1);
$b = non_zero_random(-.5,.5,.25);
$c = ($r_sq**2-($a - $d )**2)/(4*$b);

while( $a/$b == $c/$d ){
$r_sq += 1;
$c = ($r_sq**2-($a - $d )**2)/(4*$b);
}

## Eigen values are ($a+$d-$r_sq)/2 and ($a+$d+$r_sq)/2
$e_1 = ($a+$d-$r_sq)/2;
$e_2 = ($a+$d+$r_sq)/2;

## Eigen vectors are
$sol11 = 4*$b;
$sol12 = 4*$b;
$sol21 = 2*(-$a+$d-$r_sq);
$sol22 = 2*(-$a+$d+$r_sq);



$A = Matrix([[$a,$b],[$c,$d]]);

$S = Matrix([[ $sol11,$sol12],[$sol21,$sol22]]);
$D = Matrix([[$e_1,0],[0,$e_2]]);

$multians = MultiAnswer($S,$D)->with(
  singleResult => 0,
  allowBlankAnswers => 0,

  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my $S = Matrix($s[0]);
      my $D = Matrix($s[1]);

      return 0 if $S->det == Real(0);
      return 0 if ($D->element(1,2) != Real(0)) or ($D->element(2,1) != 0); 
      return $A * $S == $S * $D;
  }
);


Context()->texStrings;
BEGIN_TEXT
Let \[ A = $A. \]
Find an invertible matrix \(S\) and a diagonal matrix \(D\) such that \(S^{-1}AS=D\). 
$BR
$BR
$BCENTER
\(S=\) \{ $multians->ans_array(2,2,5) \}
$BR
$BR
\(D=\) \{ $multians->ans_array(2,2,5) \}
$ECENTER
END_TEXT
Context()->normalStrings;

ANS($multians->cmp);

;
ENDDOCUMENT();
