# DESCRIPTION
# Gives an expression and asks to translate from english into math
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Algebraic expressions)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(ConvertFromWords_ThreeBlanks)
## Level(1)
## KEYWORDS('real numbers')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserAssignment.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real",y=>"Real",c=>"Real",d=>"Real");
Context()->strings->add("="=>{}, "<="=>{}, "<"=>{}, ">"=>{}, ">="=>{});
parser::Assignment->Allow;

Context()->texStrings;
$casea = random(0,1,1);
$caseb = random(0,1,1);

$m1 = random(5,15);
$m2 = random(-15,-5);
$n1 = random(5,15);
do {$n2 = random(5,15)} until ($n2 != -$n1);
if ($casea == 0) {
  $ans11 = Formula("$m1 + x");
  $ans12 = "=";
  $ans13 = Formula("$m2");
  $question1 = "The sum of \($m1\) and \(x\) is \($m2\)";
  $dispans1 = "$m1 + x = $m2";
} else {
  $ans11 = Formula("y - $n1");
  $ans12 = "=";
  $ans13 = Formula("$n2");
  $question1 = "The difference of \(y\) and \($n1\) amounts to \($n2\)";
  $dispans1 = "y - $n1 = $n2";
}

if ($caseb == 0) {
  $p1 = random(2,5);
  $case1 = random(0,1,1);
  $word1 = ($case1 == 0) ? "plus" : "minus";
  if ($p1 == 2) {
    $disp1 = "Twice";
  } elsif ($p1 == 3) {
    $disp1 = "Three times";
  } elsif ($p1 == 4) {
    $disp1 = "Four times";
  } elsif ($p1 == 5) {
    $disp1 = "Five times";
  }
  do {$p2 = random(2,7)} until ($p2 != $p1);
  $p3 = random(2,19)*random(-1,1,2);
  $ans21 = ($case1 == 0) ? Formula("$p1*c + $p2") : Formula ("$p1 c - $p2");
  $ans22 = "=";
  $ans23 = Formula("$p3");
  $question2 = "$disp1 \(c\), $word1 \($p2\), is the same as \($p3\)";
  $dispans2 = ($case1 == 0) ? "$p1 c + $p2 = $p3" : "$p1 c - $p2 = $p3";
} else {
  Context()->texStrings;
  $q1 = random(2,5);
  $case2 = random(0,1,1);
  $word2 = ($case2 == 0) ? "more" : "less";
  if ($q1 == 2) {
    $disp2 = "Two";
  } elsif ($q1 == 3) {
    $disp2 = "Three";
  } elsif ($q1 == 4) {
    $disp2 = "Four";
  } elsif ($q1 == 5) {
    $disp2 = "Five";
  }
  do {$q2 = random(2,7)} until ($q2 != $q1);
  $q3 = random(2,19)*random(-1,1,2);
  $ans21 = ($case2 == 0) ? Formula("$q1+$q2*d") : Formula("$q2*d - $q1");
  $ans22 = "=";
  $ans23 = Formula("$q3");
  $question2 = "$disp2 $word2 than the product of \($q2\) and \(d\) is \($q3\)";
  $dispans2 = ($case2 == 0) ? "$q1 + $q2 d = $q3" : "$q2 d - $q1 = $q3";
}

$ma1 = MultiAnswer($ans11,$ans12,$ans13)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);

$ma2 = MultiAnswer($ans21,$ans22,$ans23)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write each sentence using mathematical symbols:
$PAR
(a) $question1: \{$ma1->ans_rule(10)\} \{$ma1->ans_rule(5)\} \{$ma1->ans_rule(10)\}
$PAR
(b) $question2: \{$ma2->ans_rule(10)\} \{$ma2->ans_rule(5)\} \{$ma2->ans_rule(10)\}
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($ma1->cmp());
ANS($ma2->cmp());

  BEGIN_TEXT
  $PAR
  ${BBOLD}Note:${EBOLD} The middle answer blank in both parts is for you to enter "\(<\)=", "\(<\)", "=", "\(>\)", or "\(>\)=" for "less than or equal to", "less than", "equals", "greater than", or "greater than or equal to", respectively.
  END_TEXT
Context()->normalStrings;

ENDDOCUMENT();
