# DESCRIPTION
# Gives an expression and asks to translate from english into math
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Algebraic expressions)
## MLT(ConvertFromWords_ThreeBlanks)
## MLTleader(1)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## Level(2)
## KEYWORDS('real numbers')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserAssignment.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","z","p","q","m","n","k","u","v", "a", "b");
for ($i = 1; $i < scalar(@vars); $i++) {
  Context()->variables->add($vars[$i]=>"Real");
}
Context()->strings->add("="=>{}, "<="=>{}, "<"=>{}, ">"=>{}, ">="=>{});

for ($i = 0; $i < random(10,30); $i++) { #makes more random
  $case1 = random(0,1,1);
  $case2 = random(0,1,1);
  $case21 = random(0,3,1); # determine inequality to use
  $var1 = $vars[random(0, scalar(@vars)-1, 1)];
  $var2 = $vars[random(0, scalar(@vars)-1, 1)];
}

@w = ("Twice", "Three times", "Four times", "Five times");
@s1 = ("less than or equal to", "greater than or equal to", "less than", "greater than");
@s2 = ("<=", ">=", "<", ">");
if ($case1 == 0) {
  $m1 = random(2,5);
  $m2 = random(6,10);
  $m3 = non_zero_random(-40,40);
  $question1 = "$w[$m1-2] the difference of \($var1\) and \($m2\) is \($m3\)";
  $ans11 = Formula("$m1($var1-$m2)");
  $ans12 = "=";
  $ans13 = Formula("$m3");
  $dispans1 = "\($m1($var1- $m2) = $m3\)";
} else {
  $m1 = random(2,15);
  $m2 = random(6,10);
  $m3 = non_zero_random(-40,40);
  $question1 = "$m1 times the sum of \($m2\) and \($var1\) is \($m3\)";
  $ans11 = Formula("$m1($m2+$var1)");
  $ans12 = "=";
  $ans13 = Formula("$m3");
  $dispans1 = "\($m1($m2+$var1) = $m3\)";
}

if ($case2 == 0) {
  $n1 = random(2,5,1);
  $n2 = random(6,31,1);
  $n3 = random(-1,1,2)*random(2,13,1);
  $question2 = "$w[$n1-2] the difference of \($var2\) and \($n2\) is $s1[$case21] the reciprocal of \($n3\)";
  $ans21 = Formula("$n1 ($var2 - $n2)")->reduce;
  $ans23 = Formula("1/$n3")->reduce;
  $disp23 = ($n3 < 0) ? "-\frac{1}{".abs($n3)."}" : "\frac{1}{$n3}";
} else {
  $n1 = random(2,5,1);
  $n2 = random(6,31,1);
  $n3 = random(-1,1,2)*random(2,13,1);
  $question2 = "$w[$n1-2] the sum of \($n2\) and \($var2\) is $s1[$case21] the opposite of \($n3\)";
  $ans21 = Formula("$n1 ($n2 + $var2)")->reduce;
  $ans23 = Formula("-$n3")->reduce;
  $disp23 = -$n3;
}
$ans22 = $s2[$case21];
if ($case21 == 0) {
  $dispineq = "\leq";
} elsif ($case21 == 1) {
  $dispineq = "\geq";
} elsif ($case21 == 2) {
  $dispineq = "<";
} else {
  $dispineq = ">";
}

$ma1 = MultiAnswer($ans11,$ans12,$ans13)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write each sentence using mathematical symbols:
$PAR
(a) $question1: \{$ma1->ans_rule(10)\} \{$ma1->ans_rule(5)\} \{$ma1->ans_rule(10)\}
$PAR
(b) $question2: \{ans_rule(10)\} \{ans_rule(5)\} \{ans_rule(10)\}
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($ma1->cmp());
ANS($ans21->cmp());
ANS(String($ans22)->cmp());
ANS($ans23->cmp());

  BEGIN_TEXT
  $PAR
  ${BBOLD}Note:${EBOLD} The middle answer blank in both parts is for you to enter "\(<\)=", "\(<\)", "=", "\(>\)", or "\(>\)=" for "less than or equal to", "less than", "equals", "greater than", or "greater than or equal to", respectively.
  $BR
  END_TEXT
Context()->normalStrings;

ENDDOCUMENT();
