# DESCRIPTION
# Gives an expression and asks to translate from english into math
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Algebra of real numbers and simplifying expressions)
## DBsection(Algebraic expressions)
## Institution(University of Missouri-Columbia)
## Author(Rick Lynch)
## MLT(ConvertFromWords_ThreeBlanks)
## Level(2)
## KEYWORDS('real numbers')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserAssignment.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(c=>"Real",d=>"Real");
Context()->strings->add("="=>{}, "<="=>{}, "<"=>{}, ">"=>{}, ">="=>{});
parser::Assignment->Allow;

Context()->texStrings;
$m1 = random(2,5);
$case1 = random(0,1,1);
$word1 = ($case1 == 0) ? "plus" : "minus";
if ($m1 == 2) {
  $disp1 = "Twice";
} elsif ($m1 == 3) {
  $disp1 = "Three times";
} elsif ($m1 == 4) {
  $disp1 = "Four times";
} elsif ($m1 == 5) {
  $disp1 = "Five times";
}
do {$m2 = random(2,7)} until ($m2 != $m1);
$m3 = random(2,19)*random(-1,1,2);
$ans11 = ($case1 == 0) ? Formula("$m1*c + $m2") : Formula ("$m1 c - $m2");
$ans12 = "=";
$ans13 = Formula("$m3");

Context()->texStrings;
$n1 = random(2,5);
$case2 = random(0,1,1);
$word2 = ($case2 == 0) ? "more" : "less";
if ($n1 == 2) {
  $disp2 = "Two";
} elsif ($n1 == 3) {
  $disp2 = "Three";
} elsif ($n1 == 4) {
  $disp2 = "Four";
} elsif ($n1 == 5) {
  $disp2 = "Five";
}
do {$n2 = random(2,7)} until ($n2 != $n1);
$n3 = random(2,19)*random(-1,1,2);
$ans21 = ($case2 == 0) ? Formula("$n1+$n2*d") : Formula("$n2*d - $n1");
$ans22 = "=";
$ans23 = Formula("$n3");

$dispans1 = ($case1 == 0) ? "$m1 c + $m2 = $m3" : "$m1 c - $m2 = $m3";
$dispans2 = ($case2 == 0) ? "$n1 + $n2 d = $n3" : "$n2 d - $n1 = $n3";

$ma1 = MultiAnswer($ans11,$ans12,$ans13)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);

$ma2 = MultiAnswer($ans21,$ans22,$ans23)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write each sentence using mathematical symbols:
$PAR
(a) $disp1 \(c\), $word1 \($m2\), is the same as \($m3\): \{$ma1->ans_rule(10)\} \{$ma1->ans_rule(5)\} \{$ma1->ans_rule(10)\}
$PAR
(b) $disp2 $word2 than the product of \($n2\) and \(d\) is \($n3\): \{$ma2->ans_rule(10)\} \{$ma2->ans_rule(5)\} \{$ma2->ans_rule(10)\}
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($ma1->cmp());
ANS($ma2->cmp());

  BEGIN_TEXT
  $PAR
  ${BBOLD}Note:${EBOLD} The middle answer blank in both parts is for you to enter "<=", "<", "=", ">", or ">=" for "less than or equal to", "less than", "equals", "greater than", or "greater than or equal to", respectively.
  $BR
  END_TEXT
Context()->normalStrings;

ENDDOCUMENT();
