##DESCRIPTION
##  Algebra: exponents and radicals 
##ENDDESCRIPTION

##KEYWORDS('algebra', 'radicals')

##DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## Date('10/27/2009')
## Author('Justin Sukiennik')
## Institution('University of Minnesota')
## TitleText1('Algebra for College Students')
## EditionText1('8')
## AuthorText1('Kaufmann, Schwitters')
## Section1('5.3')
## Problem1('12')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "parserMultiAnswer.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);


########################################################################

install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 0;

TEXT(beginproblem());

########################################################################

Context("Numeric");

$a = random(5,7,2);
$b = list_random(2,4);

$c = Compute("$a-2")->reduce;
$d = Compute("5*$b^2")->reduce;

$numerator = Real("2*$a-2");
$denominator = Real("$a");
$radical = Real("5");

########################################################################

$mp = MultiAnswer($numerator,$denominator)->with(
         singleResult => 0,
         separator => "/",
         tex_separator => "/",
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Your numerator (and maybe more) is incorrect") 
                unless $numeratorValid;
              Value::Error("Your denominator is incorrect") 
                unless $denominatorValid;
              Value::Error("Your fraction has the wrong sign") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );


########################################################################

Context()->texStrings;
BEGIN_TEXT
Use the distributive property to help simplify the numerical expression
\[
\frac{$c}{$a}\sqrt{5}+\frac{1}{$b}\sqrt{$d}
\]
into the form \(\frac{A}{B} \sqrt{C}\), where \(A\), \(B\), and \(C\) are all integers.
$PAR
Answer: \(A =\) \{$mp->ans_rule(15)\}, \(B =\) \{$mp->ans_rule(15)\}, and \(C =\) \{ ans_rule(15) \}
$HR
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );
ANS( $radical->cmp );

ENDDOCUMENT();       