## DESCRIPTION
## Basic Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/4/06

## DBsubject(Algebra)
## DBchapter(Absolute value expressions and functions)
## DBsection(Absolute value inequalities)
## Date(6/3/2002)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.2')
## Problem1('21 22')
## KEYWORDS('algebra, absolute value inequalities')

DOCUMENT();     
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl"
);

($a,$b,$c,$d) = random_subset(4,(2..30));

Context("ArbitraryString");

$ans1 = Compute("|x-$a| <= $b")->cmp(checker => sub {
    my ($correct,$student,$ans) = @_; 
    $correct = $correct->value; # get perl string from String object 
    $student = $student->value; # ditto 
    #remove white space 
    $correct =~ s/ //g ; 
    $student =~ s/ //g ;
    if ( "|x-$a|<=$b" eq $student || "|$a-x|<=$b" eq $student ||
         "$b>=|x-$a|" eq $student || "$b>=|$a-x|" eq $student ) {
       return 1;
    }
    else {return 0;}
});

$ans2 = Compute("|x+$c| >= $d")->cmp(checker => sub {
    my ($correct,$student,$ans) = @_; 
    $correct = $correct->value; # get perl string from String object 
    $student = $student->value; # ditto 
    #remove white space 
    $correct =~ s/ //g ; 
    $student =~ s/ //g ;
    if ( "|x+$c|>=$d" eq $student || "|$c+x|>=$d" eq $student ||
         "$d<=|x+$c|" eq $student || "$d<=|$c+x|" eq $student ) {
       return 1;
    }
    else {return 0;}
});

$ans3 = Compute("|x| >= $c")->cmp(checker => sub {
    my ($correct,$student,$ans) = @_; 
    $correct = $correct->value; # get perl string from String object 
    $student = $student->value; # ditto 
    #remove white space 
    $correct =~ s/ //g ; 
    $student =~ s/ //g ;
    if ( "|x|>=$c" eq $student || "$c<=|x|" eq $student ) {
       return 1;
    }
    else {return 0;}
});

$ans4 = Compute("|x| <= $a")->cmp(checker => sub {
    my ($correct,$student,$ans) = @_; 
    $correct = $correct->value; # get perl string from String object 
    $student = $student->value; # ditto 
    #remove white space 
    $correct =~ s/ //g ; 
    $student =~ s/ //g ;
    if ( "|x|<=$a" eq $student || "$a>=|x|" eq $student ) {
       return 1;
    }
    else {return 0;}
});

BEGIN_PGML
Translate the following sentences into equivalent mathematical expressions that use absolute values. For example "Absolute value of [`x`] is less than or equal to 100" would become *|x| <= 100*.  Your math expression should be as simple and concise as possible, for example, enter *100* instead of *99+1* or *-(-100)*. For absolute value use vertical bars, e.g. *|x|*.  For [`\geq`] enter *>=*, for [`\leq`] enter *<=*.  

a. [`\ `] The distance between [`x`] and [`[$a]`] is no more than [`[$b]`].  [_]{$ans1}{20}

b. [`\ `] The distance between [`x`] and [`- [$c]`] is at least [`[$d]`].  [_]{$ans2}{20}

c. [`\ `] [`x`] is at least [`[$c]`] units from [`0`].  [_]{$ans3}{20}

d.  [`\ `] [`x`] is at most [`[$a]`] units from [`0`].  [_]{$ans4}{20}

END_PGML

BEGIN_PGML_SOLUTION
a. "[`|x-[$a]|`]" means "the distance from [`x`] to [`[$a]`]" and "is no more than [$b]" means "is less than or equal to [$b]" so sentence a. becomes [`|x-[$a]| <= [$b]`]. 

b. "The distance between [`x`] and -[$c]" becomes "[`|x-(-[$c])|`] which simplifies to "[`|x+[$c]|`], and "is at least [$d]" means "is greater than or equal to [$d]" so sentence b. becomes [`|x+[$c]| >= [$d]`] . 

c.  The distance from [`x`] to 0 is [`|x-0| = |x|`] so sentence c. becomes [`|x| >= [$c]`] .

d.  As in part c., [`|x|`] is the distance from [`x`] to 0, and "is at most" means "is less than or equal to" so sentence d. becomes [`|x| <= [$a]`] .
END_PGML_SOLUTION

ENDDOCUMENT();
