# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION






##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPolynomial.pl",
  
);

TEXT(beginproblem());
##############################################

Context("Numeric");
$var = 'y';
Context()->variables->are($var=>'Real');
Context()->strings->add("cannot be factored"=>{},
    "cannot be factored."=>{alias=>"cannot be factored"});
Context()->flags->set(partialCredit=>0);

$a = random(1,20,1);
$b = random(1,20,1);
while (gcd($a,$b) != 1)
{$b = random(1,20,1);}


$exp = Compute("$a y^2-$b y")->reduce;

$fac1 = Compute("y");
$fac2 = Compute("($a y-$b)")->reduce;

$fctd = Compute("$fac1  $fac2");
$fctd2 = Compute("$fac2  $fac1");

##############################################

Context()->texStrings;

BEGIN_TEXT


Factor the polynomial.
$PAR


$SPACE $SPACE $SPACE \( $exp \)

$PAR


\{ans_rule(20)\}



 




END_TEXT

Context()->normalStrings;

##############################################
$showPartialCorrectAnswers = 1;


ANS($fctd->cmp(typeMatch=>Formula($var), checker => sub {
      my ($correct,$student,$ans) = @_;
      if (($student eq $correct) || ($student eq $fctd2))
      {return 1};
    }));


ENDDOCUMENT();
