## DESCRIPTION
## Linear transformations
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Properties)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('Multivariable Mathematics')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')
## KEYWORDS('linear transformations')

DOCUMENT(); 
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "parserMultiAnswer.pl",
  "contextLimitedVector.pl",
  "parserPopUp.pl",
  "unionLists.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;

Context("Numeric")->variables->are('c'=>'Real','x'=>'Real','y'=>'Real');
Context()->noreduce('(-x)-y');

$a = random(-1,1,2) * random(2,9,1);
$b = non_zero_random(-9,9,1);

$f = Formula("$a x + $b")->reduce;

$popup = PopUp(['choose','f is a linear transformation','f is not a linear transformation'],'f is not a linear transformation');
$popup_add = PopUp(['choose','Yes, they are equal','No, they are not equal'],'No, they are not equal');
$popup_mult = PopUp(['choose','Yes, they are equal','No, they are not equal'],'No, they are not equal');


$ans[1] = Formula("$a x + $a y + $b")->reduce;
$ans[2] = Formula("$a x + $b")->reduce;
$ans[3] = Formula("$a y + $b")->reduce;

$ans[4] = Formula("$a c x + $b")->reduce;
$ans[5] = Formula("c");
$ans[6] = Formula("$a x + $b")->reduce;

$multians[1] = MultiAnswer($ans[2], $ans[3])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s + %s",
  tex_format => "%s + %s",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);

$multians[2] = MultiAnswer($ans[5], $ans[6])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s ( %s )",
  tex_format => "%s ( %s )",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);



Context()->texStrings;
BEGIN_TEXT
Let \( f : \mathbb{R} \to \mathbb{R} \) be defined by \( f(x) = $f \).  Is \( f \) a linear transformation?
$BR
$BR
\{ BeginList('OL',type=>'a') \}
$ITEM
\( f(x+y) = \) \{ $ans[1]->ans_rule(20) \}.
$BR
\( f(x) + f(y) = \) \{ $multians[1]->ans_rule(10) \} \( + \) \{ $multians[1]->ans_rule(10) \}.
$BR
Does \( f(x+y) = f(x) + f(y) \) for all \( x,y \in \mathbb{R} \)? \{ $popup_add->menu \}
$ITEMSEP
$ITEM
\( f(c x) = \) \{ $ans[4]->ans_rule(20) \}.
$BR
\( c(f(x)) = \) \{ $multians[2]->ans_rule(5) \} \( \Big( \) \{ $multians[2]->ans_rule(10) \} \( \Big) \).
$BR
Does \( f(cx) = c(f(x))\) for all \( c, x \in \mathbb{R} \)? \{ $popup_mult->menu \}
$ITEMSEP
$ITEM Is \( f \) a linear transformation? \{ $popup->menu \}
\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;

ANS($ans[1]->cmp);
ANS($multians[1]->cmp);
ANS($popup_add->cmp);

ANS($ans[4]->cmp);
ANS($multians[2]->cmp);
ANS($popup_mult->cmp);

ANS($popup->cmp);

COMMENT('MathObject version. Never a linear transformation.');
ENDDOCUMENT();
