## DESCRIPTION
## Linear transformations
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Linear transformations)
## DBsection(Properties)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('Multivariable Mathematics')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')
## KEYWORDS('linear transformations')

DOCUMENT(); 
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "unionLists.pl",
  "parserFunction.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;

Context('Matrix')->variables->are('c'=>'Real',
'a11'=>'Real','a12'=>'Real','a21'=>'Real','a22'=>'Real',
'b11'=>'Real','b12'=>'Real','b21'=>'Real','b22'=>'Real',
);
Context()->noreduce('(-x)-y');
Context()->variables->set(
   'a11' => {TeX => "a_{11}"},'a12' => {TeX => "a_{12}"},'a21' => {TeX => "a_{21}"},'a22' => {TeX => "a_{22}"},
   'b11' => {TeX => "b_{11}"},'b12' => {TeX => "b_{12}"},'b21' => {TeX => "b_{21}"},'b22' => {TeX => "b_{22}"},
);

$A = Matrix([['a11','a12'],['a21','a22']]);
$B = Matrix([['b11','b12'],['b21','b22']]);

$popup = PopUp(['choose','f is a linear transformation','f is not a linear transformation'],'f is not a linear transformation');
$popup_add = PopUp(['choose','Yes, they are equal','No, they are not equal'],'No, they are not equal');
$popup_mult = PopUp(['choose','Yes, they are equal','No, they are not equal'],'No, they are not equal');

$ans[1] = Compute("(a11+b11) (a22+b22) - (a12+b12) (a21+b21)");
$ans[2] = Compute("a11 a22 - a12 a21");
$ans[3] = Compute("b11 b22 - b12 b21");

$ans[4] = Compute("c^2(a11 a22 - a12 a21)");
$ans[5] = Formula("c")->reduce;
$ans[6] = Compute("a11 a22 - a12 a21");

$multians[1] = MultiAnswer($ans[2], $ans[3])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s + %s",
  tex_format => "%s + %s",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);

$multians[2] = MultiAnswer($ans[5], $ans[6])->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format => "%s ( %s )",
  tex_format => "%s \left( %s \right)",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      my @score = ();
      foreach my $i (0..$#c) {
        if ($c[$i]->typeMatch($s[$i]) and $c[$i] == $s[$i]) { $score[$i]=1; } else { $score[$i]=0; } 
      }
      return [ @score ];
  }
);



Context()->texStrings;
BEGIN_TEXT
Let \( M_{n,n}(\mathbb{R}) \) denote the vector space of \( n \times n \) matrices with real entries.
Let \( f : M_{2,2}(\mathbb{R}) \to \mathbb{R} \) be the function defined by \( f(A) = \mathrm{det}(A) \) for any \( A \in M_{2,2}(\mathbb{R}) \).
Is \( f \) a linear transformation?
$BR
$BR
Let \( A = $A \) and \( B = $B \) be any two matrices in \( M_{2,2}(\mathbb{R}) \) and let \( c \in \mathbb{R} \).
$BR
$BR
\{ BeginList('OL',type=>'a') \}
$ITEM
\( f(A+B) = \) \{ $ans[1]->ans_rule(50) \}.  (Enter \( a_{11} \) as \(\verb+a11+\), etc.)
$BR
\( f(A) + f(B) = \) \{ $multians[1]->ans_rule(25) \} \( + \) \{ $multians[1]->ans_rule(25) \}.
$BR
Does \( f(A+B) = f(A) + f(B) \) for all \( A, B \in M_{2,2}(\mathbb{R}) \)? \{ $popup_add->menu \}
$ITEMSEP
$ITEM
\( f(c A) = \) \{ $ans[4]->ans_rule(50) \}.
$BR
\( c(f(A)) = \) \{ $multians[2]->ans_rule(5) \} \( \Big( \) \{ $multians[2]->ans_rule(25) \} \( \Big) \).
$BR
Does \( f(c A) = c(f(A))\) for all \( c \in \mathbb{R} \) and all \( A \in M_{2,2}(\mathbb{R}) \)? \{ $popup_mult->menu \}
$ITEMSEP
$ITEM Is \( f \) a linear transformation? \{ $popup->menu \}
\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;

ANS($ans[1]->cmp);
ANS($multians[1]->cmp);
ANS($popup_add->cmp);

ANS($ans[4]->cmp);
ANS($multians[2]->cmp);
ANS($popup_mult->cmp);

ANS($popup->cmp);

COMMENT('MathObject version.  Never a linear transformation.');
ENDDOCUMENT();
