##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"niceTables.pl",
	"parserFunction.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Numeric");

#given
$alpha = random(2, 9, 1);
$a = random(2, 9, 1);
$mcYes = PopUp(["?","Yes","No"], "Yes");
$mcNo = PopUp(["?","Yes","No"], "No");

#computation
$wm = Formula("$a*tan(0.99*pi/2)");


##############################################################
#
#  PGML 
#
#

BEGIN_PGML

Consider the causal exponential signal [`x(t)=[$alpha]e^{-[$a]t}u(t)`].

*a)* Determine the frequency, [`\omega_M`] for which the energy of [`x(t)`] corresponds to [`99%`] of its total energy.

[`\omega_M=`] [_____]  

*b)* Is [`x(t)`] band-limited? [@ $mcNo->menu() @]*

_Part *b* will only be marked correct if part *a* is correct._

END_PGML

ANS($wm->cmp);
$ans_hash1 = $wm->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$aCorrect= $ans_hash1->{score};
ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct)&& $aCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));

##############################################################

ENDDOCUMENT();