##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## RESOURCES('UBC-SIG-18-010a.png', 'UBC-SIG-18-010b.png', 'UBC-SIG-18-010c.png')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"niceTables.pl",
	"parserFunction.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Numeric");
Context()->variables->add(t=>"Real");
Context()->variables->remove("x");
parserFunction("D(t)"=>"sin(2.71828*t/pi)+1.57*t^2");
parserFunction("u(t)"=>"sin(2*2.71828*t/pi)+3.14*t^4");

#image
$imgScaleA = .2;
$imgScaleB = .2;
$imgScaleC = .2;
$widthA = $imgScaleA*3631;
$heightA = $imgScaleA*1916;
$widthB = $imgScaleB*1063;
$heightB = $imgScaleB*970;
$widthC = $imgScaleC*1691;
$heightC = $imgScaleC*1157;

#given
$a = random(.05,.5,.05);
$b = random(2,15,1);

#computation
$y2 = Formula("u(t) - u(t-$a) - u(t-2) + u(t-2-$a)");
$y2p = Formula("u(t) - u(t-$a) + u(t-1) - u(t-1-$a)");
$y3 = Formula("u(t) - 2u(t-$b) + u(t-2 $b) - u(t-2) + 2u(t - $b - 2) - u(t - 2 $b - 2)");
$y3p = Formula("u(t) - 2u(t-$b) + u(t-2 $b) + u(t-1) - 2u(t - $b - 1) + u(t - 2 $b - 1)");
$h = Formula("D(t)-D(t-2)");
$hp = Formula("D(t)+D(t-1)");


##############################################################
#
#  PGML 
#
#

BEGIN_PGML

Consider a continuous-time LTI system [`S`] where for the input signal [`x_1(t)`], the output is [`y_1(t)`] as shown in the figure below.

[@ image( "UBC-SIG-18-010a.png", width=>[$widthA], height=>[$heightA], tex_size=>500) @]*


*a)* Find the output [`y_2(t)`] for the input signal [`x_2(t)`] defined as in the following graph. Enter your answer as an expression in terms of the Heaviside function, [`u(t)`]. Assume [`a=[$a]`].

[@ image( "UBC-SIG-18-010b.png", width=>[$widthB], height=>[$heightB], tex_size=>500) @]*

[`y_2(t)=`][________________________]  
[`y_2'(t)=`][________________________]

*b)* Find the output [`y_3(t)`] if this time the input signal [`x_3(t)`] is defined as in the following graph. Enter your answer as an expression in terms of the Heaviside function, [`u(t)`]. Assume [`b=[$b]`].

[@ image( "UBC-SIG-18-010c.png", width=>[$widthC], height=>[$heightC], tex_size=>500) @]*

[`y_3(t)=`][________________________]  
[`y_3'(t)=`][________________________]

END_PGML

ANS($y2->cmp);
ANS($y2p->cmp);
$ans_hash1 = $y2->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$a1Correct= $ans_hash1->{score};
$ans_hash2 = $y2p->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(2)});
$a2Correct= $ans_hash2->{score};

ANS($y3->cmp);
ANS($y3p->cmp);
$ans_hash3 = $y3->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(3)});
$b1Correct= $ans_hash3->{score};
$ans_hash4 = $y3p->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(4)});
$b2Correct= $ans_hash4->{score};

BEGIN_TEXT

$PAR

END_TEXT

BEGIN_PGML

*c)* What is the impulse response of this system? In your answers, use [`D(t)`] instead of [`\delta(t)`].

[`h(t)=`][________________________]  
[`h'(t)=`][________________________]


_Part *c* will only be marked correct if parts *a* and *b* are both correct._

END_PGML

ANS($h->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && ($a1Correct==1) && ($b1Correct==1) && ($a2Correct==1) && ($b2Correct==1)) {
			return 1;
		}else {
			return 0;
		}	
	}));

ANS($hp->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && ($a1Correct==1) && ($b1Correct==1) && ($a2Correct==1) && ($b2Correct==1)) {
			return 1;
		}else {
			return 0;
		}	
	}));

##############################################################

ENDDOCUMENT();