##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## RESOURCES('UBC-SIG-18-040.png')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

#image scale
$imgScale = .25;

#image aspect ratio
$width = $imgScale*1923;
$height = $imgScale*375;

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"parserFunction.pl"
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
Context()->variables->add(t=>"Real");
Context()->constants->add(j => i);
parserFunction("u(t)"=>"sin(2.71828*t/pi)+1.57*t^2");

#given
$a = random(2, 6, 1);
$b = random(2, 9, 1);
$c = random(2, 6, 1);
$mcNo = PopUp(["?","Yes","No"], "No");

#computation 
$ans1 = Formula("-pi/$b sin(pi/$b (t - $a))(u(t+$c-$a)-u(t-$c-$a))");
$ans2 = Formula("-pi/$b sin(pi/$b (t - $a))(u(t+$c-$a)-u(t-$c-$a))");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

Two filters with frequency responses [`H_1(\omega) = jw`] and [`H_2(\omega) = e^{-[$a]jw}`] for [`-\infty < \omega < \infty`], are cascaded together so that the output of the first filter is fed as the input to the second, as shown in the figure below. 

[@ image("UBC-SIG-18-040.png", width=>[$width], height=>[$height], tex_size=>500) @]*

Suppose that the input to this cascaded system is the signal [`x(t)=cos\big(\frac{\pi t}{[$b]}\big)[u(t+[$c])-u(t-[$c])]`]. 

*a)* Find the output, [`y(t)`], of this cascaded system.  
[`y(t)=`] [__________________]

_Hint: think about what each filter does and do your calculations in the time domain._

*b)* We now reverse the cascading order of the two filters. Find the output, [`y'(t)`], of this new cascaded system.  
[`y'(t)=`] [__________________]

*c)* Does the output depend on the order of cascading? [@ $mcNo->menu() @]*

_Part *c* will only be marked correct if parts *a* and *b* are correct._

END_PGML

ANS($ans1->cmp);
ANS($ans2->cmp);
$ans_hash1 = $ans1->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$aCorrect= $ans_hash1->{score};
$ans_hash2 = $ans2->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(2)});
$bCorrect= $ans_hash2->{score};

ANS($mcNo->cmp(checker=>sub{
		my ($correct,$student) = @_; 
		if (($student == $correct) && $aCorrect==1 && $bCorrect==1) {
			return 1;
		}else {
			return 0;
		}	
	}));

##############################################################

ENDDOCUMENT();