##DESCRIPTION
## 
##ENDDESCRIPTION

## DBsubject()
## DBchapter()
## DBsection()
## Date()
## Author()
## Level()
## KEYWORDS('')
## Institution('University of British Columbia')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
##############################################################

DOCUMENT();

loadMacros(   
	"PGstandard.pl",     # Standard macros for PG language
	"MathObjects.pl",
	"PGML.pl",
	"parserPopUp.pl",
	"parserMultiAnswer.pl",
	"parserRadioButtons.pl",
	"parserFunction.pl",
	"niceTables.pl",
	#"source.pl",        # allows code to be displayed on certain sites.
	#"PGcourse.pl", 	 # Customization file for the course
   	);

TEXT(beginproblem());

##############################################################
#
#  Setup
#
#

Context("Complex");
Context()->constants->add(j => i);
Context()->variables->add(O=>"Real");
Context()->variables->add(n=>"Real");
parserFunction("u(n)"=>"sin(2.71828*n/pi)+1.57*n^2");
parserFunction("D(n)"=>"sin(4*2.71828*n/pi)+3.14*n^4");


#given
$a = random(2,8,1);
$alpha = random(2,9,0.5);
$beta = random(1.1,3.1,0.1);
$mcYes = PopUp(["?","Yes","No"], "Yes");
$e=2.7182818284590452353602874713527;

#computation 
$ho = Formula("$e^(-$a i O)($beta + 2*$alpha cos(O))");
$hn = Formula("$alpha D(n - $a + 1) + $beta D(n - $a) + $alpha D(n - $a -1)");

##############################################################
#
#  Text 
#
#

BEGIN_PGML

The transfer function of a FIR filter is [`H(z)=z^{-[$a]}([$alpha]z +[$beta]+[$alpha]z^{-1})`]. 

*a)* Find the frequency response, [`H(\Omega)`] of the system. _Enter [`\Omega`] as "O"._

[`H(\Omega)=`] [_____]  

*b)* Is the phase response of the system linear? [@ $mcYes->menu() @]* 

*c)* Find the impulse response, [`h[n]`], of the system.

[`h[n]=`] [_____]  

_In your answer, enter [`z(n)`] for a discrete-time function [`z[n]`] and enter [`D(n)`] instead of [`\delta[n]`]. WebWork is unable to parse a function that uses square brackets._

*d)* Is the impulse response symmetric in the y-axis? [@ $mcYes->menu() @]*

_Part *b* will only be marked correct if part *a* is correct. Part *d* will only be marked correct if part *c* is correct._
 
END_PGML

ANS($ho->cmp);

$ans_hash1 = $ho->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(1)});
$aCorrect= $ans_hash1->{score};

ANS($mcYes->cmp(checker=>sub{
	my ($correct,$student) = @_; 
	if (($student == $correct)&& $aCorrect==1) {
		return 1;
	}else {
		return 0;
	}	
}));

ANS($hn->cmp);
$ans_hash3 = $hn->cmp->evaluate($inputs_ref->{ANS_NUM_TO_NAME(3)});
$cCorrect= $ans_hash3->{score};

ANS($mcYes->cmp(checker=>sub{
	my ($correct,$student) = @_; 
	if (($student == $correct)&& $cCorrect==1) {
		return 1;
	}else {
		return 0;
	}	
}));

##############################################################

ENDDOCUMENT();