## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## KEYWORDS ('derivative','definition')
## Tagged by cmd6a 4/20/06

## DBsubject('Calculus')
## DBchapter('Limits and Derivatives')
## DBsection('Definition of the Derivative')
## Date('')
## Author('')
## Institution('Rochester')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
	"PGstandard.pl",
	"MathObjects.pl",
        "parserMultiAnswer.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Numeric")->variables->add(t=>'Real');

$a = random(2,9,1);
$b = random(2,7,1);

$ans1 = Formula("[($b /(t - $a))- ($b /(x - $a))] /(t - x)")->reduce;

$Num = Formula("$b *(t - x)");
$Denom = Formula("(x - t)(x - $a)(t - $a)");

#Checks against common factors in both num and denom (including negatives):
$multians2 = MultiAnswer($Num, $Denom)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $NumStu, $DenomStu ) = @{$student};      
      my $context = Context()->copy;
      
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $NumStu2 = Formula($context,$NumStu);
      $Num2 = Formula($context,"$c0 ($b *(t - x))");
            
      $DenomStu2 = Formula($context,$DenomStu);
      $Denom2 = Formula($context,"$c0 ((x - t)(x - $a)(t - $a))");

      if ((($Num2 == $NumStu2) && ($Denom2 == $DenomStu2)) &&
          ((($Num) / ($Denom)) == (($NumStu)/($DenomStu)))) {
          return 1; } 
      else {
          return 0; }          
  } );


$comFactor = Formula("x - t");

$NumC = Formula("$b");
$DenomC = Formula("($a - x)(t - $a)");

#Checks against common factors in both num and denom (including negatives):
$multians4 = MultiAnswer($NumC, $DenomC)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $NumCStu, $DenomCStu ) = @{$student};

      my $context = Context()->copy;
      
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $NumCStu2 = Formula($context,$NumCStu);
      $NumC2 = Formula($context,"$c0 ($b)");
            
      $DenomCStu2 = Formula($context,$DenomCStu);
      $DenomC2 = Formula($context,"$c0 (($a - x)(t - $a))");

      if ((($NumC2 == $NumCStu2) && ($DenomC2 == $DenomCStu2)) &&
          ((($NumC) / ($DenomC)) == (($NumCStu)/($DenomCStu)))) {
          return 1; } 
      else {
          return 0; }          
  } );

$Deriv = Formula("-$b / (x - $a)^2");

Context()->texStrings;
BEGIN_TEXT
Let \( \displaystyle f(x)=\frac{$b}{x- $a}\). $BR 
Then according to the definition of derivative $BR
\( \displaystyle f'(x) = \lim_{t\to x}\) \{ans_rule(45)\}$BR
(Your answer above and the next few answers below will involve the variables
\(t\) and \(x\).)$BR $BR

$BBOLD Note: $EBOLD In the following parts, each pair of numerators and denominators count as $BBOLD one $EBOLD answer, thus you must fill in $BBOLD both $EBOLD to receive credit. $BR $BR
The expression inside the limit simplifies to a simple fraction with $BR
numerator \(=\) \{$multians2->ans_rule(30)\} $BR
and denominator \(=\) \{$multians2->ans_rule(30)\}. $BR $BR
We can cancel the factor \{ans_rule(30)\} appearing in the denominator against
a similar factor appearing in the numerator leaving a simpler fraction with $BR
numerator \(=\) \{$multians4->ans_rule(30)\} $BR
and denominator \(=\) \{$multians4->ans_rule(30)\} $BR $BR
Taking the limit of this fractional expression gives us $BR
\(f'(x) = \) \{ans_rule(30)\}
END_TEXT
Context()->normalStrings;

ANS($ans1->cmp);
ANS($multians2->cmp());

#Accepts +- correct common factor.
ANS($comFactor->cmp( checker=>sub {
      my ( $correct, $student, $ansHash ) = @_;
      return (($correct == $student) || ($correct == -($student)));
          
   } ) );
ANS($multians4->cmp());
ANS($Deriv->cmp());


ENDDOCUMENT();        # This should be the last executable line in the problem.
