##DESCRIPTION
##   Describe how a confidence interval changes as the confidence level increases
##ENDDESCRIPTION

##KEYWORDS('statistics', 'sampling')

## DBsubject('Statistics')
## DBchapter('Estimation')
## DBsection('Confidence Intervals')  
## Level(4)
## Date('10/09/2012')
## Author('Djun Kim')
## Institution('UBC')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGnumericalmacros.pl",
  "PGstatisticsmacros.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());


##############################################################
#
#  Setup
#
#
Context("Numeric");

$num = random(30,40,1);
$n = random(330, 340, 1);

$mc1 = new_multiple_choice();
$mc1->qa(
"The proportion of the $n people who have recently had the flu, $num/$n, is a: ", 
"statistic."
);
$mc1->extra(
"parameter.",
"variable of interest.",
);


# Part ii)
# Acceptable answers use z = 2 or z = 1.96

$phat = $num / $n;

$left_endpoint = $phat - 2*sqrt($phat * (1-$phat) / $n);
$right_endpoint = $phat + 2*sqrt($phat * (1-$phat) / $n);

$left_endpoint = sprintf "%.3f", $left_endpoint;
$right_endpoint = sprintf "%.3f", $right_endpoint;

$left_ans1 = $phat - 2*sqrt($phat * (1-$phat) / $n);
$right_ans1 = $phat + 2*sqrt($phat * (1-$phat) / $n);

$left_ans1 = sprintf "%.3f", $left_ans1 ;
$right_ans1 = sprintf "%.3f", $right_ans1 ;

$left_ans2 = $phat - 1.96*sqrt($phat * (1-$phat) / $n);
$right_ans2 = $phat + 1.96*sqrt($phat * (1-$phat) / $n);

$left_ans2 = sprintf "%.3f", $left_ans2 ;
$right_ans2  = sprintf "%.3f", $right_ans2 ;

$multians = MultiAnswer($left_endpoint, $right_endpoint)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $anshash, $self ) = @_;
      my ( $stu_ans_left, $stu_ans_right ) = @{$student};
      my ( $ans1, $ans2 ) = @{$correct};
      $ansHash->{correct_ans} = "($left_ans1, $right_ans1) OR ($left_ans2, $right_ans2)";
      if ( ($stu_ans_left == $left_ans1 && $stu_ans_right == $right_ans1) ||
           ($stu_ans_left == $left_ans2 && $stu_ans_right == $right_ans2) ) 
      {
          return [1,1];
      } 
      else {
          return [0,0];
      }
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

An epidemiologist is worried about the prevalence of the flu in East Vancouver and the potential shortage of vaccines for the area. She will need to provide a recommendation for how to allocate the vaccines appropriately across the city. She takes a simple random sample of $n people living in East Vancouver and finds that $num have recently had the flu.

$PAR
$BBOLD Part i) $EBOLD
\{ $mc1->print_q() \}
$BR
\{ $mc1->print_a() \}


$PAR
$BBOLD Part ii) $EBOLD
Use the sample data to compute a 95$PERCENT confidence interval for the true proportion of East Vancouver residents that have recently had the flu.

$PAR
(Please carry answers to at least six decimal places in intermediate steps. Give your final answer to the nearest three decimal places).

$PAR
95$PERCENT confidence interval = (\{ $multians->ans_rule(6) \}, \{ $multians->ans_rule(6) \})

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 0;

ANS( radio_cmp( $mc1->correct_ans() ) );
ANS( $multians->cmp() ); 



ENDDOCUMENT();        
