# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('8.1')
## Problem1('2')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Matrix");

$a = random(2,6,1);
$mat = Matrix([[ 0, $a ], [ $a, 0 ]] );
$lambda1 = Compute( $a );
$lambda2 = Compute( -$a );
$d1 = Matrix( [[ $lambda1, 0 ], [ 0, $lambda2 ]] );
$d2 = Matrix( [[ $lambda2, 0 ], [ 0, $lambda1 ]] );
$smat = Matrix([[ "1/sqrt(2)", "-1/sqrt(2)" ], [ "1/sqrt(2)", "1/sqrt(2)" ]]);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT
Consider the matrix \( $mat \).  Find an orthogonal \(S\) such that
\( S^{-1} A S = D \), a diagonal matrix.
$BR
\( S = \) \{ $smat->ans_array() \}

END_TEXT
Context()->normalStrings;

ANS( $smat->cmp( checker=>sub {
    my ( $c, $s, $ansHash ) = @_;
    $st = $s->transpose();
    return $d1 == $st*$mat*$s || $d2 == $st*$mat*$s;
} ) );

ENDDOCUMENT();
