## DESCRIPTION
## Compute a rank decomposition of a 3x3 rank-1 matrix A.
## That is, find a 3x1 C and a 1x3 R such that A = C*R.
##   The matrix has small integer entries, and there will be a solution with small integer entries.
##   There may be infinitely many different possible solutions to this problem. All of them are accepted as correct.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Matrices')
## DBsection('Echelon form')
## Date('March 25, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('2.C: More About the Rank')
## Problem1('')
## KEYWORDS('linear algebra','matrix','reduced row echelon form','rank decomposition')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Matrix");

do {
    $c1 = random(-2,3);
    $c4 = random(-2,3);
    $c7 = random(-2,3);
    $C = Matrix([[$c1],[$c4],[$c7]]);
    $D = Matrix([[$c1,random(-2,3),random(-2,3)],[$c4,random(-2,3),random(-2,3)],[$c7,random(-2,3),random(-2,3)]]);
        
    $R = Matrix([[1,random(-2,4),random(-2,4)]]);
        
    $A = $C * $R;
} until (abs($D->det) > 0.1);

$multians = MultiAnswer($C,$R)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if($s0*$s1 == $A) {
            return 1;
        } else {
            return 0;
        }
    }
);

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Compute a rank decomposition of the (rank \( 1 \)) matrix \( \displaystyle A = $A \). That is, find \( C \in \mathcal{M}_{3,1} \) and \( R \in \mathcal{M}_{1,3} \) such that \( A = CR \).
$BR
$BR
\( \displaystyle C = \) \{ $multians->ans_array() \}
$BR
\( \displaystyle R = \) \{ $multians->ans_array() \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
