# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

## KEYWORDS('vectors', 'coordinate vector')
## Tagged by glr 06/03/09


## DBsubject('Calculus')
## DBchapter('Vectors and the Geometry of Space')
## DBsection('Three-Dimensional Coordinate Systems')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('13.1')
## Problem1('17')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')
## HHSection1('Displacement Vectors')

Context("Vector2D");
$showPartialCorrectAnswers = 1;

$v = 3*i + 2*j;
# $v = Compute("<3,2>");

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

\($v\)
$PAR
Parallel: \{ ans_rule(25) \}
$BR
Perpendicular: \{ ans_rule(25) \}

END_TEXT
Context()->normalStrings;

ANS( $v->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student->isParallel($correct);
} ) );

ANS( $v->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student.$correct == 0;
} ) );

ENDDOCUMENT();
