#########################################################
# DESCRIPTION
## DBsubject(Statistics)
## DBchapter(Hypothesis tests)
## DBsection(Two sample t test)
## level(4)

#########################################################
DOCUMENT();
## Initializations: (Required)
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "answerCustom.pl",
  "PGcourse.pl"
);

#########################################################
## Setup: this is where we use Perl and PG objects (Required)

Context("Numeric");
# warn("Student name: $studentName, login: $studentLogin, id: $studentID");
my $hash = crypt($studentLogin, 'a1');
# warn("crypt: $hash");
@nums = (0..9,'a'..'z','A'..'Z');
my %nums = map { $nums[$_] => $_ } 0..$#nums;
my $seed = 0;
$seed = $seed * 62 + $nums{$_} foreach split(//, substr($hash, -5, 5));
# warn("seed: " . $seed);

rserve_eval("set.seed($seed)");
## fix the name of the downloadable data files
rserve_eval("filename <- file.path(tempdir(), 'GPAdata_${studentLogin}.csv')");

my @remote_file = rserve_eval('
#### this is to create the "GPDdata" data for Q1
# this is the number of observations
N <- sample(x=35:50, size=1)

# this is to sample "N" GPAs
GPA <- sample(x=seq(from=2.2, to=4, by=0.01), size=N, replace=F)

# this is to create the categorical representation of GPA
GPAcat <- rep(0, N)
for (i in 1:N){
  if (GPA[i] < 2.5){GPAcat[i] <- "A"}
  if (GPA[i] >= 2.5 & GPA[i]<3){GPAcat[i] <- "B"}
  if (GPA[i] >= 3 & GPA[i]<3.5){GPAcat[i] <- "C"}
  if (GPA[i] >= 3.5){GPAcat[i] <- "D"}
}

# this is to generate salaries as a quadratic function of GPA
Salary = -147100 + 118000*GPA - 16900*(GPA^2) + rnorm(n=N, mean=0, sd=4000)

dat <- data.frame(Salary = Salary, GPA = GPA, GPAcat = GPAcat)

write.csv(dat, filename, row.names=FALSE); 
filename
');
my $local_file = rserve_get_file($remote_file[0]);
$local_url = $local_file;
$local_url =~ s|$tempDirectory|$tempURL|;


rserve_eval("set.seed($seed)");
@data = rserve_eval('
#### this is to create the "GPDdata" data for Q1
# this is the number of observations
N <- sample(x=35:50, size=1)

# this is to sample "N" GPAs
GPA <- sample(x=seq(from=2.2, to=4, by=0.01), size=N, replace=F)

# this is to create the categorical representation of GPA
GPAcat <- rep(0, N)
for (i in 1:N){
  if (GPA[i] < 2.5){GPAcat[i] <- "A"}
  if (GPA[i] >= 2.5 & GPA[i]<3){GPAcat[i] <- "B"}
  if (GPA[i] >= 3 & GPA[i]<3.5){GPAcat[i] <- "C"}
  if (GPA[i] >= 3.5){GPAcat[i] <- "D"}
}

# this is to generate salaries as a quadratic function of GPA
Salary = -147100 + 118000*GPA - 16900*(GPA^2) + rnorm(n=N, mean=0, sd=4000)

# this is to generate the value for Part (f)
Part.f <- sample(seq(1.0, 1.9, 0.1), 1)

### Solutions
# part (b)
b <- round(cor(GPA, Salary, method="pearson"), 3)
# part (c)
c <- round(cor(GPA, Salary, method="spearman"), 3)
# part(e)
# intercept
e.inter <- round(coef(lm(Salary~GPA))[1], 0)
# slope
e.slope <- round(coef(lm(Salary~GPA))[2], 0)

# part (f)
# this will calculate the predicted value without rounding
# if rounded, the answer can onluy differ by about 2, so we can
# use a tolerance of 10, just to be overly safe...
f <- coef(lm(Salary~GPA))[1] + coef(lm(Salary~GPA))[2]*Part.f

# return answers
c(Part.f, b, c, e.inter, e.slope, f);
');

$partf = $data[0];
$bans = $data[1];
$cans = $data[2];
$fans = $data[5];
$fans_integer = round($fans, 0);

$multians_e = MultiAnswer($data[3], $data[4])->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu_a1, $stu_a2 ) = @{$student};
      if ($stu_a1 == $data[3]) {   
         if ($stu_a2 == $data[4]) {
            return [1, 1];
         } else {return [1, 0];}
      } else {
                if ($stu_a2 == $data[4]) {
                   return [0, 1];
                } else {return [0, 0];}
        }
  }
);


#########################################################
## Main Text: where all text goes (Required)
Context()->texStrings;
BEGIN_TEXT
This assignment will be submitted in two parts.  Some answers will be submitted directly to the online homework system, and these questions will include a note saying ($BBOLD WebWorkiR $EBOLD) at the start.  Other questions will be submitted on paper, in class.  These questions will include a note saying ($BBOLD On-Paper $EBOLD) at the start.
$PAR

The dataset  $BBOLD GPAdata.csv (\{ htmlLink($local_url, "Download the csv file here") \}) $EBOLD contains information collected for a simple random sample of recent university graduates.   For each individual in the study, we have the initial starting salary (Salary), the cumulative grade point average (GPA), as well as a categorical representation of the GPA (GPAcat: "A" = less than 2.5, "B" = 2.5 - 2.99, "C" = 3 - 3.49, "D" = 3.5 - 4).  We would like to examine the relationship between starting salary and grade point average.  Use this data to answer the following questions. 
$PAR

$BBOLD Part (a) (On-Paper) $EBOLD Produce plots to examine the relationship between GPA and Salary as well as the categorical representation of GPA and salary.  Make sure to comment on the relationship, and anything else you observe in these plots.
$BR
$BR
$BBOLD Part (b) (WebWorkiR) $EBOLD Calculate and report Pearson's correlation between GPA and Salary.  Please round your answer to 3 decimal places. \{ ans_rule(7) \}
$BR
$BR
$BBOLD Part (c) (WebWorkiR) $EBOLD Calculate and report Spearman's correlation between GPA and Salary.  Please round your answer to 3 decimal places. \{ ans_rule(7) \}
$BR
$BR
$BBOLD Part (d) (On-Paper) $EBOLD Do you think either of Pearson's or Spearman's correlation is more appropriate for describing the relationship observed in this data, or do you think either is appropriate to report?  Make sure to justify your answer.
$BR
$BR
$BBOLD Part (e) (WebWorkiR) $EBOLD Fit a simple linear regression relating GPA to Salary, using Salary as the outcome.  Report the model slope and intercept, rounded to a whole number with 0 decimal places.
$BR
intercept = \{$multians_e->ans_rule(7)\}, $SPACE    slope= \{$multians_e->ans_rule(7)\}
$BR
$BR
$BBOLD Part (f) (WebWorkiR) $EBOLD Using the fitted regression model, what would you predict the mean starting salary to be for those with a GPA of $partf?  Please round your answer to a whole number, with 0 decimal places.
\{ ans_rule(7) \}
$BR
$BR
$BBOLD Part (g) (On-Paper) $EBOLD Why might one be cautious in believing the predicted mean starting salary for someone with a GPA of $partf, calculated in part (f)?  Can you think of a second reason one might also be cautious of this predicted value?
$BR
$BR
$BBOLD Part (h) (On-Paper) $EBOLD In words, describe what it would mean to attach a 95\(%\) prediction interval to the predicted mean salary from part (f).
$BR
$BR
$BBOLD Part (i) (On-Paper) $EBOLD  Provide an interpretation of the model slope and intercept.  What could one do to give the intercept a more meaningful interpretation?
$BR
$BR
$BBOLD Part (j) (On-Paper) $EBOLD  Provide plots to examine the model assumptions, and comment on whether or not you believe each of the model assumptions to be met. (NOTE: one really should check model assumptions before interpreting the model)
$BR
$BR
$BBOLD Part (k) (On-Paper) BONUS: $EBOLD One or more of the assumptions were likely not met.  Try fitting a model using an "X" of log(GPA) instead (in R, this can be calculated using: log(GPA)).  Does this appear to have solved the problem(s)?  Now try fitting a polynomial model that uses both GPA and GPA\(^2\).  Does this appear to have solved the problem(s)?  Provide some discussion on both of these potential solutions.
$BR
$BR
$BBOLD Part (l) (On-Paper) BONUS: $EBOLD The relationship between GPA and Salary may not have appeared linear.  In the previous part, we tried fitting a model using a "transformed X".  Here we will try another approach.  Try examining the relationship between grade point average and salary using the categorical representation of GPA instead (GPAcat).  Try fitting and ANOVA model as well as a simple linear regression that uses GPAcat instead.  What can you say about the relationship between categorical GPA and Salary based on these two models?  Can you see any of the connections between the ANOVA model and the SLR model?  See if you can make sense of the model coefficients in SLR, and how these relate to the mean salary for each GPA category.  Can you see a connection between the MSE from ANOVA and the Residual SE from SLR?

END_TEXT

#########################################################
## Answers evaluation (Required)
ANS( num_cmp($bans, tol=> 0.001) );
ANS( num_cmp($cans, tol=> 0.001) );
ANS( $multians_e->cmp() ); 
ANS( num_cmp($fans, tol=> 10) );

#########################################################
## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
$BBOLD Part (a)  $EBOLD
Answer submitted ON-PAPER
$BR
$BBOLD Part (b)  $EBOLD
$bans
$BR
$BBOLD Part (c)  $EBOLD
$cans
$BR
$BBOLD Part (d)  $EBOLD
Answer submitted ON-PAPER
$BR
$BBOLD Part (e)  $EBOLD
intercept = $data[3], slope = $data[4]
$BR
$BBOLD Part (f)  $EBOLD
$fans_integer 
$BR
$BBOLD Part (g)-(l)  $EBOLD
All answers submitted On-Paper

END_SOLUTION

ENDDOCUMENT();
