##DESCRIPTION
##  Algebra problem: state Zero Product Property
##ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Quadratic equations and functions)
## DBsection(Solve by factoring)
## Level(2)
## KEYWORDS('factoring')
## Author(K. Andrew Parker)
## Institution(CUNY City Tech)
## Language(en)

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->variables->are(A=>"Real",B=>"Real");

$ans1 = Formula("A");
$ans2 = Formula("B");

$ans = MultiAnswer($ans1,$ans2);

##############################################################
#
#  Text
#
#

BEGIN_PGML

If we are told that [` AB = 0 `], then what can we infer?

By the zero product property we know:  
[_____][`=0`] OR [_____][`=0`]

END_PGML

##############################################################
#
#  Answers
#
#

ANS($ans->cmp(checker=>sub{
     my ($correct,$student,$self) = @_;  # get the parameters
     my ($stu1,$stu2) = @{$student};      # extract the student answers
     my ($ans1,$ans2) = @{$correct};
     return ( ($ans1==$stu1 && $ans2==$stu2) || ($ans2==$stu1 && $ans1==$stu2) );
}));

BEGIN_PGML_HINT

Remember that the zero product property states:  
"If the product of two values is zero, then at least one of the values must be zero."

END_PGML_HINT
BEGIN_PGML_SOLUTION

If [` AB = 0 `], then either [`A=0`] or [`B=0`].

END_PGML_SOLUTION

COMMENT("Funded by US DoE Title V: Opening Gateways grant.");

ENDDOCUMENT();        
