## DESCRIPTION
## Quadratic functions, expressions, and equations
## ENDDESCRIPTION

## KEYWORDS('quadratic functions')

## DBsubject('Algebra') 
## BookTitle('Algebra: Form and Function')
## DBchapter('Quadratic Functions') 
## BookChapter('Quadratic Functions, Expressions, and Equations')
## DBsection('Factoring Quadratics') 
## BookSection('Solving Quadratic Equations by Factoring')
## Date('01/01/10')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Algebra: Form and Function')
## EditionText1('1')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## Section1('9.4')
## Problem1('28')



####################################
#  Initialization

DOCUMENT();
loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"PGcourse.pl",
"contextLimitedPowers.pl",
"contextPolynomialFactors.pl",
);

TEXT(beginproblem()); 

####################################
#  Setup

Context("PolynomialFactors-Strict");

Context()->functions->enable("sqrt");

LimitedPowers::OnlyIntegers(
#minPower => 0, maxPower => 2,
checker => sub {
  my $self = shift;
  my $p = shift;
  # acceptable powers
  return $p==1 || $p==1 || $p==2 || $p==0.5;
},
message => "either 0, 0.5, 1, or 2",
);


$a = list_random(2,3,5,6,7,8);
$b = list_random(2,3,5,6,7,8);
while ($a == $b) { $b = list_random(2,3,5,6,7,8); }

$answer = Formula("(x - sqrt($a)) (x + sqrt($b))");


######################################
#  Main text

BEGIN_TEXT
Write a quadratic equation in the variable \( x \) 
with the solutions \( \sqrt{$a} \) and \( -\sqrt{$b} \).
$BR
$BR 
\( y \) = \{ ans_rule(30) \}
\{ AnswerFormatHelp("equations") \}
END_TEXT


######################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"$c0 * (x - sqrt($a)) (x + sqrt($b))");
      return $correct == $student;
    }
) );



COMMENT('MathObject version');
ENDDOCUMENT();
