# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Optimization - engineering and physics)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(5)
## Static(1)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('4.3')
## Section2('4.2')
## Section3('4.2')
## Section4('4.2')
## Problem1('25')
## Problem2('30')
## Problem3('36')
## Problem4('42')
## KEYWORDS('calculus', 'derivative', 'local maxima and minima', 'maxima', 'minima')

## Textbook tags
## HHChapter1('Using the Derivative')
## HHChapter2('Using the Derivative')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Numeric")->variables->are( r=>'Real', s=>'Real', R=>'Real' );

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context()->texStrings;
BEGIN_TEXT

When an electric current passes through two resistors with
resistance \( r \) and \( s \), connected in parallel, the combined
resistance, \( R \), can be calculated from the equation
\[ {1 \over R} = {1 \over {r}} + {1 \over {s}}, \]

where  \(R\),  \(r\) , and  \(s\)  are positive.
Assume that \(s\) is constant.

$PAR

Find \( dR\over dr \):$BR
\( {dR\over dr} = \) \{ ans_rule(65) \}

$PAR

Is \( R \) and increasing or decreasing function of \( r \)?
\{ ans_rule(15) \}
$BR
(Enter ${BBOLD}increasing$EBOLD,
${BBOLD}decreasing$EBOLD, ${BBOLD}neither$EBOLD, or 
${BBOLD}both$EBOLD (write both if there are values of \( r \) for
which \( R \) is increasing, and other values for which it is
decreasing; enter neither if this is a constant function.)

$PAR

If we consider the interval \( a \le r \le b \), where does 
\( R \) take on its global maximum and minimum values?  
$BR
maximum: \( r = \) \{ ans_rule(25) \} $BR
minimum: \( r = \) \{ ans_rule(25) \} $BR

(Enter ${BBOLD}none$EBOLD if there is no global maximum or
minimum for this function.)

END_TEXT
Context()->normalStrings;

# allow implicit differentiation value
ANS( Compute( 's^2/((r+s)^2)' )->cmp( checker=>sub {
	my ( $c, $s, $ans ) = @_;
        return $c == $s ||
               Compute("R^2/r^2") == $s } ) );
ANS( str_cmp( 'increasing' ) );

Context()->variables->add( a=>'Real', b=>'Real' );
ANS( Compute( 'b' )->cmp() );
ANS( Compute( 'a' )->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

To find \( dR/dr \) we first solve for \(R\), and then differentiate
using the quotient rule.
Solving for \(R\), we get
\[
 R = \frac{1}{\frac1r + \frac1s} =
  \frac{rs}{rs(\frac1r + \frac1s)} = \frac{rs}{s + r}.
\]
Then
\[
\frac{dR}{dr} = \frac{s(s+r) - rs}{(s+r)^2} = \frac{s^2}{(s+r)^2}.
\]
This is a ratio of two squared terms which can't be zero, so the
result is clearly positive.  Thus the derivative is always positive,
and \( R \) is always increasing.  Therefore the
global maximum occurs for the largest value of \( r \),
\( r = b \), and the global minimum for the smallest, \( r = a \).

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
