# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Summary of curve sketching)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Applied Calculus')
## TitleText5('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## AuthorText5('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('4')
## EditionText5('7')
## Section1('4.3')
## Section2('4.2')
## Section3('4.2')
## Section4('4.3')
## Section5('4.2')
## Problem1('35')
## Problem2('40')
## Problem3('40')
## Problem4('31')
## Problem5('47')
## KEYWORDS('calculus', 'derivative', 'local maxima and minima', 'maxima', 'minima')

## Textbook tags
## HHChapter1('Using the Derivative')
## HHChapter2('Using the Derivative')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "extraAnswerEvaluators.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
 
Context("Numeric");

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

# we plot the curve 
#  1) h = -0.5 cos(pi x/2) - 0.5, on [-2:2] [-1.25:0.25],
# which has a minimum at x=0 of -1 and goes to zero (and has zero derivatives)
# at each of x = -2 and x = 2, or
#  2) h1 = -6 - 18x - 47x^2/2 - 13x^3 - 5x^4/2, on [-2:-1] [-1.25:0.25] and
#     h2 = -2/3 + 14x/27 + x^2/18 - x^3/9 + x^4/54, on [-1:2] (etc.)
# which has a minimum at x=-1 of -1 and the same end behavior as before, or
#  3) the same piecewise function, mirrored around x = 0 (x -> -x), or
#  4-6) these functions mirrored around y = 0.

# the factors for x and y reflection
$xfac = random(-1,1,1);
$yfac = non_zero_random(-1,1,1);

# initialize the graph accordingly
if ( $yfac == 1 ) {
    $gr = init_graph(-2,-1.25,2,0.25,axes=>[0,0],grid=>[4,1]);
} else {
    $gr = init_graph(-2,-0.25,2,1.25,axes=>[0,0],grid=>[4,1]);
}

# define the function
if ( $xfac == 0 ) {
    $h = "$yfac*(-0.5*cos(pi*x/2)-0.5) for x in <-2,2> " .
	"using color:blue and weight:2";
    add_functions( $gr, $h );
    $xinfl = 0;
} else {
    if ( $xfac == 1 ) {
        $x0 = -2;         # x0 = the starting point for h1
	$x1 = -1;         # x1 = the ending point for h1
        $x2 = -1;         # x2 = the starting point for h2
        $x3 = 2;          # x3 = the ending point for h2
        $xinfl = -1;
    } else {
	$x0 = 1;
	$x1 = 2;
        $x2 = -2;
        $x3 = 1;
        $xinfl = 1;
    }
    $xv = "($xfac*x)";
    $h1 = "$yfac*(-6 - 18*$xv - 47*$xv^2/2 - 13*$xv^3 - 5$xv^4/2)" .
	"for x in <$x0,$x1> using color:blue and weight:2";
    $h2 = "$yfac*(-2/3 + 14*$xv/27 + $xv^2/18 - $xv^3/9 + $xv^4/54)" .
	"for x in <$x2,$x3> using color:blue and weight:2";
    add_functions( $gr, $h1, $h2 );
}

# the value we say g(x) has at x=-2
$g0 = non_zero_random(-10,10,1);
# random other values
$g1 = non_zero_random(-10,5,1);
$g2 = $g1 + random(1,5,1);
$g3 = random(-5,5,1);

Context()->texStrings;
BEGIN_TEXT
The figure below gives the behavior of the derivative of \(g(x)\) on
\(-2\le x\le 2\).
$PAR
$BCENTER
\{ image( insertGraph($gr), 'tex_size'=>500 ) \}
$BR
${BITALIC}Graph of \(g'(x)\) (${BBOLD}not$EBOLD \(g(x)\))$EITALIC
$BR
${BITALIC}(Click on the graph to get a larger version.)$EITALIC
$ECENTER
$PAR
Sketch a graph of \(g(x)\) and use your sketch to answer the following
questions. 
$PAR
${BBOLD}A.$EBOLD
Where does the graph of \(g(x)\) have inflection points? 
$BR
\(x = \) \{ ans_rule(15) \}
$BR
${BITALIC}Enter your answer as a comma-separated list of values, or enter
${BBOLD}none$EBOLD if there are none.$EITALIC
$PAR

${BBOLD}B.$EBOLD
Where are the global maxima and minima of \(g\) on \([-2,2]\)?
$BR
minimum at \(x = \) \{ ans_rule(5) \}
$BR
maximum at \(x = \) \{ ans_rule(5) \}
$PAR

${BBOLD}C.$EBOLD
If \(g(-2) = $g0\), what are possible values for \(g(0)\)?
$BR
\(g(0)\) is in \{ ans_rule(15) \}
$BR
${BITALIC}(Enter your answer as an interval, or union of intervals, giving
the possible values.  Thus if you know \($g1 < g(0) \le $g2\), enter
${BBOLD}($g1,$g2]$EBOLD.  Enter ${BBOLD}infinity$EBOLD for
\(\infty\), the interval ${BBOLD}[$g3,$g3]$EBOLD to indicate a single
point).$EITALIC

$PAR
How is the value of \(g(2)\) related to the value of \(g(0)\)?
$BR
\(g(2)\) \{ ans_rule(5) \} \(g(0)\)
$BR
${BITALIC}(Enter the appropriate mathematical equality or inequality, 
\(=\), \(<\), \(>\), etc.)$EITALIC

END_TEXT
Context()->normalStrings;

# set the correct infinity
$infr = ( $yfac == 1 ) ? "(-infinity,$g0)" : "($g0,infinity)";
## note that if $yfac == 1, then the derivative is everywhere negative;
##    otherwise, everywhere positive.  and in magnitude, it appears in 
##    any case to be less than one.  so if $yfac == 1, we can actually
##    also say that g must be between $g0 - 2 and $g0, and if 
##    $yfac == -1, between g0 and g0 + 2.
$gint = Interval( $infr );

ANS(number_list_cmp( $xinfl, 'strings'=>['none'] ) );
ANS(num_cmp( "($yfac)*(2)" ) );
ANS(num_cmp( "($yfac)*(-2)" ) );
# ANS(interval_cmp($infr) );
ANS($gint->cmp( checker=>sub {
	my ($c, $s, $ansHash) = @_;
	if ( $c == $s ) { return 1; }
	else {
	    my @corCoord = $c->value();
	    my @stuCoord = $s->value();
	    if ( $yfac == 1 ) {
		if ( $corCoord[1] == $stuCoord[1] &&
		     Compute($g0-2) >= $stuCoord[0] ) {
		    return 1;
		} else {
		    return 0;
		}
	    } else {
                if ( $corCoord[0] == $stuCoord[0] &&
                     Compute($g0+2) <= $stuCoord[1] ) {
                    return 1;
                } else {
                    return 0;
                }
	    }
	}
}) );
ANS(str_cmp( ($yfac == 1) ? '<' : '>' ) );

# some useful solution variables
if ( $yfac == 1 ) {
    $sgn = 'negative';
    $incr = 'decreasing';
    $decr = 'increasing';
    $inc = 'decreases';
    $grless = 'less';
    $inf = '-\infty';
    $ltsgn = '<';
    $gtsgn = '>';

    $g02 = $g0-2;
    $addedExpl = "We note that it appears that the magnitude of the " .
	"derivative \(g'\) is also less than one for the entire domain " .
	"that we are considering, so we can actually give a smaller " .
	"interval for the possible values of \(g\); any interval " .
	"between \((-\infty,$g0)\) and \(($g02,$g0)\) is reasonable.";
} else {
    $sgn = 'positive';
    $incr = 'increasing';
    $decr = 'decreasing';
    $inc = 'increases';
    $grless = 'greater';
    $inf = '\infty';
    $ltsgn = '>';
    $gtsgn = '<';

    $g02 = $g0+2;
    $addedExpl = "We note that it appears that the magnitude of the " .
	"derivative \(g'\) is also less than one for the entire domain " .
	"that we are considering, so we can actually give a smaller " .
	"interval for the possible values of \(g\); any interval " .
	"between \(($g0,\infty)\) and \(($g0,$g02)\) is reasonable.";
}
$gmin = $yfac*(2);
$gmax = $yfac*(-2);
$infrfmt = $infr;
$infrfmt =~ s/infinity/~~~~infty/;

Context()->texStrings;
BEGIN_SOLUTION
$PAR SOLUTION $PAR

Because the graph of \(g'(x)\) is $sgn for all \(-2 < x < 2\), we know
that the graph of \(g(x)\) must always be $incr there.  The derivative
appears to be zero at \(x = \pm 2\), so the graph of \(g(x)\) must start
out flat at \(x = -2\), and then $inc more and more rapidly as it
approaches \(x = $xinfl\).  After \(x = $xinfl\) the it $inc less rapidly,
finally flattening out entirely as we get to \(x = 2\).
$PAR

${BBOLD}A.$EBOLD
The graph has an inflection point at \(x = $xinfl\), where the slope
changes from $sgn and $incr to $sgn and $decr.
$PAR

${BBOLD}B.$EBOLD
The function has a global maximum at \(x=$gmax\) and a global minimum at
\(x=$gmin\).
$PAR

${BBOLD}C.$EBOLD
Since the function is $incr over the interval \(-2 < x < 2\), we know
that \(g(0)\) must be $grless than \(g(-2)\), and so must be in the
interval \($infrfmt\).  ($addedExpl)
$PAR

${BBOLD}D.$EBOLD
Again, since the function is $incr over the interval \(-2 < x < 2\), we
know that \(g(-2) = $g0 $gtsgn g(0) $gtsgn g(2)\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
