## DESCRIPTION
## Calculus
## ENDDESCRIPTION

## Tagged by YL




## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Integral test)
## Institution(NAU)
## Author(Utah ww group)
## MLT(IntegralToLimitErrorInPartialSum)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('5e')
## Section1('11.3')
## Problem1('')
## TitleText2('Calculus: Early Transcendentals')
## AuthorText2('Stewart')
## EditionText2('6')
## Section2('11.3')
## Problem2('')
## TitleText3('Calculus')
## AuthorText3('Dale Varberg, Edwin J. Purcell, and Steve E. Rigdon')
## EditionText3('9')
## Section3('.')
## Problem3('')
## KEYWORDS('series', 'convergence', 'converge', 'partial sum', 'estimate', 'error' )

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGasu.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context('Numeric');

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$b = random(2,10,1);
do  {$a = random( 1+5*$b , 7*$b , 1)}  until  ( gcd($a,$b) == 1 );
$c = random(3,5,1);
do  {$d = random(3,5,1)}  until  ($d != $c) ;
$c1 = $c - 1 ;

if ($d==3) {
	$nth = '3rd';
} elsif ($d==4) {
	$nth = '4th';
} elsif ($d==5) {
	$nth = '5th';
} else {
	$nth = '6th';
}


##  Partial sum: do summation in reverse order, i.e., from small to larger
$sum = 0 ;
foreach  my $k  (1 .. $d) {
  $n = 1 + $d - $k ;
  $sum += 1/($n**$c) ;
}
$sum *= $a/$b ;


$err2 = $a/($b*$c1*($d+1)**$c1);
$tolbnd = $a/($b*$d**$c)/3;

$ans1 = Real($a/($b*$c1*$d**$c1));

Context()->flags->set(tolerance=>$tolbnd);
$sum1 = Real($sum);

Context()->flags->set(tolerance=>($err2/2));
$ans2 = Real($sum + $err2);
$ans3 = Real($sum + $ans1);

$mp = MultiAnswer($sum1, $ans1, $ans2, $ans3)->with(
 singleResult => 0,
 allowBlankAnswers => 1,
 checker => sub {
	 my ($correct,$student,$self) = @_;  # get the parameters
	 my ($s,$a1,$a2,$a3) = @{$student};      # extract the student answers
	 my ($sc, $a1c, $a2c, $a3c) = @{$correct}; # correct answers
	 my @res = (0,0,0,0); #results
	 $res[0] = 1 if $s == $sc;
	 $res[1] = 1 if $a1 == $a1c;
	 $res[2] = 1 if $a2 == $a2c;
	 $res[3] = 1 if $a3 == $a3c;
	 $res[3] = 1 if ($res[3] == 0 and $s==$sc and $a1 == $a1c and ($a3 == $s+$a1));
	 return @res;
 }
);
												 
	

BEGIN_TEXT

\( (a) \) Compute \(s_{$d} \) (the $nth partial sum) of
	\( \displaystyle s=\sum_{n=1}^{\infty}\frac{$a}{$b \, n^$c}  \) :
$BR
	\(s_{$d} = \) \{$mp->ans_rule(45)\} 
$BR $BR
$BBOLD Note: $EBOLD if giving a decimal approximation to this part, you may 
need to enter more digits than usual.
$BR $BR
\( (b) \) Estimate the error in using \(s_{$d}\) as an approximation of the sum
	of the series (i.e. use \( \displaystyle \int_{$d}^{\infty} f(x)\, dx \geq R_{$d} \) ): $BR
	Estimate = \{$mp->ans_rule(30)\} 
$BR
$BR
\( (c) \) Use \(n = $d\) and
\[ s_n + \int_{n+1}^{\infty} f(x)\,dx  \;   \leq  \;   s  \;  \leq  \; s_n 
  + \int_{n}^{\infty}f(x)\, dx \]
	to find a better estimate of the sum: $BR
	\{$mp->ans_rule(35)\} \( \leq s \leq \)  \{$mp->ans_rule(35)\}
$BR
$BR
${BBOLD}Note${EBOLD} give each answer accurate to (at least) five decimal places.
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();

