##DESCRIPTION
#
# File Created: 3/27/2009
# Last Modified: 3/27/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Euler equations)
## Date(3/27/2009)
## Institution(Alfred University)
## Author(Juan Marin)
## MLT(IVPinterval)
## MLTleader(1)
## Level(3)
## TitleText1('Differential Equations with Boundary-Value Problems')
## AuthorText1('Zill and Cullen')
## EditionText1('7')
## Section1('4.7')
## Problem1('')


DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);
 
## Show partial correct answers
$showPartialCorrectAnswers = 1;
## Display the problem information
TEXT(beginproblem());




  Context("Interval");
  # to allow open or closed intervals, uncomment
  #    the following line.
  # Context()->flags->set(ignoreEndpointTypes=>1);

$int1 = Compute("(-inf,0)");
$int2 = Compute("(0,inf)");

## Declare variables
Context()->variables->are(x=>'Real',t=>'Real');

$a = non_zero_random(1,10);
$b = non_zero_random(11,20);
$u = non_zero_random(11,20);
$w = non_zero_random(1,10);

$A = $u * $w;
$B = -($b * $u + $a * $w) + $A ;
$C = $a * $b;

$p = $a/$u;
$q = $b/$w;

$y1 =random(-10,10);
$Dy1 =random(-10,10);


$Yh1 = Formula("x^($p)")->reduce;
$Yh2 = Formula("x^($q)")->reduce;



$multipart = MultiAnswer($Yh1, $Yh2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


$C1 = ($Dy1 - $q *$y1)/($p - $q);
$C2 = $y1 - $C1;


$yp = Formula("$C1 x^($p)+$C2 x^($q)")->reduce;




Context()->texStrings;
BEGIN_TEXT
In this problem you will solve the initial value problem 
\[$A x^2 y'' + $B x y' + $C y =0,\ \ y(1) = $y1,\ \ y'(1) = $Dy1.\]
$PAR
$BBOLD (1) $EBOLD
The largest intervals that may be considered for the domain of the solution to 
\($A x^2 y''+ $B x y' + $C y = 0\)
are either  \{ ans_rule(15) \} or \{ ans_rule(15) \}.$BR
$BBOLD
For \( -\infty \) type -inf and for \( \infty \) type inf.
$EBOLD
$PAR
$BBOLD (2) $EBOLD
 Let \(C_1\) and \(C_2\) be arbitrary constants.  The general solution to 
the homogeneous differential equation  
\($A x^2 y''+ $B x y' + $C y = 0\) is the function
\(y(x)=C_1\ y_1(x)+C_2\ y_2(x)=C_1\) \{$multipart->ans_rule(10)\} \(+C_2\)\{$multipart->ans_rule(10)\}.
$PAR
$BBOLD (3) $EBOLD
The unique solution to 
the initial value problem   
\[$A x^2 y''+$B x y' + $C y =0,\ \ y(1) = $y1,\ \ y'(1) = $Dy1.\] 
is the function
\(y(x)=\)\{ans_rule(30)\} 
for \( x \in\) \{ ans_rule(15) \}.

END_TEXT


Context()->normalStrings;
ANS( $int1->cmp() );
ANS( $int2->cmp() );
ANS( $multipart->cmp() );
ANS( $yp->cmp() );
ANS( $int2->cmp() );

ENDDOCUMENT();
