##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Euler equations)
## Date(6/3/2002)
## MLT(texas)
## Level(5)
## TitleText1('Precalculus')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.1')
## Problem1('22')
## KEYWORDS('algebra', 'inequality', 'fraction')

########################################################################
#Source Code for calc_compound_problem2.pg:

#This problem is a compound problem with three parts. The student is given an #integral that must be done using integration by parts. In the first part of #the problem the student must give the u and dv functions. In the second part #the student must give the du and v functions. In the third part the student #must #produce the antiderivative.

    DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "compoundProblem.pl",
  "parserMultiAnswer.pl",
  "MathObjects.pl",
  "unionTables.pl",
  "Alfredmacros.pl",
  "PGcourse.pl"
);

#install_problem_grader(~~&full_partial_grader);
$showPartialCorrectAnswers = 1;

Context()->variables->are(x=>'Real',a=>'Real',b=>'Real',c=>'Real',m=>'Real',u=>'Real',w=>'Real');
Context()->variables->set(x=>{limits=>[1,9]},a=>{limits=>[1,2]},b=>{limits=>[1,2]},c=>{limits=>[10,20]});


#Let's keep all variables to the general Problem here.
$va = Compute("a");
$vb = Compute("b");
$vc = Compute("c");
$ab = Compute("a-b");
$ba = Compute("b-a");
$aa = Compute("2*a");
$auxeq = Compute("a * m^2 + (b-a) * m + c");
$D = Compute("(b-a)^2 - 4 * a * c");
$NSD = Compute("sqrt(4*a*c-(b-a)^2)/(2*a)");


$m2 = Compute("(a-b)/(2a)");
$y41 = Compute("x^((a-b)/(2a))");
$P = Compute("b/(ax)");
$g = Compute("x^(-b/a)");
$xy41 = Compute("x^((a-b)/(2a))*ln(x)");

$multipart4 = MultiAnswer($y41, $xy41)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct4, $student4, $self4 ) = @_;
        my ( $f1stu4, $f2stu4 ) = @{$student4};
        my ( $f14, $f24 ) = @{$correct4};
        if ( ($f14 == $f1stu4 && $f24 == $f2stu4) ||
             ($f14 == $f2stu4 && $f24 == $f1stu4) ) {
            return [1,1];
        } else {
            if ($f14 == $f1stu4 || $f24 == $f1stu4) {
                return [1,0];
            } elsif ($f14 == $f1stu4 || $f24 == $f2stu4) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

$g51 = Compute("cos(w*ln(x))");
$g52 = Compute("sin(w*ln(x))");
$g53 = Compute("x^(u)*cos(w*ln(x))");
$g54 = Compute("x^(u)*sin(w*ln(x))");

$g61 = Compute("cos($NSD * ln(x))");#this is used on page 6
$g62 = Compute("sin($NSD * ln(x))");#this is used on page 6
$y61 = Compute("x^((a-b)/(2*a))*cos($NSD * ln(x))");#this is used on page 6
$y62 = Compute("x^((a-b)/(2*a))*sin($NSD * ln(x))");#this is used on page 6

$y7g  = Compute("x^(m)");#this is used on page 7
$y711 = Compute("x^(u)");#this is used on page 7
$y712 = Compute("x^(w)");#this is used on page 7
$y722 = Compute("x^(u)*ln(x)");#this is used on page 7
$y731 = Compute("x^(u)*cos(w*ln(x))");#this is used on page 7
$y732 = Compute("x^(u)*sin(w*ln(x))");#this is used on page 7

$r = random(2,9,1);
$s = random(2,9,1);
$p = random(2,9,1);
$q = random(2,9,1);

# Let's keep everything related to Example 1 here.
$a1 = $r * $s;
$ba1 = $r * $q - $p *$s;
$b1= $ba1 + $a1;
$c1 = -1 * $p * $q;

$auxeq1 = Compute("$a1 * m^2 + $ba1 * m + $c1");

$m11 = Compute("(-$ba1-sqrt(($ba1)^2-4*$a1*$c1))/(2*$a1)");
$m12 = Compute("(-$ba1+sqrt(($ba1)^2-4*$a1*$c1))/(2*$a1)");

$multipartm = MultiAnswer($m11, $m12)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $f1stu, $f2stu ) = @{$student};
        my ( $f1, $f2 ) = @{$correct};
        if ( ($f1 == $f1stu && $f2 == $f2stu) ||
             ($f1 == $f2stu && $f2 == $f1stu) ) {
            return [1,1];
        } else {
            if ($f1 == $f1stu || $f2 == $f1stu) {
                return [1,0];
            } elsif ($f1 == $f1stu || $f2 == $f2stu) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


$y11 = Formula("x^($m11)")->reduce;
$y12 = Formula("x^$m12")->reduce;

$multipart2 = MultiAnswer($y11, $y12)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct2, $student2, $self2 ) = @_;
        my ( $f1stu2, $f2stu2 ) = @{$student2};
        my ( $f12, $f22 ) = @{$correct2};
        if ( ($f12 == $f1stu2 && $f22 == $f2stu2) ||
             ($f12 == $f2stu2 && $f22 == $f1stu2) ) {
            return [1,1];
        } else {
            if ($f12 == $f1stu2 || $f22 == $f1stu2) {
                return [1,0];
            } elsif ($f12 == $f1stu2 || $f22 == $f2stu2) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


$multipart2b = MultiAnswer($y11, $y12)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct2b, $student2b, $self2b ) = @_;
        my ( $f1stu2b, $f2stu2b ) = @{$student2b};
        my ( $f12b, $f22b ) = @{$correct2b};
        if ( ($f12b == $f1stu2b && $f22b == $f2stu2b) ||
             ($f12b == $f2stu2b && $f22b == $f1stu2b) ) {
            return [1,1];
        } else {
            if ($f12b == $f1stu2b || $f22b == $f1stu2b) {
                return [1,0];
            } elsif ($f12b == $f1stu2b || $f22b == $f2stu2b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

# Let's keep all variables related to Example 2 here.
$a2 = $p * $p;
$ba2 = 2 * $p * $s;
$b2= $ba2 + $a2;
$c2 = $s * $s;
$auxeq2 = Compute("$a2 * m^2 + $ba2 * m + $c2");
$m21 = -1*$ba2/(2*$a2);
$y21 = Formula("x^($m21)")->reduce;
$y22 = Formula("x^($m21)ln(x)")->reduce;

$multiparte2a = MultiAnswer($y21, $y22)->with(
    singleResult => 0,
    checker => sub {
        my ( $correcte2a, $studente2a, $selfe2a ) = @_;
        my ( $f1stue2a, $f2stue2a ) = @{$studente2a};
        my ( $f1e2a, $f2e2a ) = @{$correcte2a};
        if ( ($f1e2a == $f1stue2a && $f2e2a == $f2stue2a) ||
             ($f1e2a == $f2stue2a && $f2e2a == $f1stue2a) ) {
            return [1,1];
        } else {
            if ($f1e2a == $f1stue2a || $f2e2a == $f1stue2a) {
                return [1,0];
            } elsif ($f1e2a == $f1stue2a || $f2e2a == $f2stue2a) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

$multiparte2b = MultiAnswer($y21, $y22)->with(
    singleResult => 0,
    checker => sub {
        my ( $correcte2b, $studente2b, $selfe2b ) = @_;
        my ( $f1stue2b, $f2stue2b ) = @{$studente2b};
        my ( $f1e2b, $f2e2b ) = @{$correcte2b};
        if ( ($f1e2b == $f1stue2b && $f2e2b == $f2stue2b) ||
             ($f1e2b == $f2stue2b && $f2e2b == $f1stue2b) ) {
            return [1,1];
        } else {
            if ($f1e2b == $f1stue2b || $f2e2b == $f1stue2b) {
                return [1,0];
            } elsif ($f1e2b == $f1stue2b || $f2e2b == $f2stue2b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

# Let's keep all variables related to Example 3 here.
$a3 = $q * $q;
$ba3 = 2 * $q * $r;
$b3= $ba3 + $a3;
$c3 = $r * $r + $s * $s;

$auxeq3 = Compute("$a3 * m^2 + $ba3 * m + $c3");

$m3r = -$ba3/(2*$a3);
$m3i = sqrt(4 * $a3 * $c3 - $ba3 * $ba3)/(2*$a3);

$y31 = Formula("x^($m3r)cos($m3i*ln(x))")->reduce;
$y32 = Formula("x^($m3r)sin($m3i*ln(x))")->reduce;


$multiparte3a = MultiAnswer($y31, $y32)->with(
    singleResult => 0,
    checker => sub {
        my ( $correcte3a, $studente3a, $selfe3a ) = @_;
        my ( $f1stue3a, $f2stue3a ) = @{$studente3a};
        my ( $f1e3a, $f2e3a ) = @{$correcte3a};
        if ( ($f1e3a == $f1stue3a && $f2e3a == $f2stue3a) ||
             ($f1e3a == $f2stue3a && $f2e3a == $f1stue3a) ) {
            return [1,1];
        } else {
            if ($f1e3a == $f1stue3a || $f2e3a == $f1stue3a) {
                return [1,0];
            } elsif ($f1e3a == $f1stue3a || $f2e3a == $f2stue3a) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

$multiparte3b = MultiAnswer($y31, $y32)->with(
    singleResult => 0,
    checker => sub {
        my ( $correcte3b, $studente3b, $selfe3b ) = @_;
        my ( $f1stue3b, $f2stue3b ) = @{$studente3b};
        my ( $f1e3b, $f2e3b ) = @{$correcte3b};
        if ( ($f1e3b == $f1stue3b && $f2e3b == $f2stue3b) ||
             ($f1e3b == $f2stue3b && $f2e3b == $f1stue3b) ) {
            return [1,1];
        } else {
            if ($f1e3b == $f1stue3b || $f2e3b == $f1stue3b) {
                return [1,0];
            } elsif ($f1e3b == $f1stue3b || $f2e3b == $f2stue3b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );
###################################################################
###################################################################

$multipart3 = MultiAnswer($y41, $xy41)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct3, $student3, $self3 ) = @_;
        my ( $f1stu3, $f2stu3 ) = @{$student3};
        my ( $f13, $f23 ) = @{$correct3};
        if ( ($f13 == $f1stu3 && $f23 == $f2stu3) ||
             ($f13 == $f2stu3 && $f23 == $f1stu3) ) {
            return [1,1];
        } else {
            if ($f13 == $f1stu3 || $f23 == $f1stu3) {
                return [1,0];
            } elsif ($f13 == $f1stu3 || $f23 == $f2stu3) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


$multipart6 = MultiAnswer($y61, $y62)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct6, $student6, $self6 ) = @_;
        my ( $f1stu6, $f2stu6 ) = @{$student6};
        my ( $f16, $f26 ) = @{$correct6};
        if ( ($f16 == $f1stu6 && $f26 == $f2stu6) ||
             ($f16 == $f2stu6 && $f26 == $f1stu6) ) {
            return [1,1];
        } else {
            if ($f16 == $f1stu6 || $f26 == $f1stu6) {
                return [1,0];
            } elsif ($f16 == $f1stu6 || $f26 == $f2stu6) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


$multipart7a = MultiAnswer($y711, $y712)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct7a, $student7a, $self7a ) = @_;
        my ( $f1stu7a, $f2stu7a ) = @{$student7a};
        my ( $f17a, $f27a ) = @{$correct7a};
        if ( ($f17a == $f1stu7a && $f27a == $f2stu7a) ||
             ($f17a == $f2stu7a && $f27a == $f1stu7a) ) {
            return [1,1];
        } else {
            if ($f17a == $f1stu7a || $f27a == $f1stu7a) {
                return [1,0];
            } elsif ($f17a == $f1stu7a || $f27a == $f2stu7a) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

$multipart7b = MultiAnswer($y711, $y722)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct7b, $student7b, $self7b ) = @_;
        my ( $f1stu7b, $f2stu7b ) = @{$student7b};
        my ( $f17b, $f27b ) = @{$correct7b};
        if ( ($f17b == $f1stu7b && $f27b == $f2stu7b) ||
             ($f17b == $f2stu7b && $f27b == $f1stu7b) ) {
            return [1,1];
        } else {
            if ($f17b == $f1stu7b || $f27b == $f1stu7b) {
                return [1,0];
            } elsif ($f17b == $f1stu7b || $f27b == $f2stu7b) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );

$multipart7c = MultiAnswer($y731, $y732)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct7c, $student7c, $self7c ) = @_;
        my ( $f1stu7c, $f2stu7c ) = @{$student7c};
        my ( $f17c, $f27c ) = @{$correct7c};
        if ( ($f17c == $f1stu7c && $f27c == $f2stu7c) ||
             ($f17c == $f2stu7c && $f27c == $f1stu7c) ) {
            return [1,1];
        } else {
            if ($f17c == $f1stu7c || $f27c == $f1stu7c) {
                return [1,0];
            } elsif ($f17c == $f1stu7c || $f27c == $f2stu7c) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


  Context("Interval");
  # to allow open or closed intervals, uncomment
  #    the following line.
  # Context()->flags->set(ignoreEndpointTypes=>1);

  $int = Compute("(0,infinity)");

TEXT(beginproblem);

    ##############################################

    #This gives the professor (in this case jmarin) the ability to "go backwards" in 
    #a compound problem - useful for debugging.
    $isProfessor = ($studentLogin eq 'admin' || $studentLogin eq 'professor');

    #
    #  Start a compound problem.  See the compoundProblem.pl
    #  file for more details about the parameters you
    #  can supply.
    #
    $cp = new compoundProblem(
      parts => 7,                  # the total number of parts in this problem
      totalAnswers => 61,           # total answers in all parts combined
      parserValues => 1,           # make MathObjects from student answers
      allowReset => $isProfessor,  # professors get Reset button for testing
      nextVisible => Always,
    );
    $part = $cp->part;             # look up the current part


    ##############################################
    #
    #  Part 1
    #

    if ($part == 1) {

    BEGIN_TEXT
$PAR
This problem has multiple parts.  After you submit your answers, you will be given the option to move on to the next part.  As you will submit this problem  to me, before you move on to the next part, please print your work.

$PAR 
Throughout this WeBWork problem, let \(a,b,\) and \(c\) be real numbers with \(a \neq 0\).  In this exercise we will discover a method that will allow us to solve linear homogeneous differential equations of the form
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0.\]
An equation of this type, is called a Cauchy-Euler equation.
$PAR 

To solve this equation, we will assume that the solution \(y\) of the differential equation is of the form \(y=x^{m}\) where \(m\) is some number that we need to find. Since \(x^2 = 0\) when \(x=0\) a solution to a Cauchy -Euler equation will be valid for \(x \in (-\infty,0)\) or for \(x\in(0,\infty)\).  Throughout this assignment, we will assume that \(x\in(0,\infty)\).


 If we substitute \(y=x^{m}\) into 
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\), we find that 
$BR
$BCENTER
\{ BeginTable(center=>0).
      Row(['\(0=ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy\)', ],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 48pt =x^{m}\Big[\)',ans_rule(1),'\(m^2+\)',ans_rule(2),'\(m + \)',ans_rule(1),'\(\Big]\)' ],separation=>1).
   EndTable();
\}
$ECENTER
$BR

    END_TEXT
     ANS($va->cmp());
     ANS($ba->cmp());
     ANS($vc->cmp());

    #
    #  the useGrader() method sets the grader for this part. The student is only 
    #given credit for each part of the problem if all answer blanks in that part are 
    #correct (so no partial credit within parts of the problem).
    #
    $cp->useGrader(~~&std_problem_grader);
    $showPartialCorrectAnswers = 0;
    #Students are not shown which answer is incorrect.

    }

    ##############################################
    #
    #  Part 2
    #

    if ($part == 2) {

    BEGIN_TEXT


Good! If we assume that \(y=x^{m}\) is a solution to the differential equation  
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\), we find that 
\[
\begin{aligned}
0&=ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy\\ 
 &\\
 &=x^{m}\left[am^2+(b-a)m+c\right].
\end{aligned}
\]

Since \(x\in(0,\infty)\),
\(x^{m}\) never equals zero; thus \(am^2+(b-a)m+c=0\).  The equation \[am^2+(b-a)m+c=0\] is called the $BBOLD auxiliary equation $EBOLD of the differential equation \[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0.\]
$PAR

$BBOLD Let's look at some examples. $EBOLD
$PAR

$BBOLD Example 1: $EBOLD The auxiliary equation of the differential equation
\[ $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\] is 
\{ ans_rule(15) \}=\{ ans_rule(1) \}.
$PAR

$BBOLD Example 2: $EBOLD The auxiliary equation of the differential equation
\[ $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\] is  
\{ ans_rule(15) \}=\{ ans_rule(1) \}.
$PAR

$BBOLD Example 3: $EBOLD The auxiliary equation of the differential equation
\[ $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\] is  
\{ ans_rule(15) \}=\{ ans_rule(1) \}.






    END_TEXT

ANS($auxeq1->cmp());
ANS(Compute("0")->cmp());
ANS($auxeq2->cmp());
ANS(Compute("0")->cmp());
ANS($auxeq3->cmp());
ANS(Compute("0")->cmp());


    }

    ##############################################
    #
    #  Part 3
    #

    if ($part == 3) {

    BEGIN_TEXT

So, if we assume that \(y=x^{m}\) is a solution to the differential equation  
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\), we arrive at the auxiliary equation \(am^2+(b-a)m+c=0.\)

$PAR

From the quadratic formula, we know that \(am^2+(b-a)m+c=0\) when 
$BR
$BCENTER
\{ BeginTable(center=>0).
      Row(['\(m=\)',ans_rule(3).'\(\pm\sqrt{}\)'.ans_rule(10).$BR."-----------------------------------------".$BR.'\(\hskip 25pt\)'.ans_rule(3)],separation=>1).
   EndTable();
\}
$ECENTER
$BR
$PAR

We will need to consider three different cases; that is, when  
\((b-a)^2-4ac > 0, (b-a)^2-4ac = 0\), and \((b-a)^2-4ac<0.\)
$PAR
$BBOLD Case 1:$EBOLD If \( (b-a)^2-4ac>0\), then if \(x \in (0,\infty)\), we have found two linearly independent solutions 
\(y_1(x)= x^{\displaystyle{\frac{a-b-\sqrt{(b-a)^2-4ac}}{2a}}}\) 
and
\(y_2(x)= x^{\displaystyle{\frac{a-b+\sqrt{(b-a)^2-4ac}}{2a}}}\) 
to the differential equation \( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0.\)  In this case, the general solution to \[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0,\ \ \ x\in(0,\infty)\] is the function \[y= c_1 y_1(x) + c_2 y_2(x)\] where \(c_1\) and \(c_2\) are arbitrary constants.
$PAR
$BBOLD Example 4 :$EBOLD
In Example 1, we found that the auxiliary equation of the differential equation 
\[ $a1 x^2\frac{d^2y}{dx^2}+ $b1 x\frac{dy}{dx}+ $c1 y=0\] is 
\( $a1 m^2+ $ba1 m+ $c1 =0.\) (Notice that \( (b-a)^2-4ac > 0\).)
$BR
Since \( $a1 m^2+ $ba1 m+ $c1 =0\), we have that
\(m=\)\{ ans_rule(5) \} 
and 
\(m=\)\{ ans_rule(5) \}.
$BR
In this case, 
\(y_1(x)=\)\{ ans_rule(15) \}
and 

\(y_2(x)=\)\{ ans_rule(15) \}
$BR
are linearly independent solutions to 
\( $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\) and 
$BR
\(y = c_1\) \{ ans_rule(15) \}\( + c_2 \)\{ ans_rule(15) \} where \(c_1\) and \(c_2\) are arbitrary constants 
is the general solution to \( $a1 \frac{d^2y}{dx^2}+ $b1 \frac{dy}{dx}+ $c1 y=0\).

    END_TEXT

ANS($ab->cmp());
ANS($D->cmp());
ANS($aa->cmp());
ANS($multipartm->cmp());
ANS($multipart2->cmp());
ANS($multipart2b->cmp());


    }

   ##############################################
    #
    #  Part 4
    #

    if ($part == 4) {

    BEGIN_TEXT
So, if we assume that \(y=x^{m}\) is a solution to the differential equation  
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\), we arrive at the auxiliary equation \(am^2+(b-a)m+c=0.\)

$PAR

From the quadratic formula, we know that \(am^2+(b-a)m+c=0\) when 
\[m=\frac{a-b\pm\sqrt{(b-a)^2-4ac}}{2a}\]
$PAR

We now consider the case when  
\((b-a)^2-4ac = 0.\)
$PAR
$BBOLD Case 2:$EBOLD If \( (b-a)^2-4ac=0\), then the only solution to the auxiliary equation \(am^2+(b-a)m+c=0\) is 
\(m=\)\{ ans_rule(15) \} and in this case, we get only one solution 
\(y_1(x)=\) \{ ans_rule(15) \}
to the differential equation \( a\frac{d^2y}{dx^2}+b\frac{dy}{dx}+cy=0.\)
$PAR

Since a second order linear homogeneous differential equation must have two linearly independent solutions, we need a second solution that is linearly independent to \(y_1\).
$PAR

Recall that if \( y_1(x)\) is a solution to the differential equation 
\[ a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\]
then if \( P(x)=a_1(x)/a_2(x)\), a second solution to  
\( a_2(x)\frac{d^2y}{dx^2}+a_1(x)\frac{dy}{dx}+a_0(x)y=0,\) is given by 
\[
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx.
\]
Furthermore, \(y_1\) and \(y_2\) are linearly independent.

$PAR

For the differential equation 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0,\]
\(P=\)\{ ans_rule(15) \}, 
\(e^{-\int P(x)dx}=\)\{ ans_rule(15) \}, and 
$BR
$BCENTER
\(
y_2(x)=y_1(x)\int\frac{e^{-\int P(x)dx}}{y_1^2(x)}dx
=\)\{ ans_rule(15) \}.
$ECENTER
$BR


$PAR
So, if \( (b-a)^2-4ac=0\), then we have found two linearly independent solutions
$BR
\(y_1(x)=\)\{ ans_rule(15) \}
and
\(y_2(x)=\)\{ ans_rule(15) \}
to the differential equation \( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0.\)  In this case, the general solution to the differential equation 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
is 
$BR
$BCENTER
\(y(x) = c_1\)\{ ans_rule(15) \}\(+c_2\)\{ ans_rule(15) \}
$ECENTER
$BR
$PAR
$BBOLD Example 5 :$EBOLD
In Example 2, we found that the auxiliary equation of the differential equation 
\[ $a2 x^2\frac{d^2y}{dx^2}+ $b2 x\frac{dy}{dx}+ $c2 y=0\] is 
\( $a2 m^2+ $ba2 m+ $c2 =0.\) (Notice that \( (b-a)^2-4ac=0\).)
$BR
Since \( $a2 m^2+ $ba2 m+ $c2 =0\), we have that
\(m=\)\{ ans_rule(5) \}.
$BR
In this case, 
\(y_1(x)=\)\{ ans_rule(15) \}
and 
\(y_2(x)=\)\{ ans_rule(15) \}
$BR
are linearly independent solutions to 
\( $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\) and
$BR 
$BCENTER
\(y = c_1 \)\{ ans_rule(15) \}\(+ c_2 \)\{ ans_rule(15) \}
$ECENTER
$BR
where \(c_1\) and \(c_2\) are arbitrary constants 
is the general solution to \( $a2 \frac{d^2y}{dx^2}+ $b2 \frac{dy}{dx}+ $c2 y=0\).





    END_TEXT
ANS($m2->cmp());
ANS($y41->cmp());
ANS($P->cmp());
ANS($g->cmp());
ANS($xy41->cmp());
ANS($y41->cmp());
ANS($xy41->cmp());
ANS($multipart3->cmp());
ANS(Compute("$m21")->cmp());
ANS($multiparte2a->cmp());
ANS($multiparte2b->cmp());

    }


   ##############################################
    #
    #  Part 5
    #

    if ($part == 5) {

    BEGIN_TEXT
Before we tackle the last case, I want to remind you of a very useful identity from complex analysis.  Don't forget that in complex analysis 
\(i = \sqrt{-1}\).
$PAR

The formula I want to remind you of is Euler's formula, which states that 
\[e^{iz}=\cos(z)+i\sin(z).\]
$PAR

So, if \(u\) and \(w\) are real numbers, 
$BR
$BCENTER
\{ BeginTable(center=>0).
      Row(['\(x^{u+iw}\)','\(=\)', '\(x^{u}x^{iw}\)'],separation=>1).
      Row(['','\(=\)', '\(x^{u}x^{iw}\)'],separation=>1).
      Row(['','\(=\)', '\(x^{u}e^{iw\ln(x)}\)'],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 164pt\)','\(= x^{u}\Big(\)',ans_rule(15),'\(+ i\)',ans_rule(15),'\(\Big)\)'],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 141pt\)','\(= \)',ans_rule(15),'\(+ i\)',ans_rule(15)],separation=>1).
   EndTable();
\}
$ECENTER
$BR
$BBOLD
As this last equation will be used on the next page, please make sure that you write it down.
$EBOLD





    END_TEXT
ANS($g51->cmp());
ANS($g52->cmp());
ANS($g53->cmp());
ANS($g54->cmp());
    }

   ##############################################
    #
    #  Part 6
    #

    if ($part == 6) {

    BEGIN_TEXT
So, if we assume that \(y=x^{m}\) is a solution to the differential equation  
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\), we arrive at the auxiliary equation \(am^2+(b-a)m+c=0.\)

$PAR

From the quadratic formula, we know that \(am^2+(b-a)m+c=0\) when 
\[m=\frac{a-b\pm\sqrt{(b-a)^2-4ac}}{2a}\]
$PAR

We now consider the case when  
\((b-a)^2-4ac < 0.\)
$PAR
$BBOLD Case 3:$EBOLD If \( (b-a)^2-4ac < 0\), then the solutions to the auxiliary equation \(am^2+(b-a)m+c=0\) are 
\(m=\displaystyle{\frac{a-b-\sqrt{(b-a)^2-4ac}}{2a}}\) and
\(m=\displaystyle{\frac{a-b+\sqrt{(b-a)^2-4ac}}{2a}}\).  For now let's focus on what happens if \(m=\displaystyle{\frac{a-b+\sqrt{(b-a)^2-4ac}}{2a}}\).

Notice that
$BR
$BCENTER
\{ BeginTable(center=>0).
      Row(['\(m\)','\(=\frac{a-b+\sqrt{(b-a)^2-4ac}}{2a}\)'],separation=>1).
      Row(['','\(= \frac{a-b}{2a}+\frac{\sqrt{-1(4ac-(b-a)^2)}}{2a}\)'],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(\hskip 50pt\)','\(=\)',ans_rule(10),'\(+i\)',ans_rule(10)],separation=>1).
   EndTable();
\}
$ECENTER
$BR
So,
$BR
$BCENTER
\{ BeginTable(center=>0).
      Row(['\(\hskip 220pt y\)','\(= x^{m}\)'],separation=>1).
      Row(['','\(= x^{\frac{a-b}{2a}}\Big(\)',ans_rule(25),'\(\Big)+i\ x^{\frac{a-b}{2a}}\Big(\)',ans_rule(25),'\(\Big)\)'],separation=>1).
   EndTable();
\}
\{ BeginTable(center=>0).
      Row(['\(= y_1(x)+iy_2(x).\)'],separation=>1).
   EndTable();
\}
$ECENTER
$BR
$PAR
If \( (b-a)^2-4ac < 0\), it turns out that  
$BR
\(y_1(x)=\)\{ ans_rule(35) \}
$BR and $BR
\(y_2(x)=\)\{ ans_rule(35) \}
$BR
are linearly independent solutions 
to the differential equation \( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0.\) 
(You should check this.) 
Since a fundamental set for a second order linear differential equation has two linearly independent solutions, \(\lbrace y_1, y_2  \rbrace\) forms a fundamental set for 
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\) when \( (b-a)^2-4ac < 0\).

If \( (b-a)^2-4ac < 0\), the general solution to the differential equation 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
is 
$BR
$BCENTER
\(y(x) = c_1\) \{ ans_rule(35) \}\(+c_2\)\{ ans_rule(35) \}
$ECENTER
$PAR
$BBOLD Example 5 :$EBOLD
In Example 3, we found that the auxiliary equation of the differential equation 
\[ $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\] is 
\( $a3 x^2m^2+ $ba3 x m+ $c3 =0.\) (Notice that \( (b-a)^2-4ac < 0\).)
$BR
Since \( $a3 m^2+ $ba3 m+ $c3 =0\), we have that
\(m=\)\{ ans_rule(5) \}\(\pm i\) \{ ans_rule(5) \}.
$BR
In this case, 
\(y_1(x)=\)\{ ans_rule(17) \}
and 
\(y_2(x)=\)\{ ans_rule(17) \}
$BR
are linearly independent solutions to 
\( $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\) and 
$BR
$BCENTER
\(y = c_1 \) \{ ans_rule(17) \} \( + c_2 \) \{ ans_rule(17) \}
$ECENTER
$BR
 where \(c_1\) and \(c_2\) are arbitrary constants 
is the general solution to \( $a3 \frac{d^2y}{dx^2}+ $b3 \frac{dy}{dx}+ $c3 y=0\).
$PAR

$BBOLD
The material in this webwork assignment is covered in Section 4.3 of our text.
$EBOLD






    END_TEXT
ANS($m2->cmp());
ANS(fun_cmp( $NSD,var=>['a','b','c'],limits=>[[1,2],[1, 2],[10,20]]));
ANS(fun_cmp( $g61,var=>['a','b','c','x'],limits=>[[1,2],[1, 2],[10,20],[1,9]]));
ANS(fun_cmp( $g62,var=>['a','b','c','x'],limits=>[[1,2],[1, 2],[10,20],[1,9]]));
ANS(fun_cmp( $y61,var=>['a','b','c','x'],limits=>[[1,2],[1, 2],[10,20],[1,9]]));
ANS(fun_cmp( $y62,var=>['a','b','c','x'],limits=>[[1,2],[1, 2],[10,20],[1,9]]));
ANS($multipart6->cmp());
ANS(Compute("$m3r")->cmp());
ANS(Compute("$m3i")->cmp());
ANS($multiparte3a->cmp());
ANS($multiparte3b->cmp());
    }



   ##############################################
    #
    #  Part 7
    #

    if ($part == 7) {

    BEGIN_TEXT

Let's review what we have learned.  To solve a Cauchy-Euler equation
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
we assume that the solution is of the form \(y=\)\{ ans_rule(4) \}.

Substituting \(y\) into the equation 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
we arrive at the auxiliary equation
\{ ans_rule(12) \} \(=0\)
$PAR


$BBOLD
Case 1.
$EBOLD
If the roots of the auxiliary equation to 
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\)
are two real numbers \(u\) and \(w\), then 
the general solution to 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
is \(y=c_1\)\{ ans_rule(10) \} \(+c_2\) \{ ans_rule(10) \}
for \(x\) in the interval  \{ ans_rule(10) \}.$BR
$BBOLD 
For \(-\infty\) type -infinity and for \(\infty\) type infinity
$EBOLD
$PAR 


$BBOLD
Case 2.
$EBOLD
If the auxiliary equation to  
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\)
has only one real root \(u\), then 
the general solution to 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
is \(y=c_1\)\{ ans_rule(10) \} \(+c_2\) \{ ans_rule(10) \}
for \(x\) in the interval  \{ ans_rule(10) \}.$BR
$BBOLD 
For \(-\infty\) type -infinity and for \(\infty\) type infinity
$EBOLD
$PAR 

$BBOLD
Case 3.
$EBOLD
If the roots of the auxiliary equation to 
\( ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\)
are two complex numbers \(u + iw\) and \(u - iw\), then 
the general solution to 
\[ ax^2\frac{d^2y}{dx^2}+bx\frac{dy}{dx}+cy=0\]
is \(y=c_1\)\{ ans_rule(10) \} \(+c_2\) \{ ans_rule(10) \}
for \(x\) in the interval  \{ ans_rule(10) \}.$BR
$BBOLD 
For \(-\infty\) type -infinity and for \(\infty\) type infinity
$EBOLD



    END_TEXT

  ANS($y7g->cmp());
  ANS($auxeq->cmp());
  ANS($multipart7a->cmp());
  ANS( $int->cmp() );
  ANS($multipart7b->cmp());
  ANS( $int->cmp() );
  ANS($multipart7c->cmp());
  ANS( $int->cmp() );

    }


    ##############################################

    ENDDOCUMENT();        # This should be the last executable line in the problem.


