## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear combinations)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(2)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Fraction');
Context()->parens->set("[" => {formMatrix => 1});

do {

$Mred = Matrix([
[1,0,0,random(-5,5,1)],
[0,1,0,random(-5,5,1)],
[0,0,1,random(-5,5,1)],
]);

$U = GL3Z();

$M = $U * $Mred;

$b = Matrix([
[$M->element(1,4)],
[$M->element(2,4)],
[$M->element(3,4)],
]);

$A1 = $M->column(1);
$A2 = $M->column(2);
$A3 = $M->column(3);

$x1 = $Mred->element(1,4);
$x2 = $Mred->element(2,4);
$x3 = $Mred->element(3,4);

} until (($x1 != 0) and ($x2 != 0) and ($x3 != 0));

$multians = MultiAnswer($x1,$x2,$x3)->with(
  singleResult => 1,
  checker => sub {
    my ($correct,$student,$ansHash) = @_;  # get the parameters
    my @c = @{$correct};      # extract the correct answers
    my @s = @{$student};      # extract the student answers
    return (($s[0] == $c[0]) and ($s[1] == $c[1]) and ($s[2] == $c[2]));
  },
);

Context()->texStrings;
BEGIN_TEXT
Given that
\[ $M \sim $Mred, \]
write \( $b \) as a linear combination of the vectors \( $A1, $A2, $A3. \)
$BR
$BR
\( $b = \) 
\{ $multians->ans_rule(10) \} \( $A1 + \)
\{ $multians->ans_rule(10) \} \( $A2 + \)
\{ $multians->ans_rule(10) \} \( $A3.\)
END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS($multians->cmp);

COMMENT('Correct answer always has nonzero integer entries.');
ENDDOCUMENT();
