# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Polynomial and Rational Functions)
## DBsection(Polynomial Functions and their Graphs)
## Institution(PCC)
## Author(Alex Jordan)
## MO(1)



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PCCgraphMacros.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

##############################################


Context("Numeric");

$r[0] = random(-4,-1,1);
for my $i (1..2)
{
  $r[$i] = $r[$i-1] + random(1,5,1);
}

$p = Compute("(x-$r[0])*(x-$r[2])/(x-$r[1])")->reduce;

for my $i ($r[0]..$r[2]-1)
{$vals[$i-$r[0]] = abs($p->eval(x=>$i+0.5));}

$ymax = 1.3*max(@vals);

$xmax = max($r[2]+1,5);
$gr = init_graph(-5,-$ymax,$xmax,$ymax,axes=>[0,0],grid=>[$xmax+5,10],size=>[400,400]);
$gr -> lb('reset');


foreach my $i (1..$r[2]) {
  $gr -> lb(new Label ( $i,-0.5, $i,'black','center','middle'));
  $gr -> lb(new Label (-$i,-0.5,-$i,'black','center','middle'));
}

$r1low = $r[1] - 0.0001;
$r1high = $r[1] + 0.0001;

add_functions($gr, "$p for x in <-5,$r1low> using color:blue and weight:1");
add_functions($gr, "$p for x in <$r1high,$xmax> using color:blue and weight:1");

##############################################

$refreshCachedImages = 1;
BEGIN_PGML

Find a formula for a rational function [`f`] that might have this graph. Do not over-complicate things: assume that all degrees are as small as can be compatible with this graph.

[@EnlargeImageStatementPGML()@]**


>>[@image(insertGraph($gr), width=>400, height=>400, tex_size=>800)@]* <<


    [`f(x)={}`][________________________________]

END_PGML

##############################################

$showPartialCorrectAnswers = 1;

ANS( $p->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $context = Context()->copy;
      return 0 if $student == 0;
      $context->flags->set(no_parameters=>0);
      $context->variables->add('C0'=>'Parameter');
      my $c0 = Formula($context,'C0');
      $student = Formula($context,$student);
      $correct = Formula($context,"abs($c0) * $p");
      return $correct == $student;
    }
) );

##############################################

BEGIN_PGML_SOLUTION

We need to see from the graph that there are zeros at [`[@$r[0]@]`] and [`[@$r[2]@]`]. Further, since the graph passes through them is a straight-line fashion, they are degree-1 zeros.

We also need to see that there is a vertical asymptote at [`[@$r[1]@]`]. Since the arms of the function go in opposite directions (one up, one down) as [`x`] approaches [`[@$r[1]@]`], then the degree of the vertical asymptote must be odd: [`1`], [`3`], [`5`], etc. 

Since this is a rational function, having zeros at [`[@$r[0]@]`] and [`[@$r[2]@]`] means that [`(x-[@$r[0]@])`] and [`(x-[@$r[2]@])`] are factors of the numerator. Since they are degree-1 zeros, there is no exponent on these factors (other than [`1`]).

Having a vertical asymptote at [`[@$r[1]@]`] means that [`(x-[@$r[1]@])`] is a factor of the denominator. At first thought, the exponent on this factor could be [`1`] or [`3`] or [`5`], etc.  But we can see that the function has a slant-line asymptote, and so the polynomial long division of the numerator by the denominator must result in a degree-1 polynomial. Since we know the numerator has total degree [`2`], the denominator must therefore have total degree [`1`]. 

So far, this tells us that [`f(x)=c[$p]`] for some constant [`c`]. If [`c`] were negative, then the long-term behavior (slant-line asymptote) would have a downward trend. But the graph shows it has an upward trend. So [`c`] must be positive, and you can enter any formula that has this pattern to get credit for this problem, including just [`[$p]`] (where [`c=1`]).

END_PGML_SOLUTION

##############################################





ENDDOCUMENT();