##DESCRIPTION

## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Trigonometric substitution)
## Date(8/13/13)
## Institution(Cleveland State University)
## Author(Barbara Margolius)
## MLT(trigSub-tangent)
## Level(4)
## MO(1)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('.')
## Problem1('2')
## KEYWORDS('integrals', 'trigonometric','substitution')
##ENDDESCRIPTION

############################################################################
## development of this problem is supported in part by the National Science#
## Foundation under the grant DUE-0941388.                                 #
############################################################################

## Bug fixer's comment:  This problem originally used 
## parserFormulaUpToConstants.pl to require 
## students to enter "+C" , but for some 
## seeds the range of possible correct answers was so large that some correct 
## answers raised "can't solve for ## adaptive parameters" errors.  I found 
## no way to fix this by resetting max_adapt or anything else so I edited 
## the problem, supplying "+C" in the text, removing 
## parserFormulaUpToConstants.pl, and adding a custom answer checker.  
## -- George Jennings (GAJ)

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "AppletObjects.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = random(2,9,1);

$a2 = $a*$a;
$a4 = $a2*$a2;

Context("Numeric");  ## GAJ edit

#$funct = FormulaUpToConstant("x/4*($a2+x^2)^(3/2)-$a4/8*ln(x+sqrt($a2+x^2))-$a2/8*x*sqrt(x^2+$a2)")->with(limits => [$a+1,$a+2]); ## GAJ edit
$funct = Compute("x/4*($a2+x^2)^(3/2)-$a4/8*ln(x+sqrt($a2+x^2))-$a2/8*x*sqrt(x^2+$a2)")->with(limits => [$a+1,$a+2]); ## GAJ edit

   ###################################
    # Create  link to applet 
    ###################################
    $appletName = "trigSubWW";
    $applet =  FlashApplet(
       codebase              => findAppletCodebase("$appletName.swf"),
       appletName            => $appletName,
       appletId              => $appletName,
       setStateAlias         => 'setXML',
       getStateAlias         => 'getXML',
       setConfigAlias        => 'setConfig',
       maxInitializationAttempts => 10,   # number of attempts to initialize applet
       #answerBoxAlias        => 'answerBox',
       height                => '550',
       width                 => '595',
       bgcolor               => '#e8e8e8',
       debugMode             =>  0,
     );

###################################
# Configure applet
###################################
 
    $applet->configuration(qq{<xml><trigString>tan</trigString></xml>});
    $applet->initialState(qq{<xml><trigString>tan</trigString></xml>});

##################################
# Setup Flash applet -- this does not need to be changed
###################################

TEXT(MODES(TeX=>"", HTML=><<'END_TEXT'));
<script>
if (navigator.appVersion.indexOf("MSIE") > 0) {
	document.write("<div width='3in' align='center' style='background:yellow'>You seem to be using Internet Explorer.<br/>It is recommended that another browser be used to view this page.</div>");
}
</script>
END_TEXT
###################################
# Main text

Context()-texStrings; ## this and "+C" are GAJ edits
BEGIN_TEXT

Evaluate the indefinite integral.
$BR \[ \int x^2\sqrt{$a2 + x^2}\; dx \]
$BR \{ans_rule( 50) \} + C  

END_TEXT

Context()->normalStrings;  ## GAJ edit

$ans = $funct;
#ANS( $funct->cmp() );  ## GAJ edit
ANS( $funct->cmp(
checker => sub {
       my ($correct,$student,$ansHash) = @_; 
       return $correct->D('x') == $student->D('x'); 
     }
));  ## GAJ edit
##################################
Context()->texStrings;

###################################
TEXT($PAR, $BBOLD, $BITALIC, "Hi $studentLogin, If you don't get this in 5 tries I'll give you a hint with an applet to help you out.", $EITALIC, $EBOLD, $PAR);

$im = image( "trigsub.png", 
width=>486, height=>306, tex_size=>900 );

$showHint=5;
Context()->normalStrings;
TEXT(hint(
 $PAR,  MODES(TeX=>'object code', HTML=>$applet->insertAll(
   debug =>0, reinitialize_button => 0, includeAnswerBox=>0,
 ), TeX=>$im),"Follow the step-by-step questions in the hint in the online version of this problem."
));

##################################
Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$BBOLD Solution: $EBOLD $PAR
To evaluate this integral use a trigonometric substitution.  For this problem use the \(\tan\) substitution. \[x = {$a}\tan(\theta)\]
So:
\[dx = {$a}\sec^2(\theta) \; d\theta\]
Therefore:
\[\int x^2\sqrt{$a2 + x^2} \;dx=\int $a2\tan^2\theta\sqrt{$a2+$a2\tan\theta}($a\sec^2\theta) \; d\theta\]
\[=\int  $a4\tan^2\theta\sec^3\theta \; d\theta\]
\[=$a4\int  (\sec^2\theta-1)\sec^3\theta \; d\theta\]
\[=$a4\int  (\sec^5\theta-\sec^3\theta) \; d\theta\]

$BR$BR
Before proceeding to the several methods we can use to evaluate the integral of \(\sec^3\theta\), note that \(\tan\theta=\frac{x}{$a}\), and \(\sec\theta=\frac{\sqrt{$a2+x^2}}{$a}\).  To see this, label a right triangle so that the tangent is \(x/$a\).  We will have the opposite side with length \(x\), and the adjacent side with length \($a\), so the hypotenuse has length \(\sqrt{$a2+x^2}\).

$BR$BR
From this point, we apply the secant reduction formula to evaluate the integral.

$BR
Use the secant reduction formula.  The secant reduction formula is given by
\[\int \sec^n\theta d\theta=\frac{1}{n-1}\tan\theta\sec^{n-2}\theta+\frac{n-2}{n-1}\int\sec^{n-2}\theta d\theta.\]

$BR$BR
so
\[\int\sec^5\theta d\theta=\frac{1}{5-1}\tan\theta\sec^{5-2}\theta+\frac{5-2}{5-1}\int\sec^{5-2}\theta d\theta\]
\[=\frac{1}{4}\tan\theta\sec^{3}\theta+\frac{3}{4}\int\sec^{3}\theta d\theta\]
and
\[\int\sec^3\theta d\theta=\frac{1}{3-1}\tan\theta\sec^{3-2}\theta+\frac{3-2}{3-1}\int\sec^{3-2}\theta d\theta\]
\[=\frac{1}{2}\tan\theta\sec\theta+\frac{1}{2}\int\sec\theta d\theta\]

$BR
Substituting back in,
\[$a4\int  (\sec^5\theta-\sec^3\theta) \; d\theta
=\frac{{$a4}}{4}\tan\theta\sec^{3}\theta-\frac{{$a4}}{4}\int\sec^{3}\theta d\theta \]
\[=\frac{{$a4}}{4}\tan\theta\sec^{3}\theta-\frac{{$a4}}{4}\left(\frac{1}{2}\tan\theta\sec\theta+\frac{1}{2}\int\sec\theta d\theta\right) \]

\[=\frac{{$a4}}{4}\tan\theta\sec^{3}\theta-\frac{{$a4}}{8}\tan\theta\sec\theta-\frac{{$a4}}{8}\int\sec\theta d\theta \]

\[=\frac{{$a4}}{4}\left(\frac{x}{$a}\right)\left(\frac{\sqrt{$a2+x^2}}{$a}\right)^3-\frac{{$a4}}{8}\left(\frac{x}{$a}\right)\left(\frac{\sqrt{$a2+x^2}}{$a}\right)-\frac{{$a4}}{8}\ln\left|\frac{\sqrt{$a2+x^2}}{$a}+\frac{x}{$a}\right|+C \]

\[=\frac{x($a2+x^2)^{3/2}}{4}-\frac{{$a2 x\sqrt{$a2+x^2}}}{8}-\frac{{$a4}}{8}\ln\left|\frac{\sqrt{$a2+x^2}}{$a}+\frac{x}{$a}\right|+C \]



END_SOLUTION
Context()->normalStrings;
##################################

COMMENT('Uses Flash applet in hint.');
ENDDOCUMENT();        # This should be the last executable line in the problem.
