# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 06/17/10


## DBsubject(Calculus - multivariable)
## DBchapter(Concepts for multivariable functions)
## DBsection(Parameterized surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('17.5')
## Problem1('27')
## KEYWORDS('parametric surfaces', 'parametric equations', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are( s=>'Real', t=>'Real' );
$showPartialCorrectAnswers = 1;

$r0 = random(3,9,1);
$x0 = Compute( "$r0*cos(t)*sin(s)" );
$y0 = Compute( "$r0*sin(t)*sin(s)" );
$z0 = Compute( "$r0*cos(s)" );
$s0 = Compute( "0" );
$s1 = Compute( "2*pi" );
$t0 = Compute( "0" );
$t1 = Compute( "pi" );

$param = MultiAnswer( $x0, $y0, $z0, $s0, $s1, $t0, $t1 )->with(
    singleResult=>0,
    checker=>sub{
	my ( $c, $s, $ans ) = @_;
	my ( $sx, $sy, $sz, $ss0, $ss1, $st0, $st1 ) = @{$s};

	## the answer is correct if x^2 + y^2 + z^2 = r^2 and
	##    we cover the sphere.  we check the latter by looking at
	##    appropriate pairs of points on the sphere and checking
	##    that they are a diameter apart
	my $r0sq = Formula( "$r0*$r0" );
        my $sx0 = $sx->eval( s=>$ss0, t=>$st0 );
        my $sx1 = $sx->eval( s=>$ss1, t=>$st1 );
        my $sy0 = $sy->eval( s=>$ss0, t=>$st0 );
        my $sy1 = $sy->eval( s=>$ss1, t=>$st1 );
        my $sz0 = $sz->eval( s=>$ss0, t=>$st0 );
        my $sz1 = $sz->eval( s=>$ss1, t=>$st1 );
        my $sx2 = $sx->eval( s=>($ss0+$ss1)/2, t=>($st0+$st1)/2 );
        my $sy2 = $sy->eval( s=>($ss0+$ss1)/2, t=>($st0+$st1)/2 );
        my $sz2 = $sz->eval( s=>($ss0+$ss1)/2, t=>($st0+$st1)/2 );
        my $sx3 = $sx->eval( s=>($ss0+$ss1)/2, t=>$st0 );
        my $sy3 = $sy->eval( s=>($ss0+$ss1)/2, t=>$st0 );
        my $sz3 = $sz->eval( s=>($ss0+$ss1)/2, t=>$st0 );
        my $sx4 = $sx->eval( s=>$ss0, t=>($st0+$st1)/2 );
        my $sy4 = $sy->eval( s=>$ss0, t=>($st0+$st1)/2 );
        my $sz4 = $sz->eval( s=>$ss0, t=>($st0+$st1)/2 );

        my $sx5 = $sx->eval( s=>($ss0+$ss1)/4, t=>$st0 );
        my $sy5 = $sy->eval( s=>($ss0+$ss1)/4, t=>$st0 );
        my $sz5 = $sz->eval( s=>($ss0+$ss1)/4, t=>$st0 );
        my $sx6 = $sx->eval( s=>$ss0, t=>($st0+$st1)/4 );
        my $sy6 = $sy->eval( s=>$ss0, t=>($st0+$st1)/4 );
        my $sz6 = $sz->eval( s=>$ss0, t=>($st0+$st1)/4 );
        my $sx7 = $sx->eval( s=>($ss0+$ss1)/4, t=>$st1 );
        my $sy7 = $sy->eval( s=>($ss0+$ss1)/4, t=>$st1 );
        my $sz7 = $sz->eval( s=>($ss0+$ss1)/4, t=>$st1 );
        my $sx8 = $sx->eval( s=>$ss1, t=>($st0+$st1)/4 );
        my $sy8 = $sy->eval( s=>$ss1, t=>($st0+$st1)/4 );
        my $sz8 = $sz->eval( s=>$ss1, t=>($st0+$st1)/4 );

	my @ret = ( 0,0,0, 0,0, 0,0 );
	if ( $r0sq == $sx*$sx + $sy*$sy + $sz*$sz &&
	     ( $sx0 != $sx1 || $sx1 != $sx2 || $sy0 != $sy1 || $sy1 != $sy2 ||
	       $sz0 != $sz1 || $sz1 != $sz2 ) ) {
	    $ret[0] = 1;
	    $ret[1] = 1;
	    $ret[2] = 1;
	}
        my $dsq = Compute( "4*$r0*$r0" );
	if ( ## normal parameterization: endpoints are at poles, half circle
 	     ##    rotates all the way around the sphere
	     ( $dsq == ($sx1-$sx0)*($sx1-$sx0) + ($sy1-$sy0)*($sy1-$sy0) +
	       ($sz1-$sz0)*($sz1-$sz0) &&
	       ( $dsq == ($sx3-$sx2)*($sx3-$sx2) + ($sy3-$sy2)*($sy3-$sy2) +
		 ($sz3-$sz2)*($sz3-$sz2) ||
		 $dsq == ($sx4-$sx2)*($sx4-$sx2) + ($sy4-$sy2)*($sy4-$sy2) +
		 ($sz4-$sz2)*($sz4-$sz2) ) ) ||
	     ## parameterization sweeping out a full circle, rotated
	     ##    half way around
	     ( ( $dsq == ($sx0-$sx3)*($sx0-$sx3) + ($sy0-$sy3)*($sy0-$sy3) +
	         ($sz0-$sz3)*($sz0-$sz3) &&
	         $dsq == ($sx7-$sx5)*($sx7-$sx5) + ($sy7-$sy5)*($sy7-$sy5) +
	         ($sz7-$sz5)*($sz7-$sz5) ) ||
	       ( $dsq == ($sx4-$sx0)*($sx4-$sx0) + ($sy4-$sy0)*($sy4-$sy0) +
	         ($sz4-$sz0)*($sz4-$sz0) &&
	         $dsq == ($sx8-$sx6)*($sx8-$sx6) + ($sy8-$sy6)*($sy8-$sy6) +
	         ($sz8-$sz6)*($sz8-$sz6) ) ) ) {
	    $ret[3] = 1;
	    $ret[4] = 1;
	    $ret[5] = 1;
	    $ret[6] = 1;
	}
        return [ @ret ];
    } );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find parametric equations for the sphere centered at the
origin and with radius $r0.  Use the parameters \(s\) and \(t\)
in your answer.

$PAR
\( x(s,t) = \) \{ $param->ans_rule(45) \},$BR
\( y(s,t) = \) \{ $param->ans_rule(45) \}, and $BR
\( z(s,t) = \) \{ $param->ans_rule(45) \}, where $BR
\{ $param->ans_rule(10) \} \(\le s\le \) \{ $param->ans_rule(10) \} and $BR
\{ $param->ans_rule(10) \} \(\le t\le \) \{ $param->ans_rule(10) \}.

END_TEXT
Context()->normalStrings;

ANS($param->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

We use spherical coordinates with \(\phi = s\) and \(\theta = t\) as the two
parameters.  Since the radius is $r0,  we can take
\[
x = $x0, \quad
y = $y0, \quad \mbox{and}\quad
z = $z0,
\]
with
\[ 0\le s\le\pi\quad\mbox{and}\quad 0\le t\le 2\pi. \]

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
