# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserImplicitPlane.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/18/09


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Lines with planes)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('13.4')
## Section2('13.4')
## Section3('13.4')
## Problem1('16')
## Problem2('20')
## Problem3('24')
## KEYWORDS('cross product', 'vectors')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')

Context("Vector");
Context()->flags->set( ijk=>1 );

Context()->constants->set(
  i => {TeX => "\mathit{\vec i}"},
  j => {TeX => "\mathit{\vec j}"},
  k => {TeX => "\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;

@p1 = ( non_zero_random(-2,2,1), non_zero_random(-3,3,1),
	random(-5,-1,1), random(1,3,1) );
@p2 = ( random(1,4,1), non_zero_random(-2,2,1),
	random(1,3,1), random(4,8,1) );

$n1 = Vector( (@p1[0..2]) );
$n2 = Vector( (@p2[0..2]) );

# this case is bad
while ( $n1->isParallel( $n2 ) ) {
    @p2 = ( random(1,4,1), non_zero_random(-2,2,1),
	    random(1,3,1), random(4,8,1) );
    $n2 = Vector( (@p2[0..2]) );
}
$v = $n1 x $n2;

# we must have v[2] non-zero
@vval = $v->value;
while ( $vval[2] == 0 ) {
    $p1[1] += ( $p1[1] > 0 ) ? 1 : -1;
    $n1 = Vector( (@p1[0..2]) );
    if ( $n1->isParallel( $n2 ) ) {
	$p1[1] += ( $p1[1] > 0 ) ? 1 : -1;
	$n1 = Vector( (@p1[0..2]) );
    }
    $v = $n1 x $n2;
    @vval = $v->value;
}

$vecContext = Context()->copy();

Context("ImplicitPlane");
# Context()->variables->add( z=>'Real' );

$z1 = ($vval[0]*5 + $vval[1]*5)/$vval[2];
$z2 = ($vval[0]*5 + $vval[1]*(-5))/$vval[2];
$z3 = ($vval[0]*(-5) + $vval[1]*5)/$vval[2];
$z4 = ($vval[0]*(-5) + $vval[1]*(-5))/$vval[2];
$min = $z1;
$max = $z1;
foreach my $z ( ( $z2, $z3, $z4 ) ) {
    $min = $z if ( $z < $min );
    $max = $z if ( $z > $max );
}
$min--;
$max++;

Context()->variables->set(
  x=>{limits=>[-5,5]},
  y=>{limits=>[-5,5]},
  z=>{limits=>[$min,$max]},
);

$plane = ImplicitPlane("$vval[0]*x + $vval[1]*y + $vval[2]*z = 0");

$pl1lhs = Compute("$p1[0]*x + $p1[1]*y + $p1[2]*z")->reduce();
$pl2lhs = Compute("$p2[0]*x + $p2[1]*y + $p2[2]*z")->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Consider the planes given by the equations
\[ $pl1lhs = $p1[3], \]
\[ $pl2lhs = $p2[3]. \]

$PAR
${BBOLD}(a)$EBOLD
Find a vector \(\vec v\) parallel to the line of intersection of the planes.
$BR
\(\vec v =\) \{ ans_rule(45) \}

$PAR
${BBOLD}(b)$EBOLD
Find the equation of a plane through the origin which is perpendicular
to the line of intersection of these two planes.
$BR
This plane is \{ ans_rule(55) \}


END_TEXT
Context()->normalStrings;

ANS( $v->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $correct->isParallel($student);
} ) );
ANS( $plane->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
Normal vectors of the planes are
\( \vec{n_1} =  $n1 \) and \( \vec{n_2} = $n2 \), respectively.
The line of intersection is perpendicular
to both normals. (Picture the pages in a partially open book.)
We can use \(\vec{n_1} \times \vec{n_2} = $v\).

$PAR
${BBOLD}(b)$EBOLD
We use the same normal \(\vec v = $v\) and the
point \((0,0,0)\) to get \($plane\).

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
