# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/18/09


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Planes)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(planes_vector_parallel_to_smth)
## MLTleader(1)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('13.4')
## Section2('13.4')
## Section3('13.4')
## Problem1('21')
## Problem2('25')
## Problem3('29')
## KEYWORDS('cross product', 'vectors')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')

Context("Vector");
Context()->flags->set( ijk=>1 );

Context()->constants->set(
  i => {TeX => "\mathit{\vec i}"},
  j => {TeX => "\mathit{\vec j}"},
  k => {TeX => "\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;

do {
  @n1 = ( non_zero_random(-3,3,1), non_zero_random(-5,5,1),
  	non_zero_random(-6,6,1) );
  @n2 = ( non_zero_random(-4,4,1), non_zero_random(-3,3,1),
  	non_zero_random(-4,4,1) );
  $norm1 = Vector( @n1 );
  $norm2 = Vector( @n2 );
} until ( $norm1 != $norm2 && $norm1 != -1*$norm2 );
## make sure that we didn't get parallel planes

#if ( $norm1 == $norm2 ) {
#    $n2[1]++;
#    $n2[2]++;
#    $norm2 = Vector( @n2 );
#}

$d1 = random(1,8,1);
$d2 = random(1,8,1);

$p1lhs = Compute( "$n1[0]*x + $n1[1]*y + $n1[2]*z" )->reduce();
$p2lhs = Compute( "$n2[0]*x + $n2[1]*y + $n2[2]*z" )->reduce();

$vvec = $norm1 x $norm2;

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a vector \(\vec v\) parallel to the intersection of the planes
\( $p1lhs = $d1 \) and \( $p2lhs = 7 \).

$PAR
\(\vec v =\) \{ ans_rule(55) \}

END_TEXT
Context()->normalStrings;

ANS($vvec->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student->isParallel($correct);
} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The normal vectors to the planes are \(\vec{n_1} = $norm1\) and
\(\vec{n_2}= $norm2\).  The line of intersection is perpendicular
to both normal vectors (picture the pages in a partially open book).
Hence the vector we need is
\(\vec{n_1} \times \vec{n_2} = $vvec\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
