##DESCRIPTION
##  Statistics problem: parameter identification, hypothesis testing, p-values and interpretation
##ENDDESCRIPTION

##KEYWORDS('statistics', 'statistical inference', 'hypothesis test', 'parameter identification', 'errors', 'p-values', 'interpretation')

## DBsubject('Statistics')
## DBchapter('Hypothesis Testing')
## DBsection('Elements of a Statistical Test')
## Level(5)
## Date('7/17/2012')
## Author('Djun Kim')
## Institution('UBC')


########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "unionTables.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);



# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$mc_a = RadioButtons( 
    [
    "The difference in the mean BMI of fourth graders at Laura Secord and Charles Dickens schools.",
    "The difference in the mean BMI of the children taken in the samples from Laura Secord and Charles Dickens schools.",
    "The mean BMI of children at Laura Secord and Charles Dickens schools.",
    "The difference in the variances of BMI of the children taken in the samples from Laura Secord and Charles Dickens schools.",
    "The mean BMI of children at schools in Canada.",
    ],
    # correct answer
    "The difference in the mean BMI of fourth graders at Laura Secord and Charles Dickens schools.",
    last => [], # can be a list
    );


$ans_b = Real("3.3");

$mc_c =  RadioButtons(
    [
     "There is a difference in the mean BMI for fourth graders at the two schools.",
     "There is no difference in the mean BMI for fourth graders at the two schools.",
     "The difference between the mean BMI of fourth graders at the two schools is 3.3.",
     "The mean difference between the BMI of a fourth grader at Laura Secord and the BMI of a fourth grader at Charles Dickens is 3.3.",     
     "The mean BMI of fourth graders at Laura Secord is greater than the mean BMI of fourth graders at Charles Dickens.",
    ],
    # correct answer
    "There is no difference in the mean BMI for fourth graders at the two schools.",
    last => [], # can be a list
    );

$mc_d = RadioButtons( 
   [ 
"two-sided",

"one-sided, left-tailed",

"one-sided, right-tailed",

"it does not matter whether we take a one-sided or two-sided alternative",
    ],
    # correct answer
    "two-sided",

    last => [], # can be a list
);


#------------------------------------------------------------
# setup for part e
@eqn = (
    "\( z_0 = \frac{ \bar{y} - \mu_0 } {\sigma / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{y} - \mu_0}{s / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{d}} {s_d / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{y}_1 - \bar{y}_2} {\mathrm{SE} (\bar{y}_1 - \bar{y}_2)} \)",
    );

@letter = ("A", "B", "C", "D");

$k = random(0,3,1);

@perm = shuffle(4);
@eqn = @eqn[@perm];
@inv = invert(@perm);

# The last (3rd) equation in the list above is the correct choice.
$popup = PopUp(["?","A","B","C","D"], $letter[$inv[3]]);

#------------------------------------------------------------
# setup for part f


$mc_f =  RadioButtons( 
  [
   "There is sufficient evidence to suggest that there is a difference in the mean BMI for fourth graders at the two schools.",
   
   "There is sufficient evidence to suggest that the difference between the mean BMI of fourth graders at the two schools is 3.3.",

   "There is sufficient evidence to suggest that the mean difference between the BMI of a fourth grader at Laura Secord and the BMI of a fourth grader at Charles Dickens is 3.3.",

   "There is sufficient evidence to suggest that the mean BMI of fourth graders at Laura Secord is greater than the mean BMI of fourth graders at Charles Dickens.",

   "There is insufficient evidence to suggest that there is a difference in the mean BMI for fourth graders at the two schools.",

   "There is insufficient evidence to suggest that the difference between the mean BMI of fourth graders at the two schools is 3.3.",

   "There is insufficient evidence to suggest that the mean difference between the BMI of a fourth grader at Laura Secord and the BMI of a fourth grader at Charles Dickens is 3.3.",

   "There is insufficient evidence to suggest that the mean BMI of fourth graders at Laura Secord is greater than the mean BMI of fourth graders at Charles Dickens.",
],

    "There is insufficient evidence to suggest that there is a difference in the mean BMI for fourth graders at the two schools.", 
    last => [], # can be a list 
    );


##############################################################
#
#  Text
#
#

TEXT(beginproblem());


Context()->texStrings;
BEGIN_TEXT
$BEGIN_ONE_COLUMN
In a study to compare obesity levels between two schools, random
samples of fourth graders were taken from the two schools and body mass
indices (BMI) recorded. Here is a summary of the body mass index data from
the two schools, Laura Secord and Charles Dickens.

$PAR
$BCENTER
\{ begintable(4) \}
\{ row( "School", "Number of children measured", "Average BMI", "Standard deviations of BMI" ) \}
\{ row( "Laura Secord", "7", "24.3", "3.1" ) \}
\{ row( "Charles Dickens", "6", "21.0", "2.9" ) \}
\{ endtable() \}
$ECENTER
$END_ONE_COLUMN


$PAR
$BBOLD Part a)$EBOLD What is the parameter of interest in this study?$BR

$BR
\{ $mc_a -> buttons() \}

$BR
$BR
$BBOLD Part b)$EBOLD Based on the data provided, what is your estimate of this parameter?  \{ ans_rule(10) \}

$BR
$BR 
$BBOLD Part c)$EBOLD In testing a hypothesis about the parameter of interest, what would your null hypothesis be?

$BR 
$BR 

\{ $mc_c -> buttons() \}

$BR 
$BR 
$BBOLD Part d)$EBOLD You would take the alternative hypothesis to be:

$BR 
$BR 
\{ $mc_d -> buttons() \}

$BR 
$BR 
$BBOLD Part e)$EBOLD  Which of the test statistics (A - D) you would choose to test your hypotheses?  \{ $popup->menu() \}

$BR
$BR
$BCENTER
\{
BeginTable().
  AlignedRow(["$BBOLD A $EBOLD", $eqn[0], "$BBOLD B $EBOLD" , $eqn[1]]).
  TableSpace(5,5).
#  AlignedRow(["A","B"]).
  TableSpace(25,6).
  AlignedRow(["$BBOLD C $EBOLD" , $eqn[2], "$BBOLD D $EBOLD" , $eqn[3]]).
  TableSpace(5,5).
#  AlignedRow(["C","D"]).
EndTable();
\}
$BR
$ECENTER


$BR 
$BBOLD Part f)$EBOLD If you use a 5$PERCENT level of significance, which of the following would you conclude?

$BR
$BR
\{ $mc_f -> buttons() \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

#install_problem_grader(~~&std_problem_grader);

ANS( $mc_a->cmp() );
ANS( $ans_b->cmp( 
	 checker=>sub {
	     my ( $correct, $student, $ansHash ) = @_;
	     return $correct == abs($student);
	 }
     ) );
ANS( $mc_c->cmp() );
ANS( $mc_d->cmp() );
ANS( $popup->cmp() );  # e)
ANS( $mc_f->cmp() );

Context()->normalStrings;

COMMENT('MathObject version.');


ENDDOCUMENT();        
