##DESCRIPTION
##  Statistics problem: parameter identification, hypothesis testing, p-values and interpretation
##ENDDESCRIPTION

##KEYWORDS('statistics', 'statistical inference', 'hypothesis test', 'parameter identification', 'errors', 'p-values', 'interpretation')

## DBsubject('Statistics')
## DBchapter('Hypothesis Testing')
## DBsection('Elements of a Statistical Test')
## Level(5)
## Date('7/17/2012')
## Author('Djun Kim')
## Institution('UBC')


########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "unionTables.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);



# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$mc_a = RadioButtons( 
    [
    "The difference in the mean IQ of seventh grade grade boys and girls at Vancouver Hill Elementary School.",
    "The difference in the mean IQ of the children taken in the sample from Vancouver Hill Elementary School.",
    "The mean IQ of boys and girls at Vancouver Hill Elementary School.",
    "The difference in the variances of IQ of the boys and girls taken in the sample from Vancouver Hill Elementary School.",
    "The mean IQ of children at elementary schools in Canada.",
    ],
    # correct answer
    "The difference in the mean IQ of seventh grade grade boys and girls at Vancouver Hill Elementary School.",
    last => [], # can be a list
    );


$ans_b = Real("1.6");

$mc_c =  RadioButtons(
    [
     "There is a difference in the mean IQ for boys and girls at Vancouver Hill Elementary School.",
     "There is no difference in the mean IQ for seventh grade boys and girls at Vancouver Hill Elementary School.",
     "The difference between the mean IQ of seventh grade boys and girls is 1.6.",
     "The mean difference between the IQ of a seventh grade boy at Vancouver Hill Secondary School and the IQ of a seventh grade girl at the same school is 1.6.",     
     "The mean IQ of seventh grade boys at Vancouver Hill Elementary School is greater than the mean IQ of seventh grade girls at the same school.",
    ],
    # correct answer
    "There is no difference in the mean IQ for seventh grade boys and girls at Vancouver Hill Elementary School.",
    last => [], # can be a list
    );

$mc_d = RadioButtons( 
   [ 
"two-sided",

"one-sided, left-tailed",

"one-sided, right-tailed",

"it does not matter whether we take a one-sided or two-sided alternative",
    ],
    # correct answer
    "two-sided",

    last => [], # can be a list
);


#------------------------------------------------------------
# setup for part e
@eqn = (
    "\( z_0 = \frac{ \bar{y} - \mu_0 } {\sigma / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{y} - \mu_0}{s / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{d}} {s_d / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{y}_1 - \bar{y}_2} {\mathrm{SE} (\bar{y}_1 - \bar{y}_2)} \)",
    );

@letter = ("A", "B", "C", "D");

$k = random(0,3,1);

@perm = shuffle(4);
@eqn = @eqn[@perm];
@inv = invert(@perm);

# The last (3rd) equation in the list above is the correct choice.
$popup = PopUp(["?","A","B","C","D"], $letter[$inv[3]]);

#------------------------------------------------------------
# setup for part f


$mc_f =  RadioButtons( 
  [
   "There is sufficient evidence to suggest that there is a difference in the mean IQ for seventh grade boys and girls at the school.",
   
   "There is sufficient evidence to suggest that the difference between the mean IQ of seventh grade boys and girls at the school is 1.6.",

   "There is sufficient evidence to suggest that the mean difference between the IQ of a seventh grade boy at the school and the IQ of a seventh grade girl at the school is 1.6.",

   "There is sufficient evidence to suggest that the mean IQ of seventh grade boys at the school is greater than the mean IQ of seventh grade girls at the school.",

   "There is insufficient evidence to suggest that there is a difference in the mean IQ for seventh grade boys and girls at the school.",

   "There is insufficient evidence to suggest that the difference between the mean IQ of seventh grade boys and girls at the school is 1.6.",

   "There is insufficient evidence to suggest that the mean difference between the IQ of a seventh grade boy at the school and the IQ of a seventh grade girl at the school is 1.6.",

   "There is insufficient evidence to suggest that the mean IQ of seventh grade boys at the school is greater than the mean IQ of seventh grade girls at the school.",
],

    "There is insufficient evidence to suggest that there is a difference in the mean IQ for seventh grade boys and girls at the school.",
    last => [], # can be a list 
    );


##############################################################
#
#  Text
#
#

TEXT(beginproblem());


Context()->texStrings;
BEGIN_TEXT
$BEGIN_ONE_COLUMN
In a study to compare the IQ between boys and girls in a particular elementary school, a random sample of seventh grade boys and girls was taken from Vancouver Hill Elementary School. The girls and boys were asked to take an IQ test and their scores were recorded. Some summary statistics of the IQs of the boys and girls is given below.

$PAR
$BCENTER
\{ begintable(4) \}
\{ row( "Gender", "Number of children measured", "Average IQ", "Standard deviations of IQ" ) \}
\{ row( "Boys", "17", "107.1", "5.6" ) \}
\{ row( "Girls", "15", "105.5", "4.9" ) \}
\{ endtable() \}
$ECENTER
$END_ONE_COLUMN


$PAR
$BBOLD Part a)$EBOLD What is the parameter of interest in this study?$BR

$BR
\{ $mc_a -> buttons() \}

$BR
$BR
$BBOLD Part b)$EBOLD Based on the data provided, what is your estimate of this parameter?  \{ ans_rule(10) \}

$BR
$BR 
$BBOLD Part c)$EBOLD In testing a hypothesis about the parameter of interest, what would your null hypothesis be?

$BR 
$BR 

\{ $mc_c -> buttons() \}

$BR 
$BR 
$BBOLD Part d)$EBOLD You would take the alternative hypothesis to be:

$BR 
$BR 
\{ $mc_d -> buttons() \}

$BR 
$BR 
$BBOLD Part e)$EBOLD  Which of the test statistics (A - D) you would choose to test your hypotheses?  \{ $popup->menu() \}

$BR
$BR
$BCENTER
\{
BeginTable().
  AlignedRow(["$BBOLD A $EBOLD", $eqn[0], "$BBOLD B $EBOLD" , $eqn[1]]).
  TableSpace(5,5).
#  AlignedRow(["A","B"]).
  TableSpace(25,6).
  AlignedRow(["$BBOLD C $EBOLD" , $eqn[2], "$BBOLD D $EBOLD" , $eqn[3]]).
  TableSpace(5,5).
#  AlignedRow(["C","D"]).
EndTable();
\}
$BR
$ECENTER


$BR 
$BBOLD Part f)$EBOLD If you use a 5$PERCENT level of significance, which of the following would you conclude?

$BR
$BR
\{ $mc_f -> buttons() \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

#install_problem_grader(~~&std_problem_grader);

ANS( $mc_a->cmp() );
ANS( $ans_b->cmp( 
	 checker=>sub {
	     my ( $correct, $student, $ansHash ) = @_;
	     return $correct == abs($student);
	 }
     ) );
ANS( $mc_c->cmp() );
ANS( $mc_d->cmp() );
ANS( $popup->cmp() );  # e)
ANS( $mc_f->cmp() );


Context()->normalStrings;

COMMENT('MathObject version.');


ENDDOCUMENT();        
