##DESCRIPTION
##   Describe how a confidence interval changes as the confidence level increases
##ENDDESCRIPTION


## DBsubject(Statistics)
## DBchapter(Confidence intervals)
## DBsection(One sample proportion)
## Date(10/09/2012)
## Institution(UBC)
## Author(Djun Kim)
## Level(3)
## MO(1)
## KEYWORDS('statistics', 'sampling')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());


##############################################################
#
#  Setup
#
#
Context("Numeric");
$num = random(30,40,1);
$n = random(330,340,1);
$phat = $num/$n;

$m = random(0.03, 0.05, 0.01);

$ans1 = Real(ceil(1.96**2 * ($phat) * (1 - $phat) / $m**2));
$ans2 = Real(ceil(2.00**2 * ($phat) * (1 - $phat) / $m**2));


##############################################################
#
#  Text
#
#

Context()->texStrings;

BEGIN_TEXT
Refer to the following scenario.
$PAR
An epidemiologist is worried about the prevalence of the flu in East Vancouver and the potential shortage of vaccines for the area. She will need to provide a recommendation for how to allocate the vaccines appropriately across the city. She takes a simple random sample of $n people living in East Vancouver and finds that $num have recently had the flu. 

$PAR
Suppose that the epidemiologist wants to re-estimate the population proportion and wishes
for her 95$PERCENT confidence interval to have a margin of error no larger than $m.  How large 
a sample should she take to achieve this? Please carry answers to at least six decimal places in intermediate steps.

$PAR
Sample size = \{ ans_rule(6) \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 1;

ANS( $ans2->cmp( 
   correct_ans => "$ans1 (using a critical value of 1.96)  or $ans2 (using a critical value of 2.0)",
   checker => sub {
       my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
       return (($student == $ans1 || $student == $ans2)? 1 : 0);      # return 1 if correct, 0 otherwise
   }
));



ENDDOCUMENT();
