## DESCRIPTION
## Radicals and Rational Exponents: Simplifying Radicals
## ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Date(Spring 2012)
## Institution(University of Missouri)
## Level(2)
## MO(1)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.7')
## Problem1('')
## KEYWORDS('radical expression')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
# Subclass the numeric functions
#
package my::Function::numeric;
our @ISA = ('Parser::Function::numeric');

#
# Override sqrt() to return a special value (usually 1) when evaluated
# effectively eliminating it from the product.
#
sub sqrt {
  my $self = shift;
  my $value = $self->context->flag("setSqrt");
  return $value+1 if $value && $_[0] == 1; # force sqrt(1) to be incorrect
  return $value if $value;
  return $self->SUPER::sqrt(@_);
}

#
# end of subclass
#
package main;

###########################

Context("Numeric");
Context("Numeric")->variables->are(
x => ["Real", limits => [0,2]],    # only needed if x is used in the square roots
);
#
# make sqrt() use our subclass
#
Context()->functions->set(sqrt=>{class=>'my::Function::numeric'});
Context()->flags->set(reduceConstantFunctions=>0);
Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real");

$var = list_random('x','y','z','w','p','q','r','s','t','u','v');
$var1 = list_random('x','y','z','w','t');
$var2 = list_random('p','q','r','u','v');
Context()->variables->are($var=>"Real",$var1=>"Real", $var2=>"Real");
Context()->variables->set($var=>{limits=>[1/10,1]});
Context()->variables->set($var1=>{limits=>[1/10,1]});
Context()->variables->set($var2=>{limits=>[1/10,1]});
#
# Don't allow fractional powers (avoids 1/2 power)
# [Could subclass exponentiation to handle that as well]
#
LimitedPowers::OnlyPositiveIntegers();

$a = list_random(2,3,5,6);
do {$b = list_random(2,3,5,6)} until ($b != $a);
$ab = $a*$b;
$n1 = random(2,5,1);
$n2 = random(1,5,1);
$n3 = random(1,5,1);
$n4 = random(1,5,1);
$exp1 = 2*$n1;
$exp2 = 2*$n2+1;
$exp3 = 2*$n3+1;
$exp4 = 2*$n4+1;

$expr = " \sqrt{$a $var1^{$exp1} $var2^{$exp2}}  \sqrt{$ab $var1^{$exp3} $var2^{$exp4}} ";

$ans = Compute(" $a $var1**{$n1 + $n3} $var2**{$n2 + $n4 + 1} sqrt($b $var1) ");
$texanswer = "\($a $var1^{".($n1+$n3)."} $var2^{".($n2+$n4+1)."} \sqrt{$b $var1}\)"; ##############################################################

BEGIN_TEXT
Multiply and simplify to a form where the expression under the radical is as simple as possible. Assume that all variables represent positive quantities. For example \(2x\sqrt{x}=\sqrt{4x^3}\) but \(2x\sqrt{x}\) is the simpler form.       
$PAR
$BCENTER
\( $expr = \text{ }~\) \{ans_rule(20)\}
$ECENTER
$BR Write your answer using radical notation if necessary but do not use fractional exponents.   
END_TEXT


BEGIN_TEXT
$PAR
${BBOLD}Help:${EBOLD}

To enter the square root of a number use ${BBOLD}sqrt()${EBOLD}. For example, to give \(\sqrt{7x}\) as an answer, you would type sqrt(7x). If you mistakenly type sqrt 7x instead, Webwork will take this to mean \(x \sqrt{7}\), not \(\sqrt{7x}\). That is, you MUST put parenthesis around everything that you want to take the square root of. 
$BR
END_TEXT


$showPartialCorrectAnswers = 1;

Context()->variables->set($var=>{limits=>[5,10]});
Context()->variables->set($var1=>{limits=>[5,10]});
Context()->variables->set($var2=>{limits=>[5,10]});
ANS( $ans-> cmp( checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 if $ans->{isPreview} || $correct != $student;
  #
  #  Get parsed formula for student and correct answers
  #
  $student = $ans->{student_formula};
  $correct = $correct->{original_formula} if defined $correct->{original_formula};
  #
  # check if equal when sqrt's are replaced by 1
  #
  Context()->flags->set(setSqrt => 1);
  delete $correct->{test_values};
  delete $student->{test_values};
  my $OK = ($correct == $student);
  Context()->flags->set(setSqrt => 0);
  #
  Value::Error("Check to see if your answer is simplified.") unless $OK;
  return $OK;
  }, formatStudentAnswer=>"reduced"
  )
);

COMMENT('MathObject version');
ENDDOCUMENT();
