## DESCRIPTION
## Radicals and Rational Exponents: Simplifying Radicals
## ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Date(Spring 2012)
## Institution(University of Missouri)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.7')
## Problem1('')
## KEYWORDS('radical expression')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

###################################
# Setup

Context("Numeric");

Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real");

$var = list_random('x','y','z','w','p','q','r','s','t','u','v');
$var1 = list_random('x','y','z','w','t');
$var2 = list_random('p','q','r','u','v');
Context()->variables->are($var=>"Real",$var1=>"Real", $var2=>"Real");
Context()->variables->set($var1=>{limits=>[1/10,1]});
Context()->variables->set($var2=>{limits=>[1/10,1]});

$a = random(2,5);
$b = random(2,5);
$b3 = $b**3;
$n1 = random(2,5,1);
$n2 = random(2,5,1);
$topconst = $a*$b3;
$e1 = 2*$n1 + 1;
$e2 = 2*$n2;


$frac = "\displaystyle \frac{ \sqrt{$topconst $var1^{$e1}}  }   {  \sqrt{$a $var2^{$e2}}  }";
Context()->flags->set(limits=>[-10,10]);
$num = Formula(" $b $var1^{$n1} sqrt($b $var1)")->reduce;
$den = Formula("$var2**$n2")->reduce;
if ($b == 4) {
  $texanswer = "\(\dfrac{8 $var1^$n1 \sqrt{$var1}}{$var2^$n2} \)";
} else {
  $texanswer = "\(\dfrac{$b $var1^$n1 \sqrt{$b $var1}}{$var2^$n2} \)";
}

Context()->variables->set($var=>{limits=>[5,10]});
Context()->variables->set($var1=>{limits=>[5,10]});
Context()->variables->set($var2=>{limits=>[5,10]});
$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}
Context()->operators->undefine("/","+","-");

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Simplify the radical expression and write your answer using a single radical.  Assume that all variables represent nonnegative quantities.
$PAR
$BCENTER
$displayfrac
$ECENTER
$BR  Write your answer using radical notation if necessary.  
END_TEXT



BEGIN_TEXT
$PAR
${BBOLD}Help:${EBOLD}
To enter the square root of a number use ${BBOLD}sqrt()${EBOLD}. For example, to give \(\sqrt{7x}\) as an answer, you would type sqrt(7x). If you mistakenly type sqrt 7x instead, Webwork will take this to mean \(x \sqrt{7}\), not \(\sqrt{7x}\). That is, you MUST put parenthesis around everything that you want to take the square root of. 
$BR
END_TEXT


###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp(studentsMustReduceFractions=>1));

COMMENT('MathObject version');
ENDDOCUMENT();
