## DESCRIPTION
##  Find relation among three vectors in Fun 
##  --Easy Version--
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear combinations)
## Date(07/03/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(1)
## Static(0)
## MO(1)




##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);


TEXT(beginproblem());

COMMENT('Three vectors not spanning Fun(S).  Easy version.');

#####################################
#  Setup

sub pick {
  my ($k, @in) = @_;
  my @out = ();
  while (@out < $k) { push( @out, splice(@in, random(0,$#in), 1) ); }
  @out;
}

sub Function {
  my ($f,@v) = @_;
  my $out = "$f (1) = $v[0]";
  foreach (2..($#v+1))  {
    $out = $out . "," . (($v[$_-2]>=0) ? "\phantom{-}\ " : "\ ") . 
                  "$f ($_) = " . $v[$_-1];
  }

  return $out;
}

sub maxnorm {
  my @v = map {abs $_} @_;
  $max = $v[0];
  foreach (1..($#v)) { $max=$v[$_] if ($v[$_] > $max);}
  return $max;
}

#######################

Context("Vector");
Context()->lists->set(Vector=>{open=>"(",close=>")"});

$n=3;

$S = "1"; foreach (2..$n) { $S = $S . ",$_"; }

my @c = (random(-2,2,4),random(-3,3,6),random(1,3,2),random(1,2));

$v[0] = Vector(pick($n,(-3..-1,1..3))); 
do { $v[1] = Vector(pick($n,(-3..-1,1..3))); }
until (! $v[1]->isParallel($v[0]) );

$v[2] = $c[2]*$v[0] + $c[3]*$v[1];

if (maxnorm($v[2]->value) > 7) { 
  $c[2]=1; $c[3]=1; 
  $v[2] = $c[2]*$v[0] + $c[3]*$v[1];
}
$v[0] = $c[0]*$v[0];  $v[1] = $c[1]*$v[1];

@comb = ($c[2]*$c[1],$c[3]*$c[0],-$c[0]*$c[1]);

my @i = pick($n,(0..($n-1)));
@v = @v[@i];
my @comb = @comb[@i];  

foreach (1..($#v+1)) { $v[$_-1] = Function("f_{$_}", $v[$_-1]->value) }

$multians = MultiAnswer(@comb)->with(
 singleResult=>1,
 format=>"[%s] f1 + [%s] f2 + [%s] f3",
 tex_format=>"%s f_1 + %s f_2 + %s f_3",
 checker => sub{
   my ($correct, $student, $ansHash) = @_;
   return Vector($correct)->isParallel(Vector($student));
} );

#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Let \(S=\lbrace $S \rbrace\) and \(f_1,f_2,f_3 \in\mathrm{Fun}(S)\) be given by
\[$v[0]\] \[$v[1]\] \[$v[2]\]

$PAR
Find a nontrivial linear relation among these vectors.
$PAR
\(\quad \) \{$multians->ans_rule(2);\}\(f_1\ +\ \)
           \{$multians->ans_rule(2);\}\(f_2\ +\ \)
           \{$multians->ans_rule(2);\}\(f_3\ =\ 0\)


END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 0;

ANS( $multians->cmp() );


ENDDOCUMENT(); 
