## DESCRIPTION
##  Find relation among three vectors in P_2
##  --Easy Version--
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear combinations)
## Date(07/03/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(1)
## Static(0)
## MO(1)




##################################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);


TEXT(beginproblem());

COMMENT('Three vectors not spanning P_2. Easy version.');

#####################################
#  Setup

sub pick {
  my ($k, @in) = @_;
  my @out = ();
  while (@out < $k) { push( @out, splice(@in, random(0,$#in), 1) ); }
  @out;
}

sub Polynomial {
  my @v = @_;
  my $P = Formula("0");
  foreach (0..$#v) { $P = $P + Formula("$v[$_] * x^{$_}"); }

  return $P->reduce;
}

sub maxnorm {
  my @v = map {abs $_} @_;
  $max = $v[0];
  foreach (1..($#v)) { $max=$v[$_] if ($v[$_] > $max);}
  return $max;
}


#######################

Context("Vector");
Context()->lists->set(Vector=>{open=>"(",close=>")"});

$n=3;
my @c = (random(-2,2,4),random(-3,3,6),random(1,3,2),random(1,2));

$v[0] = Vector(pick($n,(-3..-1,1..3))); 
do { $v[1] = Vector(pick($n,(-3..-1,1..3))); }
until (! $v[1]->isParallel($v[0]) );

$v[2] = $c[2]*$v[0] + $c[3]*$v[1];
if (maxnorm($v[2]->value) > 8) { 
  $c[2]=1; $c[3]=1; 
  $v[2] = $c[2]*$v[0] + $c[3]*$v[1];
}
$v[0] = $c[0]*$v[0];  $v[1] = $c[1]*$v[1];

my @comb = ($c[2]*$c[1],$c[3]*$c[0],-$c[0]*$c[1]);

my @i = pick($n,(0..($n-1)));
@v = @v[@i];
@comb = @comb[@i];  

foreach (0..$#v) { $v[$_] = Polynomial($v[$_]->value)->TeX; }


$multians = MultiAnswer(@comb)->with(
 singleResult=>1,
 format=>"[%s] p1 + [%s] p2 + [%s] p3",
 tex_format=>"%s p_1 + %s p_2 + %s p_3",
 checker => sub{
   my ($correct, $student, $ansHash) = @_;
   return Vector($correct)->isParallel(Vector($student));
} );

#####################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Let \(p_1{\small (x)} = $v[0]\), 
    \(p_2{\small (x)} = $v[1]\), 
    \(p_3{\small (x)} = $v[2]\) in 
\(\mathcal{P}_2(\mathbb{R})\).
$PAR
Find a nontrivial linear relation among these vectors.
$PAR
\(\quad \)\{$multians->ans_rule(2);\}\(p_1{\small (x)}\ +\ \)
          \{$multians->ans_rule(2);\}\(p_2{\small (x)}\ +\ \)
          \{$multians->ans_rule(2);\}\(p_3{\small (x)}\ =\ 0\)


END_TEXT
Context()->normalStrings;

#####################################
#  Answer evaluation

$showPartialCorrectAnswers = 0;

ANS( $multians->cmp() );


ENDDOCUMENT(); 
