## DESCRIPTION
##  Basis of subspace of P_4
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Basis and dimension)
## Date(07/03/2020)
## Institution(METU-NCC)
## Author(Benjamin Walter & Satoshi Kondo)
## Level(2)
## Static(0)
## MO(1)

DOCUMENT();  

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"PGcourse.pl",
"MathObjects.pl",
"parserMultiAnswer.pl"
);

sub pick {
  my ($k, @in) = @_;
  my @out = ();
  while (@out < $k) { push( @out, splice(@in, random(0,$#in), 1) ); }
  @out;
}

TEXT(beginproblem());

COMMENT("find polynomials with given (random) zeros");

Context("Numeric");
Context()->variables->are(x=>"Real");

@c = pick(3,(-2..3));


$ans1 = Formula("(x-$c[0])(x-$c[1])(x-$c[2]) ")->reduce;
$ans2 = Formula("(x-$c[0])(x-$c[1])(x-$c[2])x")->reduce;

$MAns = MultiAnswer($ans1, $ans2
)->with(
 singleResult=>0,
 allowBlankAnswers=>1,
 checkTypes=>0,
 checker=>sub{

 my ($correct,$student,$self) = @_;
 my ($c1,$c2) = @{$correct};
 my ($s1,$s2) = @{$student};

 my @result = (1,1);

 if ($c1->typeMatch($s1)) {
   my $msg = "";
   if (!($s1->D('x')->D('x')->D('x')->D('x')->D('x')->isZero)) {
     $result[0] = 0; $msg = "Wrong degree!$BR";
   }
   if ($s1->eval(x=>$c[0]) != 0) { 
     $result[0] = 0; $msg = $msg . "not zero at $c[0] $BR"; 
   } 
   if ($s1->eval(x=>$c[1]) != 0)  {
     $result[0] = 0; $msg = $msg . "not zero at $c[1] $BR"; 
   }
   if ($s1->eval(x=>$c[2]) != 0)  {
     $result[0] = 0; $msg = $msg . "not zero at $c[2] $BR"; 
   }
   if ($s1==Formula("0")) { $result[0] = 0; }
   $self->setMessage(1,$msg) if ($result[0]==0);
 } else { $result[0] = 0; }
 if ($c2->typeMatch($s2)) {
   my $msg = "";
   if (!($s2->D('x')->D('x')->D('x')->D('x')->D('x')->isZero)) {
     $result[1] = 0; $msg = "Wrong degree!$BR";
   }
   if ($s2->eval(x=>$c[0]) != 0) { 
     $result[1] = 0; $msg = $msg . "not zero at $c[0] $BR"; 
   } 
   if ($s2->eval(x=>$c[1]) != 0)  {
     $result[1] = 0; $msg = $msg . "not zero at $c[1] $BR"; 
   }
   if ($s2->eval(x=>$c[2]) != 0)  {
     $result[1] = 0; $msg = $msg . "not zero at $c[2] $BR"; 
   }
   if ($s2==Formula("0")) { $result[1] = 0; }
   $self->setMessage(2,$msg) if ($result[1]==0);
} else { $result[1] = 0; } 

 if ($c1->typeMatch($s1) && $c2->typeMatch($s2) && ($result[0] == 1) 
  && ($s1->eval(x=>2.7)*$s2->eval(x=>3.1) == $s1->eval(x=>3.1)*$s2->eval(x=>2.7)))
   { $result[1] = 0; $self->setMessage(2,"span is too small!"); } 


 return ~~@result;
});

Context()->texStrings;
BEGIN_TEXT

Let
\( \displaystyle 
  U = \Bigl\lbrace a_0+a_1 x+\cdots+ a_4 x^4 \in\mathcal{P}_4(\mathbb{R})\ 
	  \Big|\ \ p($c[0])=p($c[1])=p($c[2])=0 \Bigr\rbrace\)
$PAR
Find \(p_1{\small \!(x)}\) and \(p_2{\small \!(x)}\) whose span is \(U\).
$PAR

\(\quad U = \mathrm{Span}\Bigl\lbrace \)
\{$MAns->ans_rule(15);\} \( , \)
\{$MAns->ans_rule(15);\} \( \Bigr\rbrace \)

END_TEXT
Context()->normalStrings;

ANS( $MAns->cmp );

ENDDOCUMENT();
