# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 06/14/10


## DBsubject(Calculus - multivariable)
## DBchapter(Vector fields)
## DBsection(Graphs, flows lines, and level surfaces)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('17.4')
## Section2('17.4')
## Section3('17.4')
## Problem1('7')
## Problem2('7')
## Problem3('7')
## KEYWORDS('vector field', 'multivariable', 'calculus')

## Textbook tags
## HHChapter1('Parameterization and Vector Fields')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( y=>'Real', t=>'Real', a=>'Real', b=>'Real' );
$showPartialCorrectAnswers = 1;

$refreshCachedImages=1;

$a = random(2,5,1);
$which = list_random(1,2,3,4);
if ( $which == 1 ) {
    $vf = "$a x\vec i + $a y\vec j";
    $xp = Compute( "$a*x" ); $locxp = Compute( "$a*x" );
    $yp = Compute( "$a*y" ); $locyp = Compute( "$a*y" );
    $xs = Compute( "a*e^($a*t)" );
    $ys = Compute( "b*e^($a*t)" );
    $flowxy = "y = \frac{b}{a}\, x";
} elsif ( $which == 2 ) {
    $vf = "$a x\vec i - $a y\vec j";
    $xp = Compute( "$a*x" );   $locxp = Compute( "$a*x" );
    $yp = Compute( "-$a*y" );  $locyp = Compute( "-$a*y" );
    $xs = Compute( "a*e^($a*t)" );
    $ys = Compute( "b*e^(-$a*t)" );
    $flowxy = "x y = ab";
} elsif ( $which == 3 ) {
    $vf = "-$a x\vec i + $a y\vec j";
    $xp = Compute( "-$a*x" );  $locxp = Compute( "-$a*x" );
    $yp = Compute( "$a*y" );   $locyp = Compute( "$a*y" );
    $xs = Compute( "a*e^(-$a*t)" );
    $ys = Compute( "b*e^($a*t)" );
    $flowxy = "x y = ab";
} elsif ( $which == 4 ) {
    $vf = "-$a x\vec i - $a y\vec j";
    $xp = Compute( "-$a*x" );  $locxp = Compute( "-$a*x" );
    $yp = Compute( "-$a*y" );  $locyp = Compute( "-$a*y" );
    $xs = Compute( "a*e^(-$a*t)" );
    $ys = Compute( "b*e^(-$a*t)" );
    $flowxy = "y = \frac{b}{a}\, x";
}


Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Consider the vector field \(\vec v = $vf\).
On a sheet of paper, sketch the vector field and the flow.
$PAR
${BBOLD}(a)$EBOLD
After sketching the vector field and flow, find the system of
differential equations associated with the vector field:
$BR
\( x' = \) \{ ans_rule(35) \} $BR
\( y' = \) \{ ans_rule(35) \} $BR

$PAR
${BBOLD}(b)$EBOLD
Solve the system you found above to find the flow
$BR
\( x(t) = \) \{ ans_rule(35) \} $BR
\( y(t) = \) \{ ans_rule(35) \} $BR
${BITALIC}(Your solution for \(x\) should involve a constant \(a\),
and for \(y\) a constant \(b\).)$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $xp->cmp() );
ANS( $yp->cmp() );
ANS( $xs->cmp( checker=>sub {
		  my ( $c, $s, $ans ) = @_;
		  $sp = $s->D('t');
		  $spa = $s->D('a');
		  if ( Compute(0) == $spa ) {
		      Value->Error("Your answer should involve a constant a.");
		      return 0;
		  } else {
		      $crhs = $locxp->substitute( x=>$s );
		      return $crhs == $sp;
		  }
		 } ) );
ANS( $ys->cmp( checker=>sub {
		  my ( $c, $s, $ans ) = @_;
		  $sp = $s->D('t');
		  $spa = $s->D('b');
		  if ( Compute(0) == $spa ) {
		      Value->Error("Your answer should involve a constant b.");
		      return 0;
		  } else {
		      $crhs = $locyp->substitute( y=>"$s" );
		      return $crhs == $sp;
		  }
		 } ) );

$gr = init_graph(-3,-3,3,3, axes=>[0,0], grid=>[6,6], size=>[250,250]);
if ( $which == 2 || $which == 3 ) {
    add_functions( $gr, "1/x for x in <.3,3> using color:black and weight:2",
		   "-1/x for x in <-3,-.3> using color:black and weight:2" );

} else {
    add_functions( $gr, "1.25*x for x in <-3,3> using color:black and " .
		   "weight:2" );
}

$fn = new VectorField( sub { $xp->eval(x=>$_[0]); },
		       sub { $yp->eval(y=>$_[1]); }, $gr );
# $fn->rf_arrow_length( sub{ my($dx,$dy)=@_;
# 			   return(0) if sqrt($dx*$dx + $dy*$dy)==0;
# 			   sqrt($dx*$dx + $dy*$dy); } );
$fn->dot_radius(2);
$fn->dot_color('blue');

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
As
\[ \vec v(t)=\frac{dx}{dt}\,\vec i+\frac{dy}{dt}\,\vec j, \]
the system of differential equations is
\[ x' = $xp \]
and
\[ y' = $yp \]
$PAR

${BBOLD}(b)$EBOLD
Note that \(x'\) and \(y'\) are proportional to \(x\) and \(y\),
respectively.  Thus they are both exponentials, and we have
\[ x(t) = $xs, \qquad y(t) = $ys, \]
for some constant \(a\).
$PAR

If we eliminate \(t\) from the expressions for \(x(t)\) and \(y(t)\),
we have the non-parametric representation \($flowyx\) of the flow.
This is shown in the figure below (vectors are shown with a dot at their
tail).
$PAR
$BCENTER
\{ image(insertGraph($gr), tex_size=>250, height=>250, width=>250,
	 extra_html_tags=>'alt="graph of vector field and flow."') \}
$ECENTER

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
