# DESCRIPTION
# Factor an expression of the form x^3 + c x^2 + bd x + bdc, where bd may be positive or negative.
# Also, bd is not a perfect square so x^2 + bd does not factor any further if bd < 0. Note that a = 1 
# here.
# ENDDESCRIPTION

## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring by grouping)
## Institution(University of Missouri-Columbia)
## MLT(mizz_group)
## Level(3)
## KEYWORDS('factoring', 'factor', 'factorization', 'grouping')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(x=>"Real", "X"=>"Real");
LimitedPowers::OnlyPositiveIntegers();

do {
  $b = random(-1,1,2)*random(2,7);
  $d = random(-1,1,2)*random(2,7);
  $tmp = sqrt(abs($b*$d));
} until (floor($tmp) != $tmp);
$c = non_zero_random(-8,8);

$e2 = $c;  $e3 = $b*$d;  $e4 = $b*$d*$c;
$fac1 = "x + $c";    $afac1 = Formula($fac1)->reduce;
$fac2 = "x^2 + $e3";    $afac2 = Formula($fac2)->reduce;

$expn = "x^3 + $e2 x^2 + $e3 x + $e4";

$answer = MultiAnswer($afac1, $afac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1x^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01x^{500}") if ($s2->type eq 'String'); #bogus answer
    if (($c1 == $s1 && $c1 == $s2) || ($c2 == $s1 && $c2 == $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c1 == -$s2) || ($c2 == -$s1 && $c2 == -$s2)) {
      return [1,0];
    } elsif (($c1 == $s1 && $c2 == $s2) || ($c1 == $s2 && $c2 == $s1)) {
      return [1,1];
    } elsif (($c1 == -$s1 && $c2 == -$s2) || ($c1 == -$s2 && $c2 == -$s1)) {
      return [1,1];
    } elsif (($c1 == $s1 && $c2 != $s2) || ($c2 == $s1 && $c1 != $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c2 == $s2) || ($c2 == -$s1 && $c1 == $s2)) {
      return [0,1];
    } elsif (($c1 == $s2 && $c2 != $s1) || ($c2 == $s2 && $c1 != $s1)) {
      return [0,1];
    }
    return [0,0];
  }
);
$texanswer = "\(($fac1)($fac2)\)";


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());



COMMENT('Desc: Factor an expression of the form x^3 + c x^2 + bd x + bdc, where bd may be positive or negative. Also, bd is not a perfect square so x^2 + bd does not factor any further if bd < 0. Note that a = 1 here.');
ENDDOCUMENT();
