## DESCRIPTION
## Rational Expressions: Complex Fractions
## ENDDESCRIPTION

## KEYWORDS('Algebraic fractions')

## DBsubject('Algebra') 
## DBchapter('Basic Algebra') 
## DBsection('Algebraic Fractions') 
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.6')
## Problem1('')
## Adapted from a problem written by Paul Pearson (Fort Lewis)


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

###################################
# Setup

Context("Numeric");

Context()->variables->add(
y=>"Real",
t=>"Real",
s=>"Real",
u=>"Real",
v=>"Real",
p=>"Real",
q=>"Real",
r=>"Real",
w=>"Real",
z=>"Real");

$var = list_random('x','y','z','w','p','q','r','t','u','v');
$var1 = list_random('x','y','z','w','v');
$var2 = list_random('p','q','r','t','u');

$a = random(2,6,1);
do {$b = random(2,6,1)} until (gcd($a,$b) ==1);
$a2 = $a**2;
$ab = 2*$a*$b;
$b2 = $b**2;

$frac = "\dfrac{~ \dfrac{$a2}{$var1^2} + \dfrac{$ab}{$var1 $var2}+ \dfrac{$b2}{$var2^2} ~}{~ \dfrac{$a2}{$var1^2} - \dfrac{$b2}{$var2^2} ~}";

$num = Formula(" $a $var2 + $b $var1 ")->reduce;
$den = Formula(" $a $var2 - $b $var1  ")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}



###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Perform the indicated operation, and simplify if possible.  Assume any factors you cancel are not zero.
$PAR
$BCENTER
$displayfrac
$ECENTER

END_TEXT

if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    \(\dfrac{$num}{$den} \)
    $PAR
    END_TEXT
}

Context()->normalStrings;

PeriodicStatus();



###################################
# Answers

$showPartialCorrectAnswers = 0;

ANS($multians->cmp());

COMMENT('MathObject version');
ENDDOCUMENT();
