## DBsubject(Calculus - single variable)
## DBchapter(Differentiation)
## DBsection(Implicit differentiation)
## Institution(UCSB)
## MLT(xy_powers_quotient_trig)
## MLTleader(1)
## Level(3)
## Static(1)
## MO(1)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('5')
## Section1('3.6')
## Problem1('18')
## KEYWORDS('Differentiation','Product','Quotient')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context('Numeric');

TEXT(&beginproblem);
$showPartialCorrectAnswers = 1;
$a=random(1,10,1);
$c=($a)**2;

BEGIN_TEXT

Find \(dy/dx\) by implicit differentiation. $BR
\[\displaystyle \tan(x-y)=\frac{y}{1+x^2} \]

$PAR

\(\dfrac{dy}{dx} =\) \{ans_rule(60)\}

END_TEXT

# It is hard to recognize all right answers since the student may, or may not
# use the relation.
Context()->variables->add(y =>'Real');
$ans = Compute(" ((1+x^2)(sec(x-y))^2+2x tan(x-y))/(1+(1+x^2)(sec(x-y))^2) ");

ANS($ans->cmp(
  checker => sub {
    my ($correct,$student,$ansHash) = @_;
	return 1 if $correct == $student;
    my $altans = Compute("(2((1+x^2)^2+2x cos(x-y)^2*y))/((1+x^2)(3+2x^2+cos(2(x - y))))");
	return 1 if $altans == $student;
	return 0;
}
));

ENDDOCUMENT();
