## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Coordinate vectors and change of basis)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## TitleText1('')
## AuthorText1('')
## EditionText1('')
## Section1('')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "unionLists.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');
Context()->noreduce('(-x)-y');
Context()->noreduce('(-x)-y','(-x)+y');

$e11 = Matrix([ [1,0], [0,0] ]);
$e12 = Matrix([ [0,1], [0,0] ]);
$e21 = Matrix([ [0,0], [1,0] ]);
$e22 = Matrix([ [0,0], [0,1] ]);

$k1 = Real(non_zero_random(-4,4,1));
$k2 = Real(non_zero_random(-4,4,1));
$k3 = Real(non_zero_random(-4,4,1));
$k4 = Real(non_zero_random(-4,4,1));

$A = $k1 * $e11 + $k2 * $e22 + $k3 * $e21 + $k4 * $e12;

@a = $A->value;

$ans1 = Vector($a[0][0],$a[1][1],$a[1][0],$a[0][1]);



#$c1 = $e11 + $e22;
#$c2 = $e12 + $e21;
#$c3 = $e11 - $e22;
#$c4 = $e21 - $e12;

#$A = $k1 * $c1 + $k2 * $c2 + $k3 * $c3 + $k4 * $c4;

# $k1, $k2, $k3, $k4


$multians = MultiAnswer($a[0][0],$a[1][1],$a[1][0],$a[0][1])->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      if ( ($s[0] == $c[0]) and ($s[1] == $c[1]) and ($s[2] == $c[2])  and ($s[3] == $c[3]) ) {
          return 1;
      } else {
          return 0;
      }
  }
);

Context()->texStrings;
BEGIN_TEXT
Let 
\[
M = $A
\]
be a vector in the vector space \( V \) of \( 2 \times 2 \) matrices with real number entries.
Let
\[
\mathcal{B} = 
\left\lbrace $e11, $e22, $e21, $e12 \right\rbrace
\]
be an ordered basis for \( V \).

\{ BeginList('OL',type=>'a') \}

$ITEM
Write \( M \) as a linear combination of elements of \( \mathcal{B} \).
$BR
$BR
\( $A = \) 
\{ $multians->ans_rule(5) \}
\( $e11 + \)
\{ $multians->ans_rule(5) \}
\( $e22 + \)
\{ $multians->ans_rule(5) \}
\( $e21 + \)
\{ $multians->ans_rule(5) \}
\( $e12. \)

$ITEMSEP
$ITEM
Let \( \lbrack M \rbrack_{\mathcal{B}} \) denote the coordinate representation of \( M \) relative to the basis \( \mathcal{B} \).
Find the coordinate vector representation for \( M \) relative to the basis \( \mathcal{B} \).  
Your answer should be a vector of the general form \( \verb+<1,2,3,4>+ \).
$BR
$BR
\( \lbrack M \rbrack_{\mathcal{B}} = \)
\{ ans_rule(20) \}.

\{ EndList('OL') \}

END_TEXT
Context()->normalStrings;

ANS( $multians->cmp() );

ANS( $ans1->cmp() );

COMMENT('MathObject version.  Basis is not randomized.');

ENDDOCUMENT();
