## tcao tagged and PAID on 3-22-2004

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Date(6/3/2002)
## Institution(Rochester)
## MLT(Separable_Implicit)
## MO(1)
## Level(4)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('6')
## Section1('9.3')
## Problem1('13')
## KEYWORDS('differential equation' 'separable')
## DESCRIPTION
##  ## Differential equations
##    normalseparable.pg  must be done  previously.
##      exp(y)dy/dx= ($p x + $q )/($c sin(y) + $d cos(y))
##   URL:http://webhost.math.rochester.edu/mth163lib/discuss/msgReader$389
##
## ENDDESCRIPTION
DOCUMENT() ;        # the first executable line

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGdiffeqmacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context('Numeric');

TEXT(beginproblem()) ;
$showPartialCorrectAnswers = 1 ;

$a = random(1,7,1) ;
$b = random(2,9,1) ;
$c = $a - $b  ;
$d = $a + $b  ;
$p = random(2,18,2) ;
$q = random(2,9,1) ;
$s = $p /2 ;
$r = $q + $s ;

BEGIN_TEXT

The differential equation 
\[ \exp(y)\frac{dy}{dx} = \frac{$p x + $q}{$c \,\sin(y) + $d \,\cos(y)} \] 
has an implicit  general solution of the form \(F(x,y) = K. \) $BR $BR 
In fact, because the differential equation is  separable,
we can define the solution curve implicitly by a function in the form 
\[F(x,y) = G(x) + H(y) =K. \]
$BR $BR 
Find such a solution and then give the related functions requested. $BR
 \( F(x,y) = G(x) + H(y) = \)  \{ans_box(3,60) \}. $BR

END_TEXT
 
#$ansF = "x*($s *x + $q )/$r " ;

#$ansG = "- exp(y)*($a *sin(y) + $b * cos(y))/$r " ;

Context("Numeric");
Context()->variables->add(y=>'Real');
Context()->variables->add(A => 'Parameter');
Context()->variables->add(B => 'Parameter');
$fg =Compute("A*(x*($s *x + $q ) - exp(y)*($a *sin(y) + $b * cos(y)))+B");

ANS($fg->cmp(
  checker=> sub {
	my ($correct,$student,$ansHash) = @_;
	return 0 if(Formula($student)->isConstant);
	return $correct == $student;
  }
));

$showHint =4;
HINT(EV3(<<'EOF'));
Reminder: To find the anti-derivative of \( e^u\sin(u) \),
the trick is to do integration by parts twice.
EOF
ENDDOCUMENT() ;        # This should be the last executable line.

##################################################
my $XML_INFORMATION = <<'END_OF_XML_TRAILER_INFO';
<?xml version="1.0"?>
<metaPGdata>
        <author>Webwork Team</author>
        <course>MTH163</course>
        <description>Differential equations
  normalseparable.pg  must be done  previously.
    exp(y)dy/dx= ($p x + $q )/($c sin(y) + $d cos(y))</description>
        <fullPath>setDESep/4.pg</fullPath>
        <institution>University of Rochester</institution>
        <keywords>Differential Equation,separable,
first order nonlinear,
normalized,general solution</keywords>
        <libraryPath>setDESep/4.pg</libraryPath>
        <libraryURL>http://webhost.math.rochester.edu/mth163lib/discuss/msgReader$389</libraryURL>
        <modified><dateTime.iso8601>20000713T09:31:51</dateTime.iso8601></modified>
        <msgNum>389</msgNum>
        <pgProblem>true</pgProblem>
        <preface></preface>
        <problemVariants></problemVariants>
        <probNum></probNum>
        <psvn></psvn>
        <revisedVersions></revisedVersions>
        <setName>DESep</setName>
        <titleRoot>4</titleRoot>
        </metaPGdata>

END_OF_XML_TRAILER_INFO
##################################################
