## tcao tagged and PAID on 3-22-2004

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Date(6/3/2002)
## Institution(Rochester)
## MLT(Separable_Implicit)
## Level(4)
## TitleText1('Calculus: Early Transcendentals')
## AuthorText1('Stewart')
## EditionText1('6')
## Section1('9.3')
## Problem1('13')
## KEYWORDS('differential equation' 'separable')
## DESCRIPTION
##  ## Differential equations
##    normalseparable.pg  must be done  previously.
##  dy/dx = ($r x + $b)/($D y^2 + $E y +$C ) )
##   URL:http://webhost.math.rochester.edu/mth163lib/discuss/msgReader$388
##
## ENDDESCRIPTION
DOCUMENT() ;        # This should be the first executable line

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGdiffeqmacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem()) ;
$showPartialCorrectAnswers = 1 ;

Context('Numeric');

#Put constants and parameters here.

$a = random(1,7,1) ;
$r = 2*$a  ;
$b = random(2,9,1) ;
$A = random(2,9,1) ;
$B = random(2,9,1) ;
$C = random(2,9,1) ;
$D = 3*$A  ;
$E = 2* $B  ;

BEGIN_TEXT
             
The differential equation 
\[  \frac{dy}{dx} = \frac{$r x + $b}{$D y^2 + $E y +$C} \] 
has an implicit  general solution of the form \(F(x,y) = K. \) $BR $BR 
In fact, because the differential equation is  separable,
we can define the solution curve implicitly by a function in the form 
\[F(x,y) = G(x) + H(y) =K. \]
$BR $BR 
Find such a solution and then give the related functions requested. $BR
 \( F(x,y) = G(x) + H(y) = \)  \{ans_box(3,60) \}. 

END_TEXT
   
#$ansF = "($a * (x**2) + $b *x )/($a + $b )" ;
#$ansG = "- ($A * (y**3) + $B *(y**2) + $C *y)/($a + $b )" ;
#$ansG = " - y*(y*($A *y + $B ) + $C )/($a + $b )" ;

Context()->variables->add(y=>'Real');
Context()->variables->add(A => 'Parameter');
Context()->variables->add(B => 'Parameter');

$ans = Compute("A*(($a * (x**2) + $b *x ) - ($A * (y**3) + $B *(y**2) + $C *y))+B");

ANS($ans->cmp( checker=> sub {
  my ($correct,$student,$ansHash) = @_;
  return 0 if(Formula($student)->isConstant);
  return $correct == $student;
}));

#ANS(fun_cmp($ans,var=>['x','y']) );
ENDDOCUMENT() ;        # This should be the last executable line.





##################################################
my $XML_INFORMATION = <<'END_OF_XML_TRAILER_INFO';
<?xml version="1.0"?>
<metaPGdata>
        <author>Webwork Team</author>
        <course>MTH163</course>
        <description>Differential equations
  normalseparable.pg  must be done  previously.
dy/dx = ($r x + $b)/($D y^2 + $E y +$C ) )</description>
        <fullPath>setDESep/3.pg</fullPath>
        <institution>University of Rochester</institution>
        <keywords>Differential Equation,
separable,
first order nonlinear,
normalized,general solution</keywords>
        <libraryPath>setDESep/3.pg</libraryPath>
        <libraryURL>http://webhost.math.rochester.edu/mth163lib/discuss/msgReader$388</libraryURL>
        <modified><dateTime.iso8601>20000713T09:15:43</dateTime.iso8601></modified>
        <msgNum>388</msgNum>
        <pgProblem>true</pgProblem>
        <preface></preface>
        <problemVariants></problemVariants>
        <probNum></probNum>
        <psvn></psvn>
        <revisedVersions></revisedVersions>
        <setName>DESep</setName>
        <titleRoot>3</titleRoot>
        </metaPGdata>

END_OF_XML_TRAILER_INFO
##################################################
