## DESCRIPTION
## DBsubject(Statistics)
## Institution(The College of Idaho)
## Author(R Cruz)
## KEYWORDS('statistic', 'regression')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGstatisticsmacros.pl",
  "answerHints.pl",
  "fontCourier.pl"
  );

TEXT(beginproblem());

#################################################
#  Set-up

$n = random(8,10,1);
$number = ("eight","nine","ten");

for($i=0; $i<$n; $i++){
	$x[$i] = random(25,80,1);
	$y[$i] = random(60,100,1);
  $studentN[$i] = $i + 1;
}

$sx =0;
$sy =0;
$sxy =0;
$sx2 =0;
$sy2 =0;
$meanx =0;
$meany =0;

for($i=0; $i<$n; $i++) {
	$sx = $sx + $x[$i];
	$sy = $sy + $y[$i];
	$sxy = $sxy + ($x[$i]*$y[$i]);
	$sx2 = $sx2 + ($x[$i]**2);
	$sy2 = $sy2 + ($y[$i]**2);
	$meanx = $meanx+$x[$i];
	$meany = $meany+$y[$i];

}
$meanx = $meanx/$n;
$meany = $meany/$n;
$ssxy = $sxy-(($sx*$sy)/$n);
$ssx = $sx2-(($sx**2)/$n);
$ssy = $sy2-(($sy**2)/$n);

$r = $ssxy/sqrt($ssx*$ssy);
$b1 = $ssxy/$ssx;
$b0 = $meany-$b1*$meanx;

$SSE = $ssy - $b1*$ssxy;
$s = sqrt($SSE/($n-2));
$stat = .001*round(1000*$b1*sqrt($ssx)/$s);

$df = $n-2;

$alpha = list_random(.05,.10);
$pval = 2*tprob($df, abs($stat));

if ( $pval>$alpha ) { $tag = 0; } else { $tag = 1; }

@ans = ('There is not sufficient evidence to reject \(H_0\).',
'There is sufficient evidence to reject \(H_0\).');

$mc1 = new_multiple_choice();
$mc1 -> qa (' the test shows: ', $ans[$tag]);
$mc1 -> extra ($ans[1-$tag]);

#################################################
#  Main

BEGIN_TEXT
A study was conducted to determine whether the final grade of a 
student in an introductory psychology course is linearly related 
to his or her performance on the verbal ability test administered before
college entrance. The researchers hypothesized the linear model \(y = \beta_0 + \beta_1 x \), where \(x\) is the score on the verbal ability test and \(y\) is the final grade in the psychology course.  The verbal scores and final grades for $number[$n-8]
students are shown in the table below. 
$BR
$BCENTER
\{begintable($n+1)\}
\{row("${BBOLD}Student${EBOLD}", @studentN )\}
\{row("${BBOLD}Verbal Score${EBOLD} \(x\)", @x )\}
\{row("${BBOLD}Final Grade${EBOLD} \(y\)", @y )\}
\{endtable()\}
$ECENTER
$BR
Find the following: 
$PAR
(a)  The least squares regression line: \(\hat{y} = \) \{ans_rule(15)\}
$PAR
(b)  State the null and alternative hypotheses:  
$BR
${BITALIC}(Note: Type ${BBOLD}b_1$EBOLD for the symbol \( \beta_1 \) for the slope of the regression line.
$BR $SPACE $SPACE $SPACE $SPACE
Enter: ${BBOLD}b_1 not = 0$EBOLD for "there is a linear relationship" 
$BR  $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE
   ${BBOLD}b_1 > 0$EBOLD for "there is a positive relationship"
$BR $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE $SPACE
   ${BBOLD}b_1 < 0$EBOLD for "there is a negative relationship" )$EITALIC
$BR
  \( H_0\): \{ ans_rule(15)\}
$BR
  \( H_A\): \{ ans_rule(15)\}
$PAR 
(c) Give the test statistic for the test: \{ ans_rule(15)\}
$PAR
(d) Using \(\alpha = $alpha\), \{$mc1->print_q()\}
$BR
\{$mc1->print_a()\}
END_TEXT

#################################################
#  Answers

$showPartialCorrectAnswers = 1;

$ans_a = Formula("$b0 + $b1 x")->reduce;
ANS($ans_a->cmp);

$ans_b1 = "b_1 = 0";
ANS(ordered_str_cmp($ans_b1));

$ans_b2 = "b_1 not = 0";
ANS(ordered_str_cmp($ans_b2));

$ans_c = Compute($stat)->with(tolType=>'absolute', tolerance=>'0.1');
ANS($ans_c->cmp( checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return abs($correct -$student)<.2;
} ) );

#->withPostFilter(AnswerHints(
#     sub {
#         my ($correct,$student,$ansHash) = @_;
#         return abs($correct-$student) <= .2;
#         } => [score => 1]

#)));

ANS(radio_cmp($mc1->correct_ans));

#$ans_c = Compute($y[$choose] - ($b0 + $b1 * $x[$choose]))-#>with(tolType=>'absolute', tolerance=>'0.3');
#ANS($ans_c->cmp->withPostFilter(AnswerHints(
#    sub {
#         my ($correct,$student,$ansHash) = @_;
#         return abs(-1*$correct-$student) <= .02;
#         } => ["Remember that the formula is Observed - Expected, order matters in subtraction", score => 0],
#     sub {

#         my ($correct,$student,$ansHash) = @_;
#         return abs($correct-$student) <= .6;
#         } => ["Close! Try keeping more accuracy in your calculation.", score => 0]

#)));

######################################
#  Solution

Context()->texStrings;
BEGIN_SOLUTION
$PAR SOLUTION $PAR
$BR
Final grade: Mean = $meany
$BR
Verbal Score: Mean = $meanx
$BR
(a)   \(\hat{y} = $ans_a\)
$BR
(b)   \( H_0\) = $ans_b1
and   \( H_a\) = $ans_b2 
$BR
(c) t = $stat, df = $df and P = $pval
$BR
(d) Conclusion for \(\alpha = $alpha\): $ans[$tag] 

END_SOLUTION
ENDDOCUMENT();       # This should be the last executable line in the problem.





