## DESCRIPTION
## College Algebra, work with an exponential equation
## ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Exponential and logarithmic expressions and functions)
## DBsection(Exponential and logarithmic equations)
## Date(07/25/2016)
## Institution(Community College of Denver, Colorado Community College System)
## Author(Steven Wray)
## MO(1)
## KEYWORDS('college algebra', 'exponential equations')


###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"contextLimitedPoint.pl",
"parserOneOf.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
);

TEXT(beginproblem());


###########################
#  Setup
Context("Numeric");
install_problem_grader(~~&std_problem_grader);

# A random integer
$a=non_zero_random(2,5,1);
$b = non_zero_random(-3,-2,1);

$abs = abs($b);
$d = $a**$abs;

# The two sides of the correct answer
$ans1=Compute("$a**$b");
$ans2=Compute("1/$d");


# Allows for the correct equation to be given in either order

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2 ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $stu1 && $f2 == $stu2) ||
           ($f1 == $stu2 && $f2 == $stu1) ) {
          return [1,1];
      } else {
          if ($f1 == $stu1 || $f2 == $stu1) {
              return [1,0];
          } elsif ($f1 == $stu2 || $f2 == $stu2) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);






 
###########################
#  Main text

BEGIN_PGML


Rewrite the following equation in exponential form. 

[``\log_[$a]\left(\frac{1}{[$d]}\right)=[$b] ``] 
 
 [____]{$multians} [`=`] [____]{$multians} [@ AnswerFormatHelp("numbers") @]*



END_PGML

 
 
############################
#  Solution

#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION

COMMENT('MathObject version. Uses PGML.');

ENDDOCUMENT();
