## DESCRIPTION
## Compute a non-zero vector orthogonal to a given vector in R^2, R^3, and R^4.
##   All vectors will have small integer entries (between -3 and 5).
##   There are infinitely many possible solutions, all of which are accepted.
## ENDDESCRIPTION

## DBsubject('Geometry')
## DBchapter('Vector geometry')
## DBsection('Dot product, length, and unit vectors')
## Date('October 9, 2019')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('1.2: Lengths, Angles, and the Dot Product')
## Problem1('')
## KEYWORDS('vectors','vector length')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "PGchoicemacros.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Vector");
Context()->lists->set(Vector => {open => "(", close => ")"});
Context()->parens->set("(" => {type => "Vector", close => ")"});

do {
    $a = random(-3,5);
    $b = random(-3,5);
    $v = Vector([$a,$b]);
} until ($a != 0 || $b != 0);

$ans1 = Vector([-$b,$a]);

do {
    $a = random(-3,5);
    $b = random(-3,5);
    $c = random(-3,5);
    $w = Vector([$a,$b,$c]);
} until ($a != 0 || $b != 0 || $c != 0);

if($c == 0) {
    $ans2 = Vector([-$b,$a,0]);
} else {
    $ans2 = Vector([0,-$c,$b]);
}

do {
    $a = random(-3,5);
    $b = random(-3,5);
    $c = random(-3,5);
    $d = random(-3,5);
    $x = Vector([$a,$b,$c,$d]);
} until ($a != 0 || $b != 0 || $c != 0 || $d != 0);

if($d == 0) {
    $ans3 = Vector([-$b,$a,0,0]);
} else {
    $ans3 = Vector([0,0,-$d,$c]);
}

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find non-zero vectors that are orthogonal to the vectors listed below:
$BR
$BR
Orthogonal to \( \displaystyle $v \): \{ $ans1->ans_array \}
$BR
Orthogonal to \( \displaystyle $w \): \{ $ans2->ans_array \}
$BR
Orthogonal to \( \displaystyle $x \): \{ $ans3->ans_array \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $ans1->cmp(checker => sub {
      my ($correct,$student) = @_;
      return (abs($v . $student) < 0.0001 && $student->norm > 0.0001);
    }));

ANS( $ans2->cmp(checker => sub {
      my ($correct,$student) = @_;
      return (abs($w . $student) < 0.0001 && $student->norm > 0.0001);
    }));
    
ANS( $ans3->cmp(checker => sub {
      my ($correct,$student) = @_;
      return (abs($x . $student) < 0.0001 && $student->norm > 0.0001);
    }));
ENDDOCUMENT();
