##DESCRIPTION
##  Flux integrals, surface integrals
##ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Vector calculus)
## DBsection(Surface integrals of vector fields)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('19.2')
## Problem1('')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('19.2')
## Problem2('')
## KEYWORDS('Flux integrals','Surface integrals')

DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(x=>"Real",y=>"Real",dx=>"Real",dy=>"Real");


$a = random(2,5,1);
$b = random(2,5,1);
$c = random(2,5,1);
$d = random(2,5,1);

$xint = random(1,3,1);
$yint = random(2,4,1);


#-----------
#  integrand
#

$bam1 = $b * ($a - 1);

$f = Formula("$bam1 x + $d");


#-----------
#  flux
#

$V = Formula("-$bam1 * $yint * $xint^2 / 3 + $d * $yint * $xint + ($bam1 * $yint - $d * $yint / $xint) * $xint^2 / 2");


#=================================
#  Multiple solutions
#
#  inner and outer differentials
#
#  arrays are (if x inner & y outer integral, if y inner & x outer integral)
# 
@id = (Formula("dx"),Formula("dy"));
@od = (Formula("dy"),Formula("dx"));

# inner limit high, inner limit low, outer limit high, outer limit low (case 0, case 1)
@ilh = (Compute("$xint - $xint/$yint * y"),Compute("$yint - $yint/$xint * x"));
@ill = (Formula("0"),Formula("0"));
@olh = (Formula("$yint"),Formula("$xint"));
@oll = (Formula("0"),Formula("0"));


$multipart = MultiAnswer( $f, $id[0], $od[0], $oll[0], $olh[0], $ill[0], $ilh[0] )->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $fstu, $idstu, $odstu, $ollstu, $olhstu, $illstu, $ilhstu ) = @{$student};
        if ( 
             ( 
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           return 1;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration.");
           return 0.94;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration and order of integration.");
           return 0.47;
        } else {
             return 0;
        }
    }
  );



##############################################
#  Main text

Context()->texStrings;
Context()->flags->set(reduceConstants=>0);
BEGIN_TEXT
Set up a double integral for calculating the flux of 
\( \vec{F} = $a x \vec{i} + y \vec{j} + z \vec{k} \)
through the part of the surface 
\( z = -$b x - $c y + $d \) above the triangle
in the xy-plane with vertices \( (0,0) \), 
\( (0,$yint) \), and \( ($xint,0) \), oriented 
upward. 
$BR
$BR
${BITALIC}Instructions:${EITALIC}
Please enter the integrand in the first answer box. 
Depending on the order of integration you choose,  
enter ${BITALIC}dx${EITALIC} and ${BITALIC}dy${EITALIC} 
in either order into the second and third answer boxes 
with only one ${BITALIC}dx${EITALIC} or 
${BITALIC}dy${EITALIC} in each box.  Then, enter the limits of 
integration and evaluate the integral to find the flux. 
$BR
$BR
Flux = 
\( \displaystyle \int_A^B \!\! \int_C^D \) 
\{ $multipart->ans_rule(40) \}
\{ $multipart->ans_rule(5) \}
\{ $multipart->ans_rule(5) \}
$BR
$BR
A = \{ $multipart->ans_rule(20) \} $BR
B = \{ $multipart->ans_rule(20) \} $BR
C = \{ $multipart->ans_rule(20) \} $BR
D = \{ $multipart->ans_rule(20) \}
$BR
$BR
Flux = 
\( \displaystyle \iint\limits_S \vec{F} \cdot d\vec{A} \) =
\{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;





##################################################
#  Answers

ANS($multipart->cmp() );
ANS(($V)->cmp());

$showPartialCorrectAnswers = 1;

COMMENT('Allows integration in either order.');


;
ENDDOCUMENT();
