## DESCRIPTION
## Jacobian matrix
## ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Change of variable)
## Date(08/24/2012)
## Institution(Hope College)
## Author(Paul Pearson)
## MLT(Jacobian_LinearApprox)
## Level(3)
## TitleText1('Multivariable Mathematics')
## AuthorText1('Darin Stephenson')
## EditionText1('Version 3.0')
## Section1('10.2')
## Problem1('6')
## KEYWORDS('Jacobian','Jacobian matrix')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


#############################
#  Setup

Context("Numeric")->variables->are(
x=>'Real',
y=>'Real',
);

# construct the function
$a = random(2,7,1);
do { $b = random(2,3,1); } until ($b != $a);
do { $c = random(2,7,1); } until ($c != $a && $c != $b);
do { $d = random(2,7,1); } until ($d != $a && $d != $b && $d != $c);

$f1 = Formula("$a x^2 + x y + $b y^3");
$f2 = Formula("$c x + $d y + sin(x y)");

# construct the Jacobian
$f1x = $f1->D('x');   $f1y = $f1->D('y');
$f2x = $f2->D('x');   $f2y = $f2->D('y');

# coordinates of the point P
$i = random(0,1,1);
if ($i == 0) { 
  $Px = non_zero_random(-5,5,1);
  $Py = 0;
} else {
  $Px = 0;
  $Py = non_zero_random(-5,5,1);
}

# values of f at P
$f1P = $f1->eval(x=>$Px,y=>$Py);
$f2P = $f2->eval(x=>$Px,y=>$Py);

# entries of Jacobian at the point
$f1Px = $f1x->eval(x=>$Px,y=>$Py);   $f1Py = $f1y->eval(x=>$Px,y=>$Py);
$f2Px = $f2x->eval(x=>$Px,y=>$Py);   $f2Py = $f2y->eval(x=>$Px,y=>$Py);

# linear approximation
$L1 = Formula("$f1P + $f1Px (x - $Px) + $f1Py (y - $Py)");
$L2 = Formula("$f2P + $f2Px (x - $Px) + $f2Py (y - $Py)");

# nearby point Q
$Qx = $Px + list_random(0.01,0.02,0.03,0.04,0.05);
$Qy = $Py - list_random(0.01,0.02,0.03,0.04,0.05);

# linear approximation at Q
$L1Q = $L1->eval(x=>$Qx,y=>$Qy);
$L2Q = $L2->eval(x=>$Qx,y=>$Qy);

# function f at Q
$f1Q = $f1->eval(x=>$Qx,y=>$Qy);
$f2Q = $f2->eval(x=>$Qx,y=>$Qy);


$multians1 = MultiAnswer($L1, $L2)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);



$multians2 = MultiAnswer($f1Q, $f2Q)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians3 = MultiAnswer($L1Q, $L2Q)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);



######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Consider the function
$PAR
\[
f(x,y) = \left( \begin{array}{c} $f1 \\ $f2 \end{array} \right).
\]
$BR
(a) Use the Jacobian \( Jf(\boldsymbol{\vec{a}}) \) to 
construct a linear (or affine) approximation
\( L(\boldsymbol{\vec{x}}) \) to \( f \) near the point 
\( \boldsymbol{\vec{a}} = \left( \begin{array}{r} $Px \\ $Py \end{array} \right). \)
$BR
\{ 
mbox(
  "\( L(\boldsymbol{\vec{x}}) = \)",
  display_matrix([ [$multians1->ans_rule(40)], [$multians1->ans_rule(40)] ], align=>'c')
);
\}
$BR
(b) Find the values of the function \( f \) and the linear approximation 
\( L \) at the nearby point \( \boldsymbol{\vec{b}} = \left( \begin{array}{r} $Qx \\ $Qy \end{array} \right). \)
$BR
\{ 
mbox(
  "\( f(\boldsymbol{\vec{b}}) = \)",
  display_matrix([ [$multians2->ans_rule(40)], [$multians2->ans_rule(40)] ], align=>'c')
);
\}
$BR
\{ 
mbox(
  "\( L(\boldsymbol{\vec{b}}) = \)",
  display_matrix([ [$multians3->ans_rule(40)], [$multians3->ans_rule(40)] ], align=>'c')
);
\}
END_TEXT
Context()->normalStrings;

########################
#  Answer evaluation

#install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );
ANS( $multians3->cmp() );


COMMENT("(Uses the same randomization as 10-02-Jacobian-03b.pg)");

ENDDOCUMENT();


