## DESCRIPTION
## Introduction to differential equations
## ENDDESCRIPTION

## KEYWORDS('differential equations')

## DBsubject('Differential Equations')
## DBchapter('First Order Differential Equations')
## DBsection('Introduction to Differential Equations')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Notes on Diffy Qs')
## EditionText1('December 9, 2010')
## AuthorText1('Jiri Lebl')
## Section1('0.2')
## Problem1('10')
## TitleText2('Differential Equations Computing and Modeling')
## EditionText1('4')
## AuthorText1('Edwards and Penney')
## Section1('1.1')
## Problem1('')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric")->variables->add(
t=>"Real"
);


$k = random(2,5,1);
$kk = $k**2;

$answer = Formula("$k cos(t)");

$expr = "(x')^2 + x^2 = $kk";

Context()->flags->set(
  reduceConstants=>0, # no decimals
  reduceConstantFunctions=>1, # combine 4+5*2?
  formatStudentAnswer=>'parsed', # no decimals
);


#############################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Find a non-constant solution to \( $expr \) using your 
knowledge of derivatives from basic calculus.
$BR
$BR
\( x(t) = \)
\{ ans_rule(30) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


############################
#  Answers

$showPartialCorrectAnswers = 1;

ANS( $answer->cmp( 
  checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    return 0 unless $student->isFormula;
    my $studentderiv = $student->D('t');
    if ( ($studentderiv)**2 + ($student)**2 == $kk ) {
      return 1;
    } else {
      return 0;
    }
  } ) );


COMMENT("MathObject version.");

ENDDOCUMENT();
