##DESCRIPTION
##ENDDESCRIPTION
## DBsubject('Statistics')
## DBchapter('Confidence intervals')
## DBsection('One sample proportion')
## Level(3)
## MO(1)
## KEYWORDS('sample proportion', 'confidence interval')
##Tagged by Nelson Chen. Jun 26th, 2016. 


########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "parserRadioButtons.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());


##############################################################
#
#  Setup
#
#
Context("Numeric");

$num = random(20,25,1);
$n = random(120, 130, 1);

$mc1 = new_multiple_choice();
$mc1->qa(
"The proportion of the $n people who are living below the poverty line, $num/$n, is a: ", 
"statistic."
);
$mc1->extra(
"parameter.",
"variable of interest.",
);

@n = rserve_eval('
n <- round(runif(1, min=120, max=130))
');

@m = rserve_eval('
m <- round(runif(1, min=20, max=25))
');

@a = rserve_eval('
a <- m/n
round(a, 3)
');

# Part ii)
# Acceptable answers use z = 2 or z = 1.96

$phat = $num / $n;

$left_endpoint = $phat - 2*sqrt($phat * (1-$phat) / $n);
$right_endpoint = $phat + 2*sqrt($phat * (1-$phat) / $n);

$left_endpoint = sprintf "%.3f", $left_endpoint;
$right_endpoint = sprintf "%.3f", $right_endpoint;

$left_ans1 = $phat - 2*sqrt($phat * (1-$phat) / $n);
$right_ans1 = $phat + 2*sqrt($phat * (1-$phat) / $n);

$left_ans1 = sprintf "%.3f", $left_ans1 ;
$right_ans1 = sprintf "%.3f", $right_ans1 ;

$left_ans2 = $phat - 1.96*sqrt($phat * (1-$phat) / $n);
$right_ans2 = $phat + 1.96*sqrt($phat * (1-$phat) / $n);

$left_ans2 = sprintf "%.3f", $left_ans2 ;
$right_ans2  = sprintf "%.3f", $right_ans2 ;

$multians = MultiAnswer($left_endpoint, $right_endpoint)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $anshash, $self ) = @_;
      my ( $stu_ans_left, $stu_ans_right ) = @{$student};
      my ( $ans1, $ans2 ) = @{$correct};
      $ansHash->{correct_ans} = "($left_ans1, $right_ans2) OR ($left_ans2, $right_ans2)";
      if ( ($stu_ans_left == $left_ans1 && $stu_ans_right == $right_ans1) ||
           ($stu_ans_left == $left_ans2 && $stu_ans_right == $right_ans2) ) 
      {
          return [1,1];
      } 
      else {
          return [0,0];
      }
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

A government official is in charge of allocating social programs throughout the city of Vancouver.  
He will decide where these social outreach programs should be located based on the percentage 
of residents living below the poverty line in each region of the city.  He takes a simple random 
sample of $n people living in Gastown and finds that $num have an annual income that is below 
the poverty line. 

$PAR
$BBOLD Part i) $EBOLD
\{ $mc1->print_q() \}
$BR
\{ $mc1->print_a() \}


$PAR
$BBOLD Part ii) $EBOLD
Use the sample data to compute a 95$PERCENT confidence interval for the true proportion 
of Gastown residents living below the poverty line. 

$PAR
(Please carry answers to at least six decimal places in intermediate steps. Give your final answer to the nearest three decimal places).

$PAR
95$PERCENT confidence interval = (\{ $multians->ans_rule(6) \}, \{ $multians->ans_rule(6) \})

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 1;

ANS( radio_cmp( $mc1->correct_ans() ) );
ANS( $multians->cmp() ); 

## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
(i) The sample proportion is a function of the data, and so
is a statistic. (ii) The sample proportion, @a[0], is an
estimate of the population proportion, \(p\) say, who live below the
poverty line. Since the sample proportion is approximately Normal with mean 
\(p\) and standard deviation \( \left( 1-p\right) p/\mathtt{n,}\) we
can approximate a 95% confidence interval for \(p\) via
\[
@m[0]/@n[0] \pm 1.96 \sqrt{\frac{\left( @m[0]/@n[0] \right) \left( 1-@m[0]/@n[0]\right) }{@n[0]}}.
\]

END_SOLUTION


ENDDOCUMENT();
