##DESCRIPTION
## DBsubject('Statistics)
## DBchapter('Confidence intervals')
## DBsection('One sample proportion)
## Level(3)
## MO(1)
##ENDDESCRIPTION
##KEYWORDS('sample proportion', 'confidence interval', 'margin of error',
##'sample size')
##Attempts: Two. 

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());


##############################################################
#
#  Setup
#
#
Context("Numeric");

@nR = rserve_eval('
nR <- round(runif(1, min=120, max=130))
');

@mR = rserve_eval('
mR <- round(runif(1, min=20, max=25))
');

@meR = rserve_eval('
meR <- sample(c(0.01, 0.05, 0.03), 1)
');

@aR = rserve_eval('
aR <- mR/nR
round(aR, 3)
');

@sdR = rserve_eval('
sdR <- sqrt((mR/nR)*(1-mR/nR)/nR)
round(sdR, 3)
');


$num = @mR[0];
$n = @nR[0];
$phat = $num/$n;

$m = @meR[0];

$ans1 = Real(ceil(1.96**2 * ($phat) * (1 - $phat) / $m**2));
$ans2 = Real(ceil(2.00**2 * ($phat) * (1 - $phat) / $m**2));

$ans1_new = Real(1.96**2 * ($phat) * (1 - $phat) / $m**2);


##############################################################
#
#  Text
#
#

Context()->texStrings;

BEGIN_TEXT
A government official is in charge of allocating social programs throughout the city of Vancouver.  
He will decide where these social outreach programs should be located based on the percentage 
of residents living below the poverty line in each region of the city.  He takes a simple random 
sample of $n people living in Gastown and finds that $num have an annual income that is below 
the poverty line. 

$PAR
Suppose that the government official wants to re-estimate the population proportion and wishes
for his 95$PERCENT confidence interval to have a margin of error no larger than $m.  How large 
a sample should he take to achieve this? Please carry answers to at least six decimal places in intermediate steps.

$PAR
Sample size = \{ ans_rule(6) \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 1;

ANS( $ans2->cmp( 
   correct_ans => "$ans1 (using a critical value of 1.96)  or $ans2 (using a critical value of 2.0)",
   checker => sub {
       my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
       return (($student == $ans1 || $student == $ans2)? 1 : 0);      # return 1 if correct, 0 otherwise
   }
));


## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
The point estimate of the population proportion \(p\)
based on the data is \( \hat{p}_{1}=@mR[0]/@nR[0]=@aR[0].\) The standard
deviation of this estimate is estimated to be \(\sqrt{\left( @mR[0]/@nR[0]
\right) \left( 1-@mR[0]/@nR[0]\right) /@nR[0]}=@sdR[0].\) It is suggested a
second sample could be taken, of size \(n\) say. Suppose the estimate
of the population proportion based on this second sample is \(\hat{p}_{2}.\) 
Before the sample is taken, we could use \(\hat{p}_{1}\) to
estimate the standard deviation of \(\hat{p}_{2},\) this estimate being 
\(\sqrt{\left( @mR[0]/@nR[0]
\right) \left( 1-@mR[0]/@nR[0]\right) /@nR[0]}.\) The margin of error for the estimate is half the width of the confidence
interval, this being 
\[
1.96\sqrt{\left( @mR[0]/@nR[0] \right) \left( 1-@mR[0]/@nR[0]\right) /@nR[0]},
\]
and we are told this should be no more than @meR. Hence
we require n such that
\[
1.96\sqrt{\left( @mR[0]/@nR[0]\right) \left( 1- @mR[0]/@nR[0]\right) /n}<@meR[0]. 
\]
Rearranging, we see that we require
\[
n>\frac{1.96^{2}\left( @mR[0]/@nR[0]\right) \left( 1-@mR[0]/@nR[0]\right) }{
\left( @meR[0]\right) ^{2}}= $ans1_new.
\]
Hence the sample size must be at least $ans1. 
END_SOLUTION


ENDDOCUMENT();
