##DESCRIPTION
##  College Algebra Review: Computing Absolute Values
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of exponents)
## Date(Spring 2012)
## Institution(University of Missouri)
## MLT(Division_exponents_multivar)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('.2')
## Problem1('')
## KEYWORDS('algebra', 'absolute value')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
###################################
# Setup

Context("Numeric");

$b = random(2,4,1);
$a = random(3,6,1)*$b;
$c = random(1,$a/$b-1,1)*$b;
$m1 = random(7,20,1);
$n1 = random(7,20,1);
$m2 = $m1 - random(2,$m1-2,1);
$n2 = $n1 - random(2,$n1-2,1);

Context("Numeric")->variables->are(x=>"Real", y=>"Real");

LimitedPowers::OnlyPositiveIntegers();
$frac = "\frac{$a x^{$m1}y^{$n1}}{$c x^{$m2} y^{$n2} } ";

#Context(LimitedPowers::OnlyPositiveIntegers());
$numconst = $a/gcd($a,$c);
$denconst = $c/gcd($a,$c);
$numexp1 = ($m1-$m2);
$numexp2 = ($n1-$n2);
if ($denconst == 1) {
  $texanswer = "\( $numconst x^{$numexp1}y^{$numexp2}  } \)";
} else {
  $texanswer = "\(\displaystyle \dfrac{$numconst x^{$numexp1}y^{$numexp2}}{$denconst  } \)";
}
$den = Formula("$denconst ")->reduce;
$num = Formula("$numconst x^{$numexp1} y^{$numexp2} ")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);
Context()->operators->undefine("+","-","/");


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER


END_TEXT
Context()->normalStrings;

###################################
# Answers


Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";
$showPartialCorrectAnswers = 1;
Context()->operators->undefine("+","-");
ANS($multians->cmp() );

Context()->texStrings;

  BEGIN_TEXT
$PAR
  ${BBOLD}Note: $EBOLD  If your answer is not a fraction, write "1" in the denominator.
$BR
  END_TEXT

Context()->normalStrings;

ENDDOCUMENT();
