# DESCRIPTION
# Factor an expression of the form a x^2 + b x + c, where a > 1, b < 0, and c < 0.
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Expanding and Factoring')
## KEYWORDS('factoring', 'factor', 'factorization', 'substitution')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedPolynomial.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("LimitedPolynomial-Strict");
$var = list_random('a','b','m','n','x','y','z','w','p','q','r','s','t','u','v');
Context()->variables->are($var=>"Real", uc($var)=>"Real");
LimitedPowers::OnlyPositiveIntegers();

do {
  $a = random(1,4);
  do {$b = non_zero_random(-4,4)} until (gcd($a,$b) == 1);
  $c = random(1,4);
  do {$d = non_zero_random(-4,4)} until (gcd($c,$d) == 1);
  $e1 = $a*$c;
  $e2 = $a*$d+$b*$c;
  $e3 = $b*$d;
} until ($e1 > 1 && $e2 < 0 && $e3 < 0 && ($a != $c && $b != $d));

$fac1 = ($a == 1) ? "$var + $b" : "$a $var + $b";    $afac1 = Formula($fac1)->reduce;
$fac2 = ($c == 1) ? "$var + $d" : "$c $var + $d";    $afac2 = Formula($fac2)->reduce;

$expn = "$e1 $var^2 + $e2 $var + $e3";
if ($e2 == 1) {
  $expn = "$e1 $var^2 + $var + $e3";
} elsif ($e2 == -1) {
  $expn = "$e1 $var^2 - $var + $e3";
}

$answer = MultiAnswer($afac1, $afac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("$var^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("$var^{500}") if ($s2->type eq 'String'); #bogus answer
    if (($c1 == $s1 && $c1 == $s2) || ($c2 == $s1 && $c2 == $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c1 == -$s2) || ($c2 == -$s1 && $c2 == -$s2)) {
      return [1,0];
    } elsif (($c1 == $s1 && $c2 == $s2) || ($c1 == $s2 && $c2 == $s1)) {
      return [1,1];
    } elsif (($c1 == -$s1 && $c2 == -$s2) || ($c1 == -$s2 && $c2 == -$s1)) {
      return [1,1];
    } elsif (($c1 == $s1 && $c2 != $s2) || ($c2 == $s1 && $c1 != $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c2 == $s2) || ($c2 == -$s1 && $c1 == $s2)) {
      return [0,1];
    } elsif (($c1 == $s2 && $c2 != $s1) || ($c2 == $s2 && $c1 != $s1)) {
      return [0,1];
    }
    return [0,0];
  }
);
$texanswer = "\(($fac1)($fac2)\)";


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
Context()->texStrings;
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} $texanswer
  $PAR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization. Edited and updated in 2012/2013.
$BR
Desc: Factor an expression of the form a x^2 + b x + c, where a > 1, b < 0, and c < 0.');
ENDDOCUMENT();
