# DESCRIPTION
# Factor an expression of the form x^(2n) + b x^n + c.
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Expanding and Factoring')
## KEYWORDS('factoring', 'factor', 'factorization', 'substitution')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
$var1 = list_random('a','b','m','n','x','y','z','w','p','q','r','s','t','u','v');
do {$var2 = list_random('a','b','m','n','x','y','z','w','p','q','r','s','t','u','v')} until ($var2 ne $var1); 
Context()->variables->are($var1=>"Real", uc($var1)=>"Real", $var2=>"Real", uc($var2)=>"Real");
LimitedPowers::OnlyPositiveIntegers();

do {
  $a = random(1,4,1);
  do {$b = random(-1,1,2)*random(2,5,1)} until (gcd($a,$b) == 1);
  do {$c = random(2,4,1)} until ($c != $a);
  do {$d = random(-1,1,2)*random(2,5,1)} until (gcd($c,$d) != 1);
} until ($a*$d + $b*$c != 0);

$e1 = $a*$c;
$e2 = $a*$d + $b*$c;
$e3 = $b*$d;
$fac1 = ($a == 1) ? "$var1 + $b $var2" : "$a $var1 + $b $var2";    $afac1 = Formula($fac1)->reduce;
$fac2 = "$c $var1 + $d $var2";    $afac2 = Formula($fac2)->reduce;

$expn = "$e1 $var1^2 + $e2 $var1 $var2 + $e3 $var2^2";
if ($e2 == 1) {
  $expn = "$e1 $var1^2 + $var1 $var2 + $e3 $var2^2";
} elsif ($e2 == -1) {
  $expn = "$e1 $var1^2 - $var1 $var2 + $e3 $var2^2";
}

$answer = MultiAnswer($afac1, $afac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1$var^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01$var^{500}") if ($s2->type eq 'String'); #bogus answer
    if (($c1 == $s1 && $c1 == $s2) || ($c2 == $s1 && $c2 == $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c1 == -$s2) || ($c2 == -$s1 && $c2 == -$s2)) {
      return [1,0];
    } elsif (($c1 == $s1 && $c2 == $s2) || ($c1 == $s2 && $c2 == $s1)) {
      return [1,1];
    } elsif (($c1 == -$s1 && $c2 == -$s2) || ($c1 == -$s2 && $c2 == -$s1)) {
      return [1,1];
    } elsif (($c1 == $s1 && $c2 != $s2) || ($c2 == $s1 && $c1 != $s2)) {
      return [1,0];
    } elsif (($c1 == -$s1 && $c2 == $s2) || ($c2 == -$s1 && $c1 == $s2)) {
      return [0,1];
    } elsif (($c1 == $s2 && $c2 != $s1) || ($c2 == $s2 && $c1 != $s1)) {
      return [0,1];
    }
    return [0,0];
  }
);
$texanswer = "\(($fac1)($fac2)\)";


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
Context()->texStrings;
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} $texanswer
  $PAR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization. Edited and updated in 2012/2013.
$BR
Desc: Factor an expression of the form x^(2n) + b x^n + c.');
ENDDOCUMENT();
