##DESCRIPTION
##  College Algebra Review: Factoring quadratics with leading term 1
##ENDDESCRIPTION

##KEYWORDS('algebra', 'factoring', 'quadratic')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Expanding and Factoring')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.4')
## Problem1('')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PeriodicRerandomization.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(x=>"Real");

$c = list_random(2,3,4,5);
$do {$n = list_random(2,3,4,5)} until (gcd($n,$c) == 1);
$m = random(2,5,1)*$c;
$a = $m*$n/$c;
$b = $m + $n;


$fac1 = Formula(" $n x + $c ");
$fac2 = Formula(" ($m/$c) x + 1  ")->reduce;
$texanswer = "\(\left($n x + $c\right)\left(".($m/$c)."x + 1\right)\)";
$expression = "$a x^2 + $b x + $c";
if ($b == -1) {
  $expression = "$a x^2 - x + $c";
} elsif ($b == 1) {
  $expression = "$a x^2 + x + $c"; 
}

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
   
      if ( ( $f1==$f1stu  &&   $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ||
           ( $f1==$f2stu  &&   $f2==$f1stu) || 
           (-$f1==$f2stu && -$f2==$f1stu) ) {
          return [1,1];}
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text
Context()->texStrings;
BEGIN_TEXT

Factor the expression completely:
$PAR
\( $expression = \big( \) \{ $multians->ans_rule(10) \}
\( \big) \big( \) \{ $multians->ans_rule(10) \}
\( \big) \)
END_TEXT


if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    $texanswer
    $PAR
    END_TEXT
} else {
  BEGIN_TEXT
  $BR
  END_TEXT
}

Context()->normalStrings;

PeriodicStatus(); 


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

COMMENT('MathObject version', 'Features PeriodicRerandomization');
ENDDOCUMENT();
