##DESCRIPTION
## solve by integrating factor and separation of variables 
##ENDDESCRIPTION

##KEYWORDS('linear' 'differential' 'equation')

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Linear Equations')
## Date('1/7/2010')
## Author('Darwyn Cook')
## Institution('Alfred University')
## TitleText1('Differential Equations: with Boundary Value Problems')
## EditionText1('6')
## AuthorText1('Zill')
## Section1('2.3')
## Problem1('')

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserFormulaUpToConstant.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

#######################################################################
## 
## Students that use a TI-89 to do this problem are going to get into 
## trouble, give them a hint.

$showHint = 5; 

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;
$context = Context("Numeric")->copy;
$context->variables->are("r"=>"Real","theta"=>"Real",C=>"Real");
$context->variables->set(theta => {TeX => '\theta'});

# The TI-89 gives an answer for the integrating factor that is not at all 
# helpful, disable this answer which involves a fraction. Also the domain
# is set up so that absolute values are not necessary, so remove those from the 
# context.

$context->functions->disable("ln","abs");
$context->parens->remove("|");
$context->operators->undefine("/");
$context->{error}{msg}{"Function 'ln' is not allowed in this context"} 
= "e raised to a logarithmic power can be reduced";
$context->{error}{msg}{"Unexpected character '|'"} 
= "Look at graph of sec and tan over the given interval and reconsider the absolute value signs.";
$context->{error}{msg}{"Can't use '/' in this context"} 
= "If you have a fraction in your answer you need to simplify to terms of sec and tan. If you don't have sec and tan in your answer, you should.";
$reduce = Context("Numeric")->copy;
$reduce->functions->undefine("sec","tan");
$reduce->variables->are("r"=>"Real","theta"=>"Real",C=>"Real");
$reduce->variables->set(theta => {TeX => '\theta'});
$reduce->{error}{msg}{"Function 'sec' is not allowed in this context"} 
= "the correct expression can be reduced to sines and cosines";
$reduce->{error}{msg}{"Function 'tan' is not allowed in this context"} 
= "the correct expression can be reduced to sines and cosines";

$contextC = Context("Numeric")->copy;
$contextC->variables->are("theta"=>"Real");
$contextC->variables->set(theta => {TeX => '\theta'});
##############################################################
#
#  Setup
#
#

$integratingfactor = Formula($context,"sec(theta)+tan(theta)");# If you change this you must change the answer check down below
$newleft = Formula($context,"r*(sec(theta)+tan(theta))");
$integrateleft = $newleft;
$solution = Formula($reduce,"(theta-cos(theta)+C)*cos(theta)/(1+sin(theta))");
$newright = Formula($reduce,"1+sin(theta)");
$integrateright = FormulaUpToConstant($contextC,"theta-cos(theta)");

Context()->texStrings;
BEGIN_TEXT
Note that you must simplify your answers for this problem. \(\theta\) can be typed in as theta.
$BR
The linear differential equation \(\frac{dr}{d\theta}+r\sec(\theta) = \cos(\theta)\), for \( 0\leq\theta < \frac{\pi}{2}^*\), has integrating factor \{$integratingfactor->ans_rule\} 
$BR
After multiplying both sides by the integrating factor, simplifying, and unapplying the product rule we get the new differential equation:
$BR
\(\frac{d}{d\theta}\big[\) \{$newleft->ans_rule(20)\} \(\big]\) \(=\) \{$newright->ans_rule(20)\}
$BR 
Integrating both sides we get the algebraic equation
\{$integrateleft->ans_rule(30)\}=\{$integrateright->ans_rule(30)\}
$BR
Solving for r, the solution to the differential equation is r = \{$solution->ans_rule(20)\} (use C for the constant)
$BR
$BR
\({}^*\) That interval is NOT chosen at random, it has important implications for the absolute values you are about to get.
END_TEXT
Context()->normalStrings;

ANS($integratingfactor->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $contextIF = $context->copy;
      $contextIF->flags->set(no_parameters=>0);
      $contextIF->variables->add('C0'=>'Parameter');
      my $c0 = Formula($contextIF,'C0');
      $student = Formula($contextIF,$student);
      $correct = Formula($contextIF,"$c0*(sec(theta)+tan(theta))");
      return $correct == $student;
    }));
 # Any constant multiple of the integrating factor will work.
ANS($newleft->cmp());
ANS($newright->cmp());
ANS($integrateleft->cmp());
ANS($integrateright->cmp());
ANS($solution->cmp());

BEGIN_HINT
$BR
The odds are you struggling with this problem because you did the integral of sec on your TI-89 and it gave you a not useful answer: \(\int\sec(\theta)d\theta=\ln(|\sec(\theta)+\tan(\theta)|)+C\) is simpler. If you pay close attention to the interval for \(\theta\) then the integrating factor simplifies greatly. You don't need the +C either.
END_HINT
ENDDOCUMENT();
