## DESCRIPTION
## Calculus
## ENDDESCRIPTION


## Tagged by tda2d

## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Change of variable)
## Date(01/30/2011)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## KEYWORDS('Multivariable','change of variables')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGmatrixmacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


###########################
#  Setup

Context("Point")->variables->are(u=>'Real',v=>'Real');

$hshift = non_zero_random(-5,5,1);
do { $vshift = non_zero_random(-5,5,1); } until ( $hshift != $vshift );
do { $hscale = random(2,7,1); } until ( $hscale != $hshift && $hscale != $vshift);
do { $vscale = random(2,6,1); } until ( $vscale != $hshift && $vscale != $vshift && $vscale != $hscale);

$x = Formula("$hshift + $hscale u");
$y = Formula("$vshift + $vscale v");

$P = Point("($hshift,$vshift)");
$Q = Point("($hshift + $hscale, $vshift)");
$R = Point("($hshift, $vshift + $vscale)");
$S = Point("($hshift + $hscale, $vshift + $vscale)");



$multians1 = MultiAnswer($hscale, 0, 0, $vscale)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$det = Real("$hscale * $vscale");



###########################
#  Main text

Context()->texStrings;
BEGIN_TEXT
(a) Find an affine change of coordinates that takes the unit 
square with vertices 
$PAR
\[
P^* = (0,0), \ 
Q^* = (1,0), \
R^* = (0,1), \
S^* = (1,1) 
\]
$BR
in the \( uv \)-plane to the rectangle with vertices 
$PAR
\[
P = $P, \ 
Q = $Q, \
R = $R, \
S = $S 
\]
$BR
in the \( xy \)-plane.
$BR
$BR
\( x = x(u,v) = \) \{ ans_rule(20) \}
$BR
\( y = y(u,v) = \) \{ ans_rule(20) \}
$BR
$BR
(b) Find the absolute value of the determinant of the 
Jacobian for this change of coordinates.
$PAR
$BCENTER
\{ 
mbox(
  "\( \displaystyle \left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \bigg| \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  "\( \bigg| = \)".$SPACE.ans_rule(20)
);
\}
$ECENTER
$PAR
END_TEXT
Context()->normalStrings;


############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $x->cmp() );
ANS( $y->cmp() );

ANS( $multians1->cmp() );
ANS( $det->cmp() );




ENDDOCUMENT();
