## DESCRIPTION
## Calculus
## ENDDESCRIPTION


## Tagged by tda2d

## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Change of variable)
## Date(01/30/2011)
## Institution(Hope College)
## Author(Paul Pearson)
## MLT(JacobDeterm)
## MLTleader(1)
## Level(3)
## KEYWORDS('Multivariable','change of variables')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


###########################
#  Setup

Context("Numeric")->variables->are(u=>'Real',v=>'Real');
Context()->noreduce('(-x)-y','(-x)+y');

do { 

  $a = non_zero_random(-5,5,1);
  $b = non_zero_random(-5,5,1);
  $c = non_zero_random(-5,5,1);
  $d = non_zero_random(-5,5,1);

} until ( $a * $d != $b * $c );

$e = non_zero_random(-5,5,1);
$f = non_zero_random(-5,5,1);

$x = Formula("$b v + $a u + $e")->reduce();
$y = Formula("$f + $c u + $d v")->reduce();

$det = Real("abs(($a)*($d)-1*($b)*($c))");

$areaUV = random(7.01,7.49,0.01);

$area = Compute("$areaUV * $det");


$multians1 = MultiAnswer($a, $b, $c, $d)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);





###########################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose a change of coordinates 
\( T : \mathbb{R}^2 \to \mathbb{R}^2 \) 
from the \( uv \)-plane to the \( xy \)-plane is given by
$PAR
\[
x = $x, \ \ \ y = $y.
\]
$BR
(a) Find the  absolute value of the determinant of the
Jacobian for this change of coordinates.
$PAR
$BCENTER
\{ 
mbox(
  "\( \displaystyle \left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \bigg| \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  "\( \bigg| = \)".$SPACE.ans_rule(20)
);
\}
$ECENTER
$BR
(b) If a region \( D^* \) in the \( uv \)-plane has area 
\( $areaUV \), find the area of the region \( T(D^*) \)
in the \( xy \)-plane.
$BR
$BR
Area = \{ ans_rule(20) \}
END_TEXT
Context()->normalStrings;


############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $det->cmp() );

ANS( $area->cmp() );




ENDDOCUMENT();
