# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 11/18/2008

## Author('Gavin LaRose');
## DBsubject(Calculus - single variable)
## DBchapter(Limits and continuity)
## DBsection(Estimating limits numerically)
## Institution(University of Michigan)
## MLT(limit_numerical_fill_table_and_guess)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('1.8')
## Problem1('15')
## KEYWORDS('calculus','derivative','limits','continuity','piecewise functions')

## Textbook tags
## HHChapter('A Library of Functions')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
$showPartialCorrectAnswers = 1;

$a = random(2,7,1);
$asq = $a*$a;
$func = Compute("(x^2 - $asq)/(x - $a)");

@xvals = ();
@vals = ();
foreach my $x ( -0.1, -0.01, -0.001, -0.0001, 0.0001, 0.001, 0.01, 0.1 ) {
    push( @xvals, $a+$x );
    push( @vals, $func->eval( x=>$a+$x ) );
}

$limit = Compute("2*$a");

## this time the window is easy to calculate; the function is, for x ne a,
##    linear: f(x) = x + a.
$delta = Compute("0.01");
$y0 = $limit - $delta;
$y1 = $limit + $delta;

$mp = MultiAnswer( $a-$delta, $a+$delta, $y0, $y1 )->with(
    singleAnswer=>0,
    checker=>sub {
	my ( $correct, $student, $self ) = @_;

	## make sure we don't get fooled by small interval lengths
	my $currentContext = Context();
	Context()->flags->set( tolType=>"absolute", tolerance=>0.000000001 );

	## first check: if we match -delta, delta we're ok
	if ( $student->[0] == $correct->[0] &&
	     $student->[1] == $correct->[1] &&
	     $student->[2] == $correct->[2] && 
	     $student->[3] == $correct->[3] ) {
	    return [ 1, 1, 1, 1 ];

	## otherwise, we check that the function values match the
	##    requirement
	} else {
	    if ( ($student->[0] - $student->[1]) >= 0 ) {
		$self->setMessage(1, "Your interval doesn't make sense.  " .
				  "Check that your first value is less " .
				  "than the second.");
		Context( $currentContext );
		return [ 0, 0, 0, 0 ];
	    } elsif ( $student->[2] >= $student->[3] ) {
		$self->setMessage(3, "Your interval doesn't make sense.  " .
				  "Check that your first value is less " .
				  "than the second.");
		Context( $currentContext );
		return [ 0, 0, 0, 0 ];
	    } elsif ( $student->[0] >= $a || $student->[1] <= $a ) {
		$self->setMessage(1, "Your interval doesn't make sense.  " .
				  "Check that your interval spans x=$a. ");
		Context( $currentContext );
		return [ 0, 0, 0, 0 ];
	    } else {
		$fx0 = $func->eval( x=>$student->[0] );
		$fx1 = $func->eval( x=>$student->[1] );

		my @ret = ( 0, 0, 0, 0 );

		## the y-range makes sense if it's of height <= .02 and
		##    includes the limit
		my ( $sy0, $sy1 ) = ( Compute($student->[2]),
				      Compute($student->[3]) );
		if ( ($sy0 < $limit) && ($sy1 > $limit) &&
		     (($sy1 - $sy0) <= 0.02) ) {
		    $ret[2] = 1;
		    $ret[3] = 1;
		} elsif ( $limit <= $sy0 ) {
		    $self->setMessage(3,"Check that your interval [$sy0] includes " .
				      "the limiting value ($limit).");
		} elsif ( $limit >= $sy1 ) {
		    $self->setMessage(3,"Check that your interval [$sy1] includes " .
				      "the limiting value ($limit).");
		} else {
		    $self->setMessage(3,"Your y range must be 0.02 or less " .
				      "in height.");
		}
		## the x-range makes sense if the function values at the
		##    endpoints are between the indicated y-values.
		##    because we know the function is concave down at x=0
		if ( $fx0 >= $sy0 && $fx1 <= $sy1 ) {
		    $ret[0] = 1;
		    $ret[1] = 1;
		} elsif ( $fx0 >= $student->[2] ) {
		    $ret[0] = 1;
		} elsif ( $fx1 >= $student->[2] ) {
		    $ret[1] = 1;
		} else {
		    if ( $fx0 < $student->[2] ) {
			$self->setMessage(3, "Check that the function " .
					  "really leaves through the " .
					  "left side of your window.");
		    }
		    if ( $fx0 < $student->[3] ) {
			$self->setMessage(4, "Check that the function " .
					  "really leaves through the " .
					  "right side of your window.");
		    }
		}

		Context( $currentContext );
		return [ @ret ];
	    }
	}
    }
);


TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider the function \(f(x) = $func\).
$PAR
${BBOLD}(a)$EBOLD
Fill in the following table of values for \(f(x)\):
$BR
\{ begintable(9) \}
\{ row( "\(x =\)", @xvals ) \}
\{ row( "\(f(x) = \)", ans_rule(10), ans_rule(10), ans_rule(10), ans_rule(10),
	ans_rule(10), ans_rule(10), ans_rule(10), ans_rule(10) ) \}
\{ endtable() \}
$PAR

${BBOLD}(b)$EBOLD
Based on your table of values, what would you expect the limit of \(f(x)\)
as \(x\) approaches $a to be?
$BR
\(\lim\limits_{x\to $a}\,$func = \) \{ ans_rule(10) \}

$PAR
${BBOLD}(c)$EBOLD
Graph the function to see if it is consistent with your answers to
parts (a) and (b).  By graphing, find an interval for \( x \) near \( $a \)
such that the difference between your conjectured limit and the
value of the function is less than 0.01.  In other words, find a
window of height 0.02 such that the graph exits the sides of the window
and not the top or bottom.  What is the window?
$BR
\{ $mp->ans_rule(10) \} \( \le x \le \) \{ $mp->ans_rule(10) \},
$BR
\{ $mp->ans_rule(10) \} \( \le y \le \) \{ $mp->ans_rule(10) \}.

END_TEXT
Context()->normalStrings;
Context()->{format}{number} = "%.7f#";

ANS($vals[0]->cmp(tolType => 'absolute', tolerance => .00000001 ) );    
ANS($vals[1]->cmp(tolType => 'absolute', tolerance => .00000001 ) );
ANS($vals[2]->cmp(tolType => 'absolute', tolerance => .00000001 ) );	   
ANS($vals[3]->cmp(tolType => 'absolute', tolerance => .00000001 ) );
ANS($vals[4]->cmp(tolType => 'absolute', tolerance => .00000001 ) );	   
ANS($vals[5]->cmp(tolType => 'absolute', tolerance => .00000001 ) );
ANS($vals[6]->cmp(tolType => 'absolute', tolerance => .00000001 ) );	   
ANS($vals[7]->cmp(tolType => 'absolute', tolerance => .00000001 ) );

ANS($limit->cmp() );
ANS($mp->cmp() );

$gr = init_graph( $a-$delta, 2*$a-$delta, $a+$delta, 2*$a+$delta,
		  axes=>[$a,$limit-.005],
		  ticks=>[4,4], size=>[250,250] );
$amd = $a - $delta;      $ame = $a - $delta/100;
$apd = $a + $delta;      $ape = $a + $delta/100;
$lmd = $limit - $delta;  $lpd = $limit + $delta;
add_functions($gr,
	      "$func for x in <$amd,$ame> using color:blue " .
	        "and weight:2",
	      "$func for x in <$ape,$apd> using color:blue " .
	        "and weight:2");
$gr->stamps( open_circle($a,$limit,'blue') );
$gr->lb( new Label($amd,$limit-0.005,"$amd",'black',"left","top") );
$gr->lb( new Label($apd,$limit-0.005,"$apd",'black',"right","top") );
$gr->lb( new Label($a,$limit+0.01,"$lpd",'black',"left","top") );
$gr->lb( new Label($a,$limit-0.01,"$lmd",'black',"right","bottom") );

$lmd = $limit - $delta;
$lpd = $limit + $delta;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The values of \(f(x)\) are
$BR
\{ begintable(9) \}
\{ row( "\(x =\)", @xvals ) \}
\{ row( "\(f(x) = \)", $vals[0], $vals[1], $vals[2], $vals[3],
	$vals[4], $vals[5], $vals[6], $vals[7] ) \}
\{ endtable() \}
$PAR

${BBOLD}(b)$EBOLD
From the values in this table, it appears that the limit is $limit.
$PAR

${BBOLD}(c)$EBOLD
When we graph the function \(f(x)\) we see the graph:
$BR
${BCENTER}
\{ image( insertGraph( $gr ), tex_size=>250, height=>250, width=>250,
          extra_html_tags=>'alt="graph of a line ' .
          'with a positive slope, on ' .
          'the x domain [-' . $amd . ',' . $apd . '] and y range [' .
          $lmd . ',' . $lpd . ']; the curve has a discontinuity at (' .
          $a . ',' . $limit . '), and function leaves the graphed window ' .
	  'at the bottom-left and top-right corners of the graph."' ) \}
$ECENTER
$BR
This confirms our estimate of the limit.  A reasonable range for \(x\)
so that the graph of \(f(x)\) enters and leaves a window of height 0.02
around \(y = $limit\) is \($amd \le x \le $apd\) and
\($lmd\le y\le $lpd\).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
