##DESCRIPTION
## Word problem
## 
##ENDDESCRIPTION
## DBsubject('Intermediate Algebra')
## DBchapter('Ch 01: Algebra, Mathematical Models and Problem Solving')
## DBsection('Algebra, Mathematical Models, and Problem Solving')
## KEYWORDS('mathematical models','problem solving','formulas','algebra')
## TitleText1('Essentials of Intermediate Algebra')
## EditionText1('1')
## AuthorText1('Blitzer')
## Section1('1.5')
## Problem1('')
## Author('RA Cruz')
## Institution('The College of Idaho')
## Date: 2007/09

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

######################################
#  Setup

@vars = ("x","y","z","p","q","m","n","k","u","v","c","h","r","s");
$c1 = random(0,scalar(@vars)-1,1);
do {$c2 = random(0,scalar(@vars)-1,1)} until ($c2 != $c1);
do {$c3 = random(0,scalar(@vars)-1,1)} until ($c3 != $c1 && $c3 != $c2);
do {$c4 = random(0,scalar(@vars)-1,1)} until ($c4 != $c1 && $c4 != $c2 && $c4 != $c3);
$var1 = $vars[$c1];
$var2 = $vars[$c2];
$var3 = $vars[$c3];
$var4 = $vars[$c4];
Context()->variables->are($var1=>"Real", $var2=>"Real", $var3=>"Real", $var4=>"Real");

for ($i = 0; $i < random(10,50); $i++) { #more random??
  $case = random(0,1);
  $a = random(2,10);
  $b = random(3,10);
}
if ($case == 0) {
  $dispeqn = "$var1 = \dfrac{1}{$a} $var2 $var3^2 + $b $var4 $var3";
  $num = Formula("$a ($var1 - $b $var4 $var3)")->reduce;
  $den = Formula("$var3^2")->reduce;
  $var = $var2;
} else {
  $dispeqn = "$var1 + \dfrac{1}{$b} $var4 $var3 = \dfrac{1}{$a} $var2 $var3^2";
  $num = Formula("$b ($var2 $var3^2 - $a $var1)")->reduce;
  $den = Formula("$a $var3")->reduce;
  $var = $var4;
}
Context()->texStrings;
$dispans = "$var = \frac{$num}{$den}";


$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $var = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $var = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

######################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Solve the equation for \( $var\): 
$PAR
\[ 
$dispeqn 
\]
$PAR
$displayfrac
END_TEXT

######################################
#  Answers

ANS($multians->cmp);

$showPartialCorrectAnswers = 1;

###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle $dispans\)
  $PAR
  END_TEXT
} 
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
