# WeBWorK problem written by Alex Jordan 
# Portland Community College
#
# Template:
# Use rational exponents to write the expression. 
#
#       n-th-root(x)n-th-root(x)
#
#
# Last updated: Spoddeck 11/11/15
# ENDDESCRIPTION


##############################################


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedPowers.pl",
  "PGML.pl",
  "PCCmacros.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);


##############################################


TEXT(beginproblem());

Context("Numeric");
Context()->operators->set(
'/' => {class => 'bizarro::BOP::divide', isCommand => 1},
' /' => {class => 'bizarro::BOP::divide', isCommand => 1},
'/ ' => {class => 'bizarro::BOP::divide', isCommand => 1},
'//' => {class => 'bizarro::BOP::divide', isCommand => 1},
);
$var = list_random('x','y','z','t','r','m','n','a','b','c');
Context()->variables->are($var=>'Real',limits=>[0,1]);
Context()->flags->set(reduceConstants=>0);

$n = random(3,11,2);



Context()->functions->undefine('sqrt','abs');



$ans = Formula("$var^(2/$n)");

$evaluator = $ans->cmp(
     checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroDiv=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is not in the expected form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroDiv=>0);
      return 1;
  }
);

##############################################


BEGIN_PGML

Simplify the expression, answering with rational exponents and not radicals. [@KeyboardInstructions('To enter [`x^{m/n}`], type [|x^(m/n)|]*.')@]**

[``\sqrt[[$n]]{[$var]}\,\sqrt[[$n]]{[$var]}=``] [______]{$evaluator}

END_PGML


##############################################


BEGIN_PGML_SOLUTION

We will use the following properties:

    [`` \sqrt[m]{x^n}=x^{\frac{n}{m}} ``]

    [`` x^m\cdot x^n=x^{m+n} ``]

The solution is:

    [``
\begin{aligned}
   \sqrt[[$n]]{[$var]}\sqrt[[$n]]{[$var]} &= [$var]^\frac{1}{[$n]}[$var]^\frac{1}{[$n]} \\
                                                         &= [$var]^\left(\frac{1}{[$n]}+\frac{1}{[$n]}\right) \\
                                                         &= [$var]^\frac{2}{[$n]} \\
\end{aligned}
    ``]

So the answer in calculator/WeBWorK notation is [`[$ans]`].

END_PGML_SOLUTION


##############################################


ENDDOCUMENT();
