##DESCRIPTION
#
# File Created: 3/19/2009
# Last Modified: 3/19/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##

## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Variation of parameters)
## Date(3/19/2009)
## Institution(Alfred University)
## Author(Juan Marin)
## MLT(vp_multistep)
## Level(3)
## TitleText1('Differential Equations with Boundary-Value Problems')
## AuthorText1('Zill and Cullen')
## EditionText1('7')
## Section1('4.6')
## Problem1('6')


DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);
 
## Show partial correct answers
$showPartialCorrectAnswers = 1;
## Display the problem information
TEXT(beginproblem());

## Declare variables
Context()->variables->are(x=>'Real',y=>'Real');

$a   = random(2,11);
$a2 = $a * $a;
$b = 3 * $a;
$c = 2 * $a * $a;
$Yh1 = Formula("e^(-$a x)")->reduce;
$Yh2 = Formula("e^(-2 * $a x)")->reduce;
$g = Formula("sin(e^($a x))")->reduce;





$DYh1 = $Yh1->D('x')->reduce;
$DYh2 = $Yh2->D('x')->reduce;
$W = Formula("$Yh1 * $DYh2 - $Yh2 * $DYh1")->reduce;
$W1 = Formula("-$g * $Yh2")->reduce;
$W2 = Formula("$g * $Yh1")->reduce;


$Du1 = Formula("$W1/$W")->reduce;
$Du2 = Formula("$W2/$W")->reduce;
$u1 = Formula("-cos(e^($a x))/$a2")->reduce;
$u2 = Formula("(e^($a x)cos(e^($a x))-sin(e^($a x)))/$a2")->reduce;






  $Yp = Formula("$Yh1 *$u1 +$Yh2 * $u2")->reduce;

  $multipart4 = MultiAnswer($Yh1, $Yh2)->with(
    singleResult => 0,
    checker => sub {
        my ( $correct4, $student4, $self4 ) = @_;
        my ( $f1stu4, $f2stu4 ) = @{$student4};
        my ( $f14, $f24 ) = @{$correct4};
        if ( ($f14 == $f1stu4 && $f24 == $f2stu4) ||
             ($f14 == $f2stu4 && $f24 == $f1stu4) ) {
            return [1,1];
        } else {
            if ($f14 == $f1stu4 || $f24 == $f1stu4) {
                return [1,0];
            } elsif ($f14 == $f1stu4 || $f24 == $f2stu4) {
                return [0,1];
            } else {
                return [0,0];
            }
        }
    }
  );


Context()->texStrings;
BEGIN_TEXT
In this problem you will solve the non-homogeneous differential equation
\[y''+ $b y'+ $c y =$g \] 
$PAR
(1) Let \(C_1\) and \(C_2\) be arbitrary constants.  The general solution to the related homogeneous differential equation  
\(y''+ $b y'+ $c y =0\) is the function 
\(y_h(x)=C_1\ y_1(x)+C_2\ y_2(x)=C_1\) \{ans_rule(10)\} \(+C_2\)\{ans_rule(10)\}.
$BR $BR
$BBOLD
NOTE: The order in which you enter the answers is important; that is,
\(C_1 f(x) + C_2 g(x) \neq C_1 g(x) + C_2 f(x)\).
$EBOLD
$PAR
(2) The particular solution \(y_p(x)\) to the differential equation 
\[y''+ $b y'+ $c y =$g \]  is of the form 
\( y_p(x)=y_1(x)\ u_1(x)+y_2(x)\ u_2(x)\) where 
\(u_1'(x)=\) 
\{ans_rule(25)\}
and \(u_2'(x)=\)
\{ans_rule(25)\}.
$PAR
(3) It follows that 
\(u_1(x)=\)
\{ans_rule(30)\}
and \(u_2(x)=\)
\{ans_rule(40)\}; thus 
\(y_p(x)=\) \{ans_rule(40)\}.
$PAR
(4) The most general solution to the non-homogeneous differential equation
\(y''+ $b y'+ $c y =$g \)  is $BR
\(y=C_1\)
\{$multipart4->ans_rule(6)\}
\(+C_2\)\{$multipart4->ans_rule(6)\}\(+\)\{ans_rule(40)\}.




END_TEXT



Context()->normalStrings;

ANS($Yh1->cmp);
ANS($Yh2->cmp);
ANS($Du1->cmp);
ANS($Du2->cmp);
ANS($u1->cmp);
ANS($u2->cmp);
ANS($Yp->cmp);
ANS( $multipart4->cmp() );
ANS($Yp->cmp);
ENDDOCUMENT();
