## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Matrix algebra)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(2)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.2')
## Problem1('')
## KEYWORDS('linear algebra','matrix','multiplication')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context('Matrix');

$A = Matrix([[0,1,0],[1,0,0],[0,0,1]]);
$B = Matrix([[1,0,0],[0,0,1],[0,1,0]]);

$multians = MultiAnswer($A, $B)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      $s0 = Matrix($s[0]); $s1 = Matrix($s[1]);
      return $s0 * $s1 != $s1 * $s0;
  }
);

Context()->texStrings;
BEGIN_TEXT
Give an example of two \( 3 \times 3 \) matrices \( A \) and \( B \)
such that \( AB \ne BA \).
$BR
$BR
\( A = \)
\{ $multians->ans_array(3,3,5) \}
$BR
$BR
\( B = \)
\{ $multians->ans_array(3,3,5) \}
END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $multians->cmp() );



ENDDOCUMENT();
