# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('1.3')
## Problem1('2')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserPopUp.pl",
);

Context("Numeric");

$nEqn = random(2,5,1);
$nVar = random(2,$nEqn+1,1);

$rank1 = Compute(1);
$rank2 = Compute( $nVar );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

$PAR
${BBOLD}(a)$EBOLD
Suppose we have a linear system with $nEqn equations.
If the system is inconsistent, what is a possible value for the
rank of the coefficient matrix of the system?
$BR
rank = \{ $rank1->ans_rule() \}

$PAR
${BBOLD}(b)$EBOLD
Suppose we have a linear system with $nVar variables.
If the system has exactly one solution, what is a possible value
for the rank of the coefficient matrix of the system?
$BR
rank = \{ $rank2->ans_rule() \}


END_TEXT
Context()->normalStrings;

ANS( $rank1->cmp( checker=>sub{ my ($c, $s, $a) = @_; return $s < $nEqn && $s >= 0; } ) );
ANS( $rank2->cmp() );

ENDDOCUMENT();
