# DESCRIPTION
# Reading homework problem from Linear Algebra with Applications
# Bretscher, 5e
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('reading homework')
## Tagged by glr 14/9/1

  ## DBsubject('Linear Algebra')
  ## DBchapter('')
  ## DBsection('')
  ## Date('09/01/14')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Linear Algebra with Applications')
## EditionText1('5')
## AuthorText1('Bretscher')
## Section1('1.1')
## Problem1('2')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
);

Context("Numeric");
Context()->variables->are( x=>'Real', y=>'Real' );

$a11 = random(-5,5,1);
$a12 = random(-5,5,1);
$a13 = random(-5,5,1);
$a21 = random(-5,5,1);
$a22 = random(-5,5,1);
$a23 = random(-5,5,1);

$lhs1 = Compute( "$a11 x + $a12 y" );
$lhs2 = Compute( "$a21 x + $a22 y" );
$rhs1 = Formula( "$a13" );
$rhs2 = Formula( "$a23" );

$sys = MultiAnswer( $lhs1, $rhs1, $lhs2, $rhs2 )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ansHash ) = @_;
	my @cor = @{$c};
	my @stu = @{$s};
	if ( $cor[0] == $stu[0] && $cor[1] == $stu[1] && $cor[2] == $stu[2]
	     && $cor[3] == $stu[3] ) {
	    return [1, 1, 1, 1];
	} else {
	    @res = ( 0, 0, 0, 0 );
	    if ( ( $cor[0] == $stu[0] && $cor[1] == $stu[1] ) ||
		 ( $cor[0] == $stu[1] && $cor[1] == $stu[0] ) ) {
		if ( ( $cor[2] == $stu[2] && $cor[3] == $stu[3] ) ||
		     ( $cor[2] == $stu[3] && $cor[3] == $stu[2] ) ) {
		    return [1, 1, 1, 1];
		} else {
		    return [1, 1, 0, 0];
		}
	    } elsif ( ( $cor[0] == $stu[2] && $cor[1] == $stu[3] ) ||
		      ( $cor[0] == $stu[3] && $cor[1] == $stu[2] ) ) {
		if ( ( $cor[2] == $stu[0] && $cor[3] == $stu[1] ) ||
		     ( $cor[2] == $stu[1] && $cor[3] == $stu[0] ) ) {
		    return [1, 1, 1, 1];
		} else {
		    return [0, 0, 1, 1];
		}
	    } else {
		return [0, 0, 0, 0];
	    }
	}
} );

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Consider the augmented matrix
\[\begin{bmatrix} $a11 & $a12 & $a13\\ $a21 & $a22 & $a23 \end{bmatrix}.\]
What system of equations (in \(x\) and \(y\)) does this represent?
$BR
\{ $sys->ans_rule() \} = \{ $sys->ans_rule() \}, $BR
\{ $sys->ans_rule() \} = \{ $sys->ans_rule() \}

END_TEXT
Context()->normalStrings;

ANS( $sys->cmp() );

ENDDOCUMENT();
