#########################################################
# DESCRIPTION
## DBsubject('Probability')
## DBchapter('Sample space')
## DBsection('Bayes theorem - inverse probability')
## Level(3)
## Date('2016/08/23')
## Author('Nelson Chen')
## Institution('University of British Columbia')
##KEYWORDS('Probability', 'Bayes theorem', 'tree diagram')

#########################################################
DOCUMENT();
## Initializations: (Required)
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "unionTables.pl",
  "answerCustom.pl",
  "PGcourse.pl"
);

#########################################################
## Setup: this is where we use Perl and PG objects (Required)

Context("Numeric");
# warn("Student name: $studentName, login: $studentLogin, id: $studentID");
my $hash = crypt($studentLogin, 'a1');
# warn("crypt: $hash");
@nums = (0..9,'a'..'z','A'..'Z');
my %nums = map { $nums[$_] => $_ } 0..$#nums;
my $seed = 0;
$seed = $seed * 62 + $nums{$_} foreach split(//, substr($hash, -5, 5));
# warn("seed: " . $seed);

rserve_eval("set.seed($seed)");

$ans = Real("0.19");

#########################################################
## Main Text: where all text goes (Required)
Context()->texStrings;
BEGIN_TEXT
Suppose we are looking at a large database of the English language
such as the IPhOD dictionary (Vaden \(et~al\). 2009). For any word we
select, there is an 81% probability that it starts with a consonant and a
19% probability that it starts with a vowel. If it starts with a consonant,
there is a 21% chance that the second sound is also a consonant, while if
the word starts with a vowel, there is a 99% chance that the second sound
is a consonant. (These probabilities were calculated using the
"Phonological Search" function of the
free software Phonological CorpusTools, by Hall \(et~al\). 2015). Give
answers to the following to two decimals places.
$BR
$BR
(a) Construct a tree diagram here to illustrate the outcomes and their
associated probabilities. Hence or otherwise find the percentage of words in
English having both the first two sounds as vowel sounds. \{ ans_rule(6) \}  
$BR
$BR
(b) A word chosen at random is found to have a consonant as its second
sound. What is the probability that its first sound was a consonant? \{ ans_rule(6) \}   
$BR
$BR

Munhall, K. G., C. Kroos, G. Jozan, and E. Vatikiotis-Bateson (2004):
Spatial frequency requirements for audiovisual speech perception. \(
Perception ~\&~ Psychophysics\) $BBOLD 66 $EBOLD, 574-583. 
$BR
$BR
Hall, Kathleen Currie, Blake Allen, Michael Fry, Scott Mackie, and Michael
McAuliffe. (2015): \(Phonological~ CorpusTools\), Version 1.1.0. [Computer
program]. Available from
https://github.com/PhonologicalCorpusTools/CorpusTools/releases. 
$BR
$BR
Vaden, K. I., H. R. Halpin & G. S. Hickok. (2009): \(Irvine~ Phonotactic~
Online~ Dictionary\), Version 2.0. [Data file.] Available from:
http://www.iphod.com.
$BR
$BR
END_TEXT
#########################################################
## Answers evaluation (Required)

$showPartialCorrectAnswers = 1;

ANS( $ans->cmp( checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
return $correct == $student || $correct+0.01 == $student ;
} ) );

ANS( num_cmp(0.475, tol=> 0.0051));

#########################################################
## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
 (a)  With C denoting consonant and V vowel,
we have
$BCENTER
\{
image ( "./14719916018848.png", width =>450, height => 450)
\}
$ECENTER
$BR
Hence, or otherwise, the percentage of words having their first two
sounds as vowels is 0.2%. 
$BR
$BR

(b) What we want to know is the P(Sound1 is a C | Sound 2 is a C),
that is, what is the probability that
sound 1 is a consonant, given that we know that sound 2 is a consonant? We
can use Bayes' Theorem, for which we need to know the probabilities that
sound 1 is a C or V, and the conditional probabilities of {Sound2 is a C}, given either that sound 1 is a C or that sound 1 is a V. We get the basic
probabilities of sound 1's being a C or a V from the text or the diagram: 
\begin{eqnarray*}
P(\text{Sound1 is a C}) &=& 0.81, \\
P(\text{Sound1 is a }V) &=& 0.19.
\end{eqnarray*}
We get the relevant conditional probabilities from reading the diagram:
\begin{eqnarray*}
P(\text{Sound2 is a C}|\text{Sound1 is a C}) &=& 0.21 \\
P(\text{Sound2 is a C}|\text{Sound1 is a V}) &=& 0.99.
\end{eqnarray*}
Then we use Bayes' Theorem to find P(Sound1 is a C} | Sound2 is a C) as follows

\begin{align*}
& \frac{P(\text{Sound2 is a C} |\text{Sound1 is a C} )\times P(
\text{Sound1 is a C})}{P(\text{ Sound2 is a C} |\text{Sound1
is a C})\times P(\text{Sound1 is a C})+P(\text{Sound2 is a C} |\text{Sound1 is a V})\times P(\text{Sound1 is a V})} \\
&= \frac{0.21\times
0.81}{0.21\times 0.81+0.99\times 0.19} \\
&= 0.4749.
\end{align*}

Hence there is an 47.5% chance that the first sound was a consonant if
we know that the second sound is a consonant.



END_SOLUTION
ENDDOCUMENT();
