##DESCRIPTION
##  College Algebra Review: Substitution
##ENDDESCRIPTION

##KEYWORDS('algebra', 'factoring', 'quadratic')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Expanding and Factoring')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.4')
## Problem1('')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PeriodicRerandomization.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

#######################################
#  Set-up

Context("Numeric");
Context()->variables->are(x=>"Real",t=>"Real");

$a = list_random(-11,-7,-5,-3,-2,2,3,5,7,11);
do {$b = list_random(-11,-7,-5,-3,-2,2,3,5,7,11)} until (gcd($a,$b)==1 && $b != -$a);
$e2 = $a + $b;
$e3 = $a*$b;

$fac1 = Formula(" x^2 + $a ")->reduce;
$fac2 = Formula(" x^2+ $b ")->reduce;
$texanswer = "\(\left(x^2 + $a\right)\left(x^2 + $b \right) \)";
$expression = "\( x^4 + $e2 x^2 + $e3  \)";
if ($e2 == 1) {
  $expression = "\( x^4 + x^2 + $e3  \)";
} elsif ($e3 == -1) {
  $expression = "\( x^4 - x^2 + $e3  \)";
}

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);
################################################
#  Main text
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your 
answer as much as possible:
$PAR
${BCENTER}
$expression \(= \) 
\( \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \big( \)
\{ $multians->ans_rule(20) \}
\( \big) \)
${ECENTER}
END_TEXT


################################################
#  Answer evaluation


$texanswer = "\(\left($fac1\right)\left($fac2\right)\)";
if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    $texanswer
    $PAR
    END_TEXT
} 

Context()->normalStrings;

PeriodicStatus(); 


################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS($multians->cmp());

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
