# DESCRIPTION
# Gives a rational expression to simplify
# ENDDESCRIPTION

## DBsubject('Intermediate Algebra')
## DBchapter('Ch 06: Rational Expressions, Functions and Equations')
## DBsection('Rational Expressions and Functions: Multipying and Dividing')
## KEYWORDS('rational','expressions')
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","z","s","r","t","a","b","c");
$ch = random(0,8,1);
$v = $vars[$ch];
Context()->variables->are($v=>"Real");

$c = random(1,7,1);
do {$d = non_zero_random(-7,7,1)} until (abs($d) != abs($c));
$k = random(2,7,1);
$a = $k*$c;
$b = $k*$d;

do {
  $q = non_zero_random(-4,4);
  $p = random(1,3);
  $e1 = $c*$p;
  $e2 = $c*$q + $d*$p;
  $e3 = $d*$q;
} until (gcd($p,$q) == 1 && $e2 != 0);

$num = "$e1 $v^2";
if ($e2 == -1) {
  $num = $num."-$v + $e3";
} elsif ($e2 == 1) {
  $num = $num."+$v + $e3";
} else {
  $num = $num."+$e2 $v + $e3";
}

$den = "$a $v^3 + $b $v^2";
if ($c == -1) {
  $den = $den."-$v + $d";
} elsif ($c == 1) {
  $den = $den."+$v + $d";
} else {
  $den = $den."+$c $v + $d";
}
$disp = "\frac{".$num."}{".$den."}";

$numans = Formula("$p $v + $q")->reduce;
$denans = Formula("$k $v^2 + 1")->reduce;

$multians = MultiAnswer($numans, $denans)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"type an answer first");
          $self->setMessage(2,"type an answer first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

$dispfrac = ColumnTable(
  "\( \displaystyle $disp = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Simplify. Assume all cancelled factors are not zero.
$PAR
$BCENTER
$dispfrac
$ECENTER
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($multians->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle \frac{$numans}{$denans}\)
  $PAR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
