# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('chain rule', 'multivariable', 'functions', 'calculus', 'derivatives ')
## Tagged by glr 01/11/10


## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('The Chain Rule')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('14.6')
## Problem1('14')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')
## HHSection1('The Chain Rule')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( y=>'Real', u=>'Real', v=>'Real' );
$showPartialCorrectAnswers = 1;

$whichz = random(0,1,1);
$a = random(1,5,1);
$whichu = random(0,1,1);
if ( $whichz ) {
    $zfunc = Compute( "cos(x^2 + y^2)" );
} else {
    $zfunc = Compute( "sin(x^2 + y^2)" );
}
if ( $whichu ) {
    $xfunc = Compute( "u*cos(v)" );
    $yfunc = Compute( "u*sin(v)" );

    if ( $whichz ) {
        $zsub = Compute( "cos(($xfunc)^2 + ($yfunc)^2)" );
	$zsimp = Compute( "cos(u^2)" );

        $dzdu = Compute( "-sin(u^2)*2*u" );
        $dzdv = Compute( "0" );
    } else {
        $zsub = Compute( "sin(($xfunc)^2 + ($yfunc)^2)" );
	$zsimp = Compute( "sin(u^2)" );

        $dzdu = Compute( "cos(u^2)*2*u" );
        $dzdv = Compute( "0" );
    }
} else {
    $xfunc = Compute( "v*cos(u)" );
    $yfunc = Compute( "v*sin(u)" );

    if ( $whichz ) {
        $zsub = Compute( "cos(($xfunc)^2 + ($yfunc)^2)" );
	$zsimp = Compute( "cos(v^2)" );

        $dzdu = Compute( "0" );
        $dzdv = Compute( "-sin(v^2)*2*v" );
    } else {
        $zsub = Compute( "sin(($xfunc)^2 + ($yfunc)^2)" );
	$zsimp = Compute( "sin(v^2)" );

        $dzdu = Compute( "0" );
        $dzdv = Compute( "cos(v^2)*2*v" );
    }
}

$dzdx = $zfunc->D('x');
$dzdy = $zfunc->D('y');
$dxdu = $xfunc->D('u');
$dxdv = $xfunc->D('v');
$dydu = $yfunc->D('u');
$dydv = $yfunc->D('v');

$dzdxsub = $dzdx->substitute( x=>$xfunc, y=>$yfunc );
$dzdusub = Compute( "($dzdxsub)*($dxdu)" );
$dzdysub = $dzdy->substitute( x=>$xfunc, y=>$yfunc );
$dzdvsub = Compute( "($dzdysub)*($dydv)" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

If
\[ z = $zfunc,\qquad x = $xfunc,\qquad y = $yfunc, \]
find \( \partial z/\partial u \) and \( \partial z/\partial v \).
The variables are restricted to domains on which
the functions are defined.
$PAR
\( \partial z/\partial u = \) \{ ans_rule(55) \}
$PAR
\( \partial z/\partial v = \) \{ ans_rule(55) \}

END_TEXT
Context()->normalStrings;

ANS( $dzdu->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( $correct == $student ) {
        return 1;
    } else {
        return ( Compute("($dzdx)*($dxdu) + ($dzdy)*($dydu)") == $student ) ?
	    n1 : 0;
    } } ) );
ANS( $dzdv->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    if ( $correct == $student ) {
        return 1;
    } else {
        return ( Compute("($dzdx)*($dxdv) + ($dzdy)*($dydv)") == $student ) ?
	    1 : 0;
    } } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Since \(z\) is a function of two variables \(x\) and \(y\) which are
functions of two variables \(u\) and \(v\), the two chain rule identities
which apply are:
\[
\frac{\partial z}{\partial u}  = 
       \frac{\partial z}{\partial x}\frac{\partial x}{\partial u}+
       \frac{\partial z}{\partial y}\frac{\partial y}{\partial u}
\]
and
\[
\frac{\partial z}{\partial v} = 
       \frac{\partial z}{\partial x}\frac{\partial x}{\partial v}+
       \frac{\partial z}{\partial y}\frac{\partial y}{\partial v} \]

This problem is most easily solved by substitution:
\[
z = $zsub = $zsimp
\]
so
\[ \frac{\partial z}{\partial u} = $dzdu, \]
and
\[ \frac{\partial z}{\partial v} = $dzdv. \]

This problem can also be solved using the chain rule  but
it is  more difficult.



END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
