##DESCRIPTION
##  
##ENDDESCRIPTION

##KEYWORDS('algebra', 'inequality', 'fraction')

## DBsubject('Algebra')
## DBchapter('Operations on polynomial and rational expressions')
## DBsection('Polynomials: divide')
## Date('6/23/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "PGgraders.pl",
   "niceTables.pl",
   "scaffold.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

# Give full credit if the student is able to complete the
# quotient without doing the scaffold steps...
install_problem_grader(~~&full_partial_grader);

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');

$q[0] = non_zero_random(-10,10,1);
$q[1] = non_zero_random(-7,7,1);
$q[2] = random(-1,1,2)*random(2,5,1);

$quot = Formula("$q[2] x^2 + $q[1] x + $q[0]")->reduce;

# divisor: ($cDiv x - $rDiv)
$cDiv = random(-5,-2,1);
# non-negative coefficient for divisor (problem will still work if negative)
do { $rDiv = random(-1,1,2)*random(2,8,1); } until (gcf(abs($rDiv),abs($cDiv))==1);
$divisor = Formula("$cDiv x - $rDiv")->reduce;
$divisorTeX = $divisor->TeX;

# set up the dividend (in reverse degree order #unimporant)
$int0 = -1*$rDiv*$q[0];
$rem0 = 0; # non_zero_random(-10,10,1); if you want a remainder...
$d[0] = $int0 + $rem0;
$int1 = -1*$rDiv*$q[1]; # for showing work (from q*divisor, to be subtracted)
$rem1 = $cDiv*$q[0]; # for showing work (remaining after subtraction)
$d[1] = $int1 + $rem1;
$int2 = -1*$rDiv*$q[2]; # for showing work (from q*divisor, to be subtracted) 
$rem2 = $cDiv*$q[1]; # for showing work (remaining after subtraction)
$d[2] = $int2 + $rem2;
$d[3] = $cDiv*$q[2];

$dividend = Formula("$d[3] x^3 + $d[2] x^2 + $d[1] x + $d[0]")->reduce;

##############################################################
#
#  Text
#
#

BEGIN_PGML

>> ### Compute the quotient of [` \left([$dividend]\right) \div \left([$divisor]\right) `] ### <<

END_PGML

Scaffold::Begin();
Section::Begin("Long Division: Eliminating the first term");
$part1Ans = MultiAnswer(Formula("$q[2] x^2"),Formula("$d[3] x^3"),Formula("$int2 x^2"),Formula("$rem2 x^2"))->with(singleResult=>1, checker=>sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q2stu, $int3stu, $int2stu, $q1stu ) = @{$student};
      my ( $q2, $int3p, $int2p, $q1 ) = @{$correct};
      $int3stuTeX = $int3stu->TeX;
      # make sure the student's quotient * $divisor = intermediate binomial
      $self->setMessage(2,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q2stu + Formula("$rDiv * $q2stu") != $int3stu);
      $self->setMessage(3,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q2stu - Formula("$cDiv x * $q2stu") != $int2stu);
      # because the product with the divisor is correct, if there's an error
      # it's in the student's selection for the quotient
      $self->setMessage(1,"\( $d[3] x^3 - $int3stuTeX \) will not be zero. <BR> You may need to try a different quotient.") if ($int3stu != $int3p && $q2stu != $q2);
      # if the remainder is incorrect, by this point, it's because they
      # subtracted incorrectly...
      $self->setMessage(4,"You have subtracted incorrectly.") if ($q1stu != Formula("$d[2] x^2 - $int2stu"));
      return [$q2==$q2stu,$int3p==$int3stu,$int2p==$int2stu,$q1==$q1stu];
});

BEGIN_PGML
[@
LayoutTable(
    [
        ['','',[ $part1Ans->ans_rule(5), noencase=>1],'',''],
        ["$divisorTeX",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',['\\( \\large - ( \\)'.$part1Ans->ans_rule(5),noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;'],[ $part1Ans->ans_rule(5).'\\( \\large ) \\)',noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;'],'',''],
        ['','0',[ $part1Ans->ans_rule(5),noencase=>1],'',''],
    ],
    midrules=>0,
    align=>'ccccc',
    encase=>['\(','\)'],
);
@]***
END_PGML

ANS($part1Ans->cmp);

Section::End();

Section::Begin("Long Division: Eliminating the second term");

$part2Ans = MultiAnswer(Formula("$q[1] x"),Formula("$rem2 x^2"),Formula("$int1 x"),Formula("$rem1 x"))->with(singleResult=>1, checker=>sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q1stu, $int2stu, $int1stu, $q0stu ) = @{$student};
      my ( $q1, $int2p, $int1p, $q0 ) = @{$correct};
      $int2stuTeX = $int2stu->TeX;
      # make sure the student's quotient * $divisor = intermediate binomial
      $self->setMessage(2,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q1stu + Formula("$rDiv * $q1stu") != $int2stu);
      $self->setMessage(3,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q1stu - Formula("$cDiv x * $q1stu") != $int1stu);
      # because the product with the divisor is correct, if there's an error
      # it's in the student's selection for the quotient
      $self->setMessage(1,"\( $q[1] x^2 - $int2stuTeX \) will not be zero. <BR> You may need to try a different quotient.") if ($int2stu != $int2p && $q1stu != $q1);
      # if the remainder is incorrect, by this point, it's because they
      # subtracted incorrectly...
      $self->setMessage(4,"You have subtracted incorrectly.") if ($q0stu != Formula("$d[1] x - $int1stu"));
      return [$q1==$q1stu,$int2p==$int2stu,$int1p==$int1stu,$q0==$q0stu];
});

BEGIN_PGML
[@
LayoutTable(
    [
        ['','',"$q[2] x^2",[ $part2Ans->ans_rule(5), noencase=>1],''],
        ["$divisorTeX",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow',''],
        ['','0',"$rem2 x^2","+$d[1] x",''],
        ['','',['\\( \\large - ( \\)'.$part2Ans->ans_rule(5),noencase=>1,cellcss=>'color:red; border-bottom:2px solid black;'],[ $part2Ans->ans_rule(5).'\\( \\large ) \\)', noencase=>1,cellcss=>'color:red; border-bottom:2px solid black;'],''],
        ['','','0',[ $part2Ans->ans_rule(5),noencase=>1],''],
    ],
    midrules=>0,
    align=>'ccccc',
    encase=>['\(','\)'],
);
@]***
END_PGML

ANS($part2Ans->cmp);

Section::End();

Section::Begin("Long Division: Eliminating the third term");

$part3Ans = MultiAnswer(Formula("$q[0]"),Formula("$rem1 x"), Formula("$int0"),Formula("$rem0"))->with(singleResult=>1, checker=>sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q0stu, $int1stu, $int0stu, $r0stu ) = @{$student};
      my ( $q0, $int1p, $int0p, $r0 ) = @{$correct};
      $int1stuTeX = $int1stu->TeX;
      # make sure the student's quotient * $divisor = intermediate binomial
      $self->setMessage(2,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q0stu + Formula("$rDiv * $q0stu") != $int1stu);
      $self->setMessage(3,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q0stu - Formula("$cDiv x * $q0stu") != $int0stu);
      # because the product with the divisor is correct, if there's an error
      # it's in the student's selection for the quotient
      $self->setMessage(1,"\( $q[0] x - $int1stuTeX \) will not be zero. <BR> You may need to try a different quotient.") if ($int1stu != $int1p && $q0stu != $q0);
      # if the remainder is incorrect, by this point, it's because they
      # subtracted incorrectly...
      $self->setMessage(4,"You have subtracted incorrectly.") if ($r0stu != Formula("$d[0] - $int0stu"));
      return [$q0==$q0stu,$int1p==$int1stu,$int0p==$int0stu,$r0==$r0stu];
});

BEGIN_PGML
[@
LayoutTable(
    [
        ['','',"$q[2] x^2","+ $q[1] x",[ $part3Ans->ans_rule(5),noencase=>1]],
        ["$divisorTeX",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow','\downarrow'],
        ['','0',"$rem2 x^2","+$d[1] x",'\downarrow'],
        ['','',["- ( $rem2 x^2",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int1 x )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow'],
        ['','','0',"$rem1 x","+$d[0]"],
        ['','','',['\\( \\large - ( \\)'.$part3Ans->ans_rule(5), noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;'],[ $part3Ans->ans_rule(5).'\\( \\large ) \\)', noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;']],
        ['','','','0',[ $part3Ans->ans_rule(5),noencase=>1]]
    ],
    midrules=>0,
    align=>'ccccc',
    encase=>['\(','\)'],
);
@]***

END_PGML

ANS($part3Ans->cmp);

Section::End();

Scaffold::End();

$finalAns = MultiAnswer($quot, Formula($rem0))->with(
    checkTypes=>0,
    singleResult=>1,
    separator=>' with remainder ',
    tex_separator=>'\text{ with remainder }',
    checker=> sub{
        my ($correct, $student, $self) = @_;
        my ($quotStu, $remStu) = @{$student};
        my ($quotCor, $remCor) = @{$correct};
        $self->setMessage(1,"Your quotient is incorrect.") if ($quotStu != $quotCor);
        $self->setMessage(2,"Your remainder is incorrect.") if ($remStu != $remCor);
        $self->setMessage(2,"Your remainder should be a constant (no variables).") unless (Value::isReal($remStu));
        return [ $quotCor==$quotStu,$remCor==$remStu ];
    }
);

BEGIN_PGML

The steps above are provided to help you compute the quotient and remainder.

If you can find the quotient and remainder without using the steps above, you may input your answer here for full credit:

The quotient of [` \left([$dividend]\right) \div \left([$divisor]\right) `] is [_______________]{$finalAns} with remainder [_____]{$finalAns}

END_PGML

##############################################################
#
#  Answers
#
#

##############################################################
#
#  Hints
#
#



BEGIN_PGML_HINT

*Eliminating the first term*

Divide the highest term of the dividend, [`[$dividend]`], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$d[3]] x^3`] by [`[$cDiv]x`]. The result is the highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. Enter the answer in the two middle boxes.  Next, subtract this product from [`[$d[3]] x^3 + [$d[2]] x^2`]. Enter the answer in the last box. The first term, [`[$d[3]] x^3`], must be eliminated at the end of this process.

*Eliminating the second term*

Divide the highest term of the polynomial you found in the last subtraction by the highest term of the divisor, [`[$divisor]`]. The result is the second highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. Enter the answer in the two middle boxes.  Next, subtract this product from [`[$rem2] x^2+[$d[1]] x`]. Enter the answer in the last box. The second term, [`[$rem2] x^2`], must be eliminated at the end of this process.

*Eliminating the third term*

Divide the highest term of the polynomial you found in the last subtraction  by the highest term of the divisor, [`[$divisor]`].   The result is the last term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. Enter the answer in the two middle boxes.  Next, subtract this product from [`[$rem1] x+[$d[0]] `]. Enter the answer, which is the remainder of the division, in the last box. The third term, [`[$rem1] x`], must be eliminated at the end of this process. 

END_PGML_HINT

##############################################################
#
#  Solutions
#
#




#$divisor = Formula("$cDiv x - $rDiv") 

BEGIN_PGML_SOLUTION
    
*Eliminating the first term*

Divide the highest term of the dividend, [`[$dividend]`], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$d[3]] x^3`] by [`[$cDiv] x`]. 

>>[`\dfrac{[$d[3]] x^3}{[$cDiv] x} = [$q[2]] x^2`]<<

The result is the highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. 

>>[`[$q[2]] x^2 ([$divisor]) = [$d[3]] x^3 + [$int2] x^2`]<<

Enter the answer in the two middle boxes.  Next, subtract this product from [`[$d[3]] x^3 + [$d[2]] x^2`]. Enter the answer in the last box. 

>>[`[$d[3]] x^3 + [$d[2]] x^2-([$d[3]] x^3 + [$int2] x^2) = [$rem2] x^2 `]<<

Note that the first (highest) term of the dividend, [`[$d[3]] x^3`], was eliminated in  this process.

[@
LayoutTable(
    [
        ['','',"$q[2] x^2",'',''],
        ["$divisorTeX",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'',''],
        ['','0',"$rem2 x^2","",''],
    ],
    midrules=>0,
    align=>'ccccc',
    encase=>['\(','\)'],
)@]***



*Eliminating the second term*

Divide the highest term of the polynomial you found in the last subtraction, [`[$rem2] x^2`], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$rem2] x^2`] by [`[$cDiv]x`]. 

>>[`\dfrac{[$rem2] x^2}{[$cDiv]x} = [$q[1]] x`]<<

The result is the second highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. 

>>[`[$q[1]] x ([$divisor]) = [$rem2] x^2+[$int1]x`]<<

Enter the answer in the two middle boxes.  Next, subtract this product from [`[$rem2] x^2+[$d[1]] x`]. Enter the answer in the last box. 

>>[`[$rem2] x^2+[$d[1]] x-([$rem2] x^2+[$int1]x) = [$rem1] x `]<<

Note that the second term, [`[$rem2] x^2`], was eliminated in this process.

[@
LayoutTable(
    [
        ['','',"$q[2] x^2","+ $q[1] x",''],
        ["$divisorTeX",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow',''],
        ['','0',"$rem2 x^2","+$d[1] x",''],
        ['','',["- ( $rem2 x^2",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int1 x )",cellcss=>'color:red; border-bottom:2px solid black;'],''],
        ['','','0',"$rem1 x",''],
    ],
    midrules=>0,
    align=>'ccccc',
    encase=>['\(','\)'],
)@]***



*Eliminating the third term*

Divide the highest term of the polynomial you found in the last subtraction, [`[$rem1] x `], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$rem1] x`] by [`[$cDiv]x`]. 

>>[`\dfrac{[$rem1] x}{[$cDiv] x} = [$q[0]]`]<<

The result is the last term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. 

>>[`[$q[0]] ([$divisor]) = [$rem1] x + [$d[0]] `]<<

Enter the answer in the two middle boxes.  Next, subtract this product from [`[$rem1] x+[$d[0]] `]. Enter the answer, which is the remainder of the division, in the last box. 

>>[` [$rem1] x + [$d[0]]-([$rem1] x + [$d[0]]) = 0 `]<<

Note that the third term, [`[$rem1] x`], was eliminated in this process. 



[@
LayoutTable(
    [
        ['','',"$q[2] x^2","+ $q[1] x","+ $q[0]"],
        ["$divisorTeX",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow','\downarrow'],
        ['','0',"$rem2 x^2","+$d[1] x",'\downarrow'],
        ['','',["- ( $rem2 x^2",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int1 x )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow'],
        ['','','0',"$rem1 x","+$d[0]"],
        ['','','',["- ( $rem1 x",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $d[0] )",cellcss=>'color:red; border-bottom:2px solid black;']],
        ['','','','0','0']
    ],
    midrules=>0,
    align=>'ccccc',
    encase=>['\(','\)'],
)@]***


The quotient of [` \left([$dividend]\right) \div \left([$divisor]\right) `] is [`[$quot]`] with remainder [`0`].

END_PGML_SOLUTION





ENDDOCUMENT();        
