## DESCRIPTION
## Simplify a rational expression
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################


Context("RationalFunction");
$var = list_random('y', 'z', 'a', 'b', 'c');
Context()->variables->add($var =>'Real');

#common factors is (x + a)
$a = non_zero_random(2, 9);
do{
$b = non_zero_random(-9, 9);
}until($b != $a);


#Question:  (x^2 - a^2)/((x - a)(x + b))
$g = Formula("($var^2 - $a**2)/($var^2 + ($b - $a)*$var - $b*$a)")->reduce;
#answer
$f = Formula("($var + $a)/($var + $b)")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the rational expression.

[``[$g] = ``]  [________]{$cmp}
 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
We can only cancel common factors, so you must factor before you can simplify.

Both the numerators and denominators have GCFS

[`\frac{([$var] + [$a])([$var] - [$a])}{([$var] - [$a])([$var] + [$b])}`]

[`([$var] - [$a])`] is a common factor.  If we reduce this we are left with [`[$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();




