## DESCRIPTION
## Simplify a rational expression
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################

#Product rule for exponents, negative exponent and a positive power of 1 (student must enter in the denominator)

Context("RationalFunction");
$var = list_random('p', 'q', 'r', 't', 'w', 'v');
Context()->variables->add($var =>'Real');

#common factors is: (x + a)
$a = non_zero_random(-9, 9);
$b = non_zero_random(-9, 9);
do{
$c = non_zero_random(-9, 9);
}until($b != $c);

#Expression is ((x+a)(x+b))/((x+a)(x+c))
#Question
$g = Formula("($var^2 + ($a+$b)*$var + $a*$b)/($var^2 + ($a+$c)*$var +$a*$c)")->reduce;

#answer
$f = Formula("($var + $b)/($var + $c)")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the rational expression.

[``[$g]=``]  [________]{$cmp}

 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
We can only cancel common factors, so you must factor before you can simplify.


[`\frac{([$var] + [$a])([$var] + [$b])}{([$var] + [$a])([$var] + [$c])}`]

[`([$var] + [$a])`] is a common factor.  If we reduce this we are left with [`[$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();




