## DESCRIPTION
## Divide Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
Context()->variables->add(y =>'Real');

#constants
$a = random(2, 9);
$b = random(2, 9);

#exponents
$n = random(2, 6);
$m1 = random(1, 6);
$m2 = random(2, 6);

#Question
$g1 = Formula("($a*$b*x^$n)/($b*x*y^$m1)");
$g2 = Formula("$a**2*y^$m2");

#FOR SOLUTION
$gflip = Formula("1/($a**2*y^$m2)");
$gcf = $a*$b;
$mult = Formula("(x^$n)/($a*x*y^($m1 + $m2))")->reduce;

#answer
$f = Formula("(x^($n-1))/($a*y^($m1+$m2))");

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Multiply and simplify. Use only positive exponents in your answer.

[``[$g1]\div [$g2] = ``]  [________]{$cmp}

 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
To divide, multiply by the reciprocal.

[`\displaystyle [$g1] \cdot [$gflip]`]

Look for common factors before multiplying. We can reduce by both [$a] and [$b] to get:

[`\displaystyle [$mult] `]

Finally, we can reduce to [`\displaystyle [$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();