## DESCRIPTION
## Multiply Rational Expressions
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################
Context("RationalFunction");
$var = list_random('a','b','c','d','m','n','p','r','s','t','w','y');
Context()->variables->are($var =>'Real');

#Problem is of the form:
#(ax(x+b))/(x+c)^2) TIMES (x+c)/(a*d*x)
#
$a = random(2,10);
$b = non_zero_random(-8, 8);
do{
$c = non_zero_random(-8, 8);
}until($b != $c);
$d = random(2, 10);

#Question
$g1 = Formula("($a*$var^2 + $b*$a*$var)/($var^2 + 2*$c*$var + $c*$c)")->reduce;
$g2 = Formula("($var + $c)/($a*$d*$var)")->reduce;

$g1new = Formula("($a*$var($var + $b))/(($var + $c)($var + $c))")->reduce;
$g2new = Formula("($var + $c)/($a*$d*$var)")->reduce;

$gcf = Formula("$a*$var*($var + $c)")->reduce;


#answer
$f = Formula("($var + $b)/($d*($var + $c))")->reduce;

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless $cnum == $snum;
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Multiply and simplify.

[``[$g1]\cdot[$g2] = ``]  [________]{$cmp}

 
END_PGML

##############################################

BEGIN_PGML_SOLUTION
Factor both expressions before multiplying.

In factored form, the problem is now:

[`\displaystyle [$g1new] \cdot [$g2new]`]

[`[$gcf]`] is a common factor.  If we reduce this and multiply straight across, we are left with [`\displaystyle [$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();