## DESCRIPTION
## Simplify a rational expression
## ENDDESCRIPTION


## DBsubject(Elementary Algebra)
## DBchapter()
## DBsection()
## Date(05/24/2019)
## Institution(Red Rocks Community College)
## Author(Martha Stevens)
## MO(1)
## KEYWORDS('elementary algebra')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "PGgraphmacros.pl",
  "PGcourse.pl",
  "contextRationalFunction.pl",
 
);

##############################################


Context("RationalFunction");
$var = list_random('x', 'y', 'z', 'a', 'b', 'c', 't', 'm', 'n', 'p');
Context()->variables->are($var =>'Real');

#problem is of the form:
# [-a(x - b)]/[(x + b)(x - b)]
$a = random(3, 8);
$b = random(2, 9);
$gcf1 = -$a;

#Question
$g = Formula("($b*$a - $a*$var)/($var^2 - $b**2)")->reduce;

#Answer
#The negative must go in numerator!
$f = Formula("(-$a)/($var + $b)");

#custom checker to make sure it is reduced.  It checks to see that answers are equivalent and then checks to see that the numerators are the same.  Custom checker doesn't enforce polynomial format though... 
$cmp = $f->cmp(checker => sub {
  my ($correct,$student,$ans) = @_;
  return 0 unless $correct == $student;
  return 0 unless $student->{tree}->class eq "BOP" && $student->{tree}{bop} eq "/";
  my $cnum = Formula($correct->{tree}{lop});
  my $snum = Formula($student->{tree}{lop});
Value->Error("Your answer isn't reduced") unless (($cnum == $snum)||($cnum == -$snum));
  return 1;
});


##############################################

TEXT(beginproblem());

BEGIN_PGML

Simplify the rational expression.

[``[$g] = ``]  [________]{$cmp}
 
END_PGML

##############################################



BEGIN_PGML_SOLUTION
We can only cancel common factors, so you must factor before you can simplify.

To factor the numerator, take out the GCF of [`[$gcf1]`].

[`[$gcf1]([$var] - [$b])`]

The denominator is a difference of perfect squares and can be factored as:

[`([$var] - [$b])([$var] + [$b])`]

In factored form, the problem is now:

[`\displaystyle \frac{[$gcf1]([$var] - [$b])}{([$var] - [$b])([$var] + [$b])}`]

[`([$var] - [$b])`] is a common factor.  If we reduce this we are left with [`\displaystyle [$f]`].
END_PGML_SOLUTION

ENDDOCUMENT();