## DESCRIPTION
## Find a matrix E such that E*A is the RREF of the given 3x2 matrix A.
##   The matrix and its RREF both have small integer entries.
##   There may be infinitely many different possible solutions to this problem. All of them are accepted as correct.
## ENDDESCRIPTION

## DBsubject(Linear algebra)
## DBchapter(Matrices)
## DBsection(Echelon form)
## Date(February 26, 2020)
## Institution(Mount Allison University)
## Author(Nathaniel Johnston)
## MO(1)
## TitleText1('Introduction to Linear and Matrix Algebra')
## AuthorText1('Nathaniel Johnston')
## EditionText1('1')
## Section1('2.2')
## Problem1('')
## KEYWORDS('matrix','reduced row echelon form')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "contextFraction.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Matrix");

$flg = random(1,3);

if($flg == 1){
    do {
        $c1 = random(-2,3);
        $c2 = random(-2,3);
        $c3 = random(-2,3);
        $c4 = random(-2,3);
        $c5 = random(-2,3);
        $c6 = random(-2,3);
        $c7 = random(-2,3);
        $c8 = random(-2,3);
        $c9 = random(-2,3);
        $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
        
        $R = Matrix([[1,0],[0,1],[0,0]]);
        
        $A = $C * $R;
    } until (abs($C->det) > 0.1);
}elsif($flg == 2){
    do {
        $c1 = random(-2,3);
        $c2 = random(-2,3);
        $c3 = random(-2,3);
        $c4 = random(-2,3);
        $c5 = random(-2,3);
        $c6 = random(-2,3);
        $c7 = random(-2,3);
        $c8 = random(-2,3);
        $c9 = random(-2,3);
        $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
        
        $r1 = random(-2,4);
        $R = Matrix([[1,$r1],[0,0],[0,0]]);
        
        $A = $C * $R;
    } until (abs($C->det) > 0.1);
}else{
    do {
        $c1 = random(-2,3);
        $c2 = random(-2,3);
        $c3 = random(-2,3);
        $c4 = random(-2,3);
        $c5 = random(-2,3);
        $c6 = random(-2,3);
        $c7 = random(-2,3);
        $c8 = random(-2,3);
        $c9 = random(-2,3);
        $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
        
        $R = Matrix([[0,1],[0,0],[0,0]]);
        
        $A = $C * $R;
    } until (abs($C->det) > 0.1);
}

$ans = $C->inverse;
   
##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Find a matrix \( \displaystyle E \) with the property that \( \displaystyle EA \) equals the reduced row echelon form of \( \displaystyle A = $A \).
$BR
$BR
\( \displaystyle E = \) \{ $ans->ans_array() \}
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $ans->cmp( checker => sub {
    my ( $correct, $student, $answerHash ) = @_;
    
    if( $student*$A == $R )
        { return 1; } else { return 0; }
}));

ENDDOCUMENT();
