##DESCRIPTION
##  College Algebra Review: Computing Absolute Values
##ENDDESCRIPTION

##KEYWORDS('algebra', 'absolute value')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Real Numbers')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.2')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

###################################
# Setup

Context("Numeric");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(-12,-2,1);
$b = random(2,12,1);
$c = random(-12,-2,1);

Context("Numeric")->variables->add(y=>"Real",z=>"Real");

$frac = " \frac{x^{$a} y^{$b}}{z^{$c}}";

#Context(LimitedPowers::OnlyPositiveIntegers());
$num = Formula("y^{$b}z^{-$c}")->reduce;
$den = Formula("x^{-$a}")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
#      } elsif ( ($numbogus==$f1stu ) ||
#                ($denbogus==$f2stu ) ) {
#          $self->setMessage(1,"Find a common denominator first");
#          $self->setMessage(2,"Find a common denominator first");
#          return [0,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER
$PAR
$BR

END_TEXT
Context()->normalStrings;

###################################
# Answers

COMMENT('TRASH! REPEAT');
ANS($multians->cmp() );

ENDDOCUMENT();
