##DESCRIPTION
##  College Algebra Review: Computing Absolute Values
##ENDDESCRIPTION

##KEYWORDS('algebra', 'absolute value')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Real Numbers')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.2')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "problemRandomize.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$guest   = ($effectivePermissionLevel == -5) ;	####	practice user
$observe = ($effectivePermissionLevel >=  5) ;	####	TA or Professor (permission 10)
ProblemRandomize( when => "Always", onlyAfterDue => 0, style => "Button" )
 	if	($guest or $observe) ;
$when = (time >= $main::dueDate ? "Always" : "Correct");
ProblemRandomize( when => $when, onlyAfterDue => 0, style => "Button" )
 	if	not ($guest or $observe) ;



Context("Numeric");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$m = random(-7,-3,2);
$n = random(2,20,2);
$a = random(2,5,1);
$b = random(1,5,1);
$c = $b*$a;

$frac = " \frac{$a \cdot x^{$n}}{$c \cdot x^{$m}}";

$den = $b;
$num = Formula("x^{$n-$m}")->reduce;

$multians = MultiAnswer($num, $den)
->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      #my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
#      } elsif ( ($numbogus==$f1stu ) ||
#                ($denbogus==$f2stu ) ) {
#          $self->setMessage(1,"Find a common denominator first");
#          $self->setMessage(2,"Find a common denominator first");
#          return [0,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

$problemRandomize = 3;

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER

END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 0;

ANS($multians->cmp() );



COMMENT('TRASH! REPEAT');
ENDDOCUMENT();
