##DESCRIPTION
##  College Algebra Review: Computing Absolute Values
##ENDDESCRIPTION

##KEYWORDS('algebra', 'absolute value')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Real Numbers')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.2')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");


###################################
# Setup

Context("Numeric");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(-12,-2,1);
$b = random(2,4,1);
$c = $b*$a;
$exp = random(-3,-2,1);
$m = random(-5,-2,1);
$n = random(6,10,1);
$k = random(-3,-7,1);

Context("Numeric")->variables->add(y=>"Real",z=>"Real");

$frac = " \left( \frac{$c x^{$m}y^{$n}}{$a z^{$k}} \right)^{$exp}";

#Context(LimitedPowers::OnlyPositiveIntegers());
$num = Formula("x^{$m*$exp}")->reduce;
$den = Formula("($b)^{-$exp} y^{-$n*$exp}z^{$k*$exp}")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
#      } elsif ( ($numbogus==$f1stu ) ||
#                ($denbogus==$f2stu ) ) {
#          $self->setMessage(1,"Find a common denominator first");
#          $self->setMessage(2,"Find a common denominator first");
#          return [0,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Simplify.  Write your answer using only positive exponents.
$PAR
$BCENTER
$displayfrac
$ECENTER

END_TEXT

$ansnumexp = $m*$exp;
$ansdenconst = ($b)**(-$exp);
$ansdenexp1 = -$n*$exp;
$ansdenexp2 = $k*$exp;

if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    \( \dfrac{x^{$ansnumexp}}{$ansdenconst y^{$ansdenexp1} z^{$ansdenexp2}} \)
    $PAR
    END_TEXT
}

Context()->normalStrings;

PeriodicStatus(); 

###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );

COMMENT('TRASH! REPEAT');
ENDDOCUMENT();
