##DESCRIPTION
##  College Algebra Review: Computing Absolute Values
##ENDDESCRIPTION

##KEYWORDS('algebra', 'absolute value')

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Real Numbers')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('R.2')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "problemRandomize.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$guest   = ($effectivePermissionLevel == -5) ;	####	practice user
$observe = ($effectivePermissionLevel >=  5) ;	####	TA or Professor (permission 10)
ProblemRandomize( when => "Always", onlyAfterDue => 0, style => "Button" )
 	if	($guest or $observe) ;
$when = (time >= $main::dueDate ? "Always" : "Correct");
ProblemRandomize( when => $when, onlyAfterDue => 0, style => "Button" )
 	if	not ($guest or $observe) ;

###################################
# Setup

Context("Numeric");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = list_random(-12,-11,-10,-8,-7,-6,-5,-4,-3,-2,2,3,4,5,6);
$b = random(2,4,1);
$c = $b*$a;
$exp = random(-3,-2,1);
$m1 = random(-5,-2,1);
$n1 = random(6,10,1);
$k1 = random(-3,-7,1);
$m2 = random(2,5,1);
$n2 = random(-4,-6,1);
$k2 = random(-12,-9,1);

Context("Numeric")->variables->add(y=>"Real",z=>"Real");

$frac = " \left( \frac{$a x^{$m1}y^{$n1}z^{$k1}}{$c x^{$m2} y^{$n2} z^{$k2}} \right)^{$exp}";

#Context(LimitedPowers::OnlyPositiveIntegers());
$num = Formula("($b)^{-$exp} x^{$exp*($m2 - $m1)}")->reduce;
$den = Formula("y^{$exp*($n1 - $n2)}z^{$exp*($k1 - $k2)}")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
#      } elsif ( ($numbogus==$f1stu ) ||
#                ($denbogus==$f2stu ) ) {
#          $self->setMessage(1,"Find a common denominator first");
#          $self->setMessage(2,"Find a common denominator first");
#          return [0,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

$problemRandomize = 3;

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER
$PAR
(${BBOLD}Hint: $EBOLD Simplify the inside of the parentheses first.)
END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );

COMMENT('TRASH! REPEAT');
ENDDOCUMENT();
