## DESCRIPTION
## Algebraic fractions
## ENDDESCRIPTION

## KEYWORDS('Algebraic fractions')

## DBsubject('Algebra') 
## BookTitle('Algebra: Form and Function')
## DBchapter('Basic Algebra') 
## BookChapter('Rules for Expressions and the Reasons for Them')
## DBsection('Algebraic Fractions') 
## BookSection('Algebraic Fractions')
## Date('01/01/10')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Algebra: Form and Function')
## EditionText1('1')
## AuthorText1('McCallum, Connally, and Hughes-Hallett')
## Section1('2.4')
## Problem1('12')


##################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "problemRandomize.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$guest   = ($effectivePermissionLevel == -5) ;	####	practice user
$observe = ($effectivePermissionLevel >=  5) ;	####	TA or Professor (permission 10)
ProblemRandomize( when => "Always", onlyAfterDue => 0, style => "Button" )
 	if	($guest or $observe) ;
$when = (time >= $main::dueDate ? "Always" : "Correct");
ProblemRandomize( when => $when, onlyAfterDue => 0, style => "Button" )
 	if	not ($guest or $observe) ;

###################################
# Setup

Context("Numeric");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";

$a = random(-12,-2,1);
$b = random(2,12,1);
$c = random(-12,-2,1);

Context("Numeric")->variables->add(y=>"Real",z=>"Real");

$frac = " \frac{x^{$a} y^{$b}}{z^{$c}}";

#Context(LimitedPowers::OnlyPositiveIntegers());
$num = Formula("y^{$b}z^{-$c}")->reduce;
$den = Formula("x^{-$a}")->reduce;

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
#      } elsif ( ($numbogus==$f1stu ) ||
#                ($denbogus==$f2stu ) ) {
#          $self->setMessage(1,"Find a common denominator first");
#          $self->setMessage(2,"Find a common denominator first");
#          return [0,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);


# 
#  Display the fraction and answer blanks nicely
#
if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

$problemRandomize = 3;

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only:
$PAR
$BCENTER
$displayfrac
$ECENTER

END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp() );


COMMENT('TRASH! REPEAT');
ENDDOCUMENT();
