##DESCRIPTION
##  Algebra problem: combining functions
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Functions)
## DBsection(Combining Functions)
## Institution(Portland Community College)
## Author(Alex Jordan)
## MO(1)
## KEYWORDS('algebra', 'combining functions', 'domain')

########################################

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "parserAssignment.pl",
  "PGML.pl",
  "PGcourse.pl",
);

########################################

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

do {
$f = RandomVariableName(type=>'function');
$g = RandomVariableName(type=>'function');
} until ($f ne $g);

Context("Numeric");
Context()->variables->add($f=>'Real', $g=>'Real');
Context()->strings->add(
  "independent"=>{},
  "dependent"=>{},
  "input"=>{},
  "output"=>{},
  "explanatory"=>{},
  "response"=>{},
  "argument"=>{},
  "domain"=>{},
  "range"=>{},
);

$multians = MultiAnswer(Compute($f), Compute($g))->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $fstu, $gstu ) = @{$student};
      my ( $fcor, $gcor ) = @{$correct};
      if ( ($fcor == $fstu && $gcor == $gstu) ||
           ($fcor == $gstu && $gcor == $fstu) ) {
          return 1;
      } else  {return 0;}
  }
);

parser::Assignment->Function("$f");
parser::Assignment->Function("$g");

$gfun = Compute("$g(x)");

$multians2 = MultiAnswer(Compute($gfun), Formula(0))->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self, $ansHash ) = @_;
      $ansHash->{ans_preview} = $ansHash->{original_student_ans};
      my ( $gstu, $zstu ) = @{$student};
      my ( $gcor, $zcor ) = @{$correct};
      if ( ($gcor == $gstu && $zcor == $zstu) ||
           ($gcor == $zstu && $zcor == $gstu) ) {
          return 1;
      } else  {return 0;}
  }
);



########################################

Context()->texStrings;
BEGIN_TEXT

If \($f\) and \($g\) are functions, the domain of \(\frac{$f}{$g}\) is the collection of all numbers \(x\) such that \(x\) is in the \{ans_rule(10)\} of both \{$multians->ans_rule(5)\} and \{$multians->ans_rule(5)\} and for which \{$multians2->ans_rule(5)\}\(\neq\)\{$multians2->ans_rule(5)\}.

END_TEXT
Context()->normalStrings;

ANS(Compute('domain')->cmp);
ANS($multians->cmp());
ANS($multians2->cmp(correct_ans_latex_string => "$g(x);\,0"));

########################################


BEGIN_PGML_SOLUTION

If [`[$f]`] and [`[$g]`] are functions, the domain of [`\frac{[$f]}{[$g]}`] is the collection of all numbers [`x`] such that [`x`] is in the _domain_ of both [`[$f]`] and [`[$g]`] and for which [`[$g](x)\neq0`].


END_PGML_SOLUTION

ENDDOCUMENT();        # This should be the last executable line in the problem.

