##DESCRIPTION
##  Introduction to Probability: Unions, Intersections, and Complements
##ENDDESCRIPTION

##KEYWORDS('probability', 'union', 'intersection', 'complement')

## DBsubject('Probability')
## DBchapter('Discrete Probability')
## DBsection('Events/Probability Calculation')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Jason Aubrey')
## Institution('University of Missouri')

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedNumeric.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Numeric");
$p1 = Real(list_random(0.11,0.13,0.17,0.19,0.23,0.29,0.31,0.37,0.41,0.43,0.47));
$p2 = 100*$p1;
$a = random(1,3,1);
do {$b = random(4,7,1)} until (gcd($a, $b) == 1);

$ans1a = $p2;
$ans1b = 100 - $p2;
$ans2a = 100 - $p2;
$ans2b = $p2;
$ans3a = $a;
$ans3b = $b - $a;
$ans4a = $b- $a;
$ans4b = $a;

$multians1 = MultiAnswer($ans1a, $ans1b)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $f1/($f1+$f2) == $f1stu/($f1stu+$f2stu)) {
          return [1,1];
      } else {
          return [0,0];
      }
  }
);
$multians2 = MultiAnswer($ans2a, $ans2b)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $f2/($f1+$f2) == $f2stu/($f1stu+$f2stu)) {
          return [1,1];
      } else {
          return [0,0];
      }
  }
);
$multians3 = MultiAnswer($ans3a, $ans3b)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $f1/($f1+$f2) == $f1stu/($f1stu+$f2stu)) {
          return [1,1];
      } else {
          return [0,0];
      }
  }
);
$multians4 = MultiAnswer($ans4a, $ans4b)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $f2/($f1+$f2) == $f2stu/($f1stu+$f2stu)) {
          return [1,1];
      } else {
          return [0,0];
      }
  }
);
##############################################################
#
#  Text
#

Context()->texStrings;
BEGIN_TEXT
Given the following probabilities for an event \(E\), find the odds for and against \(E\):
$PAR
$SPACE $SPACE a) \(P(E) = $p1\)
$PAR
Odds for: $SPACE $SPACE $SPACE \{$multians1->ans_rule(5)\} $SPACE \(:\) $SPACE \{$multians1->ans_rule(5)\}
$PAR
Odds against: \{$multians2->ans_rule(5)\} $SPACE \(:\) $SPACE \{$multians2->ans_rule(5)\}
$PAR
$SPACE $SPACE b) \(P(E) = \dfrac{$a}{$b}\)
$PAR
Odds for: $SPACE $SPACE $SPACE \{$multians3->ans_rule(5)\} $SPACE \(:\) $SPACE \{$multians3->ans_rule(5)\}
$PAR
Odds against: \{$multians4->ans_rule(5)\} $SPACE \(:\) $SPACE \{$multians4->ans_rule(5)\}



END_TEXT
Context()->normalStrings;


##############################################################
#
#  Answers
#

ANS($multians1->cmp);
ANS($multians2->cmp);
ANS($multians3->cmp);
ANS($multians4->cmp);

COMMENT('MathObject version');
ENDDOCUMENT();
