# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

### Tagged by glr 12/16/08

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Mean value theorem)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('3.10')
## Section2('3.10')
## Section3('3.10')
## Section4('3.10')
## Problem1('15')
## Problem2('19')
## Problem3('20')
## Problem4('20')
## KEYWORDS('calculus','derivative','mean value theorem','theory')

## Textbook tags
## HHChapter1('Short-Cuts to Differentiation')
## HHChapter2('Short-Cuts to Differentiation')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add( t=>'Real' );

$showPartialCorrectAnswers = 1;

$a = random(3,9,2);

$es = list_random( "\le", "\ge" );
$ineq = ( $es eq "\le" ) ? "greater than" : "less than";

$line = Compute("$a*t");
$quad = Compute("($a/2)*t^2");

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Suppose that \(f(t)\) is continuous and twice-differentiable for
\(t\ge 0\).  Further suppose \(f''(t) $es $a\) for all \(t\ge 0\) and
\(f(0) = f'(0) = 0\).

$PAR
Using the Racetrack Principle, what linear function \(g(t)\) can we
prove is $ineq \(f'(t)\) (for \(t\ge 0\))?
$BR
\( g(t) = \) \{ ans_rule(35) \}

$PAR
Then, also using the Racetrack Principle, what quadratic function
\(h(t)\) can we prove is $ineq than \(f(t)\) (for \(t\ge 0\))?
$BR
\( h(t) = \) \{ ans_rule(35) \}

$PAR
${BITALIC}For both parts of this problem, be sure you can clearly
state how the theorem is applied to prove the indicated
inequalities.$EITALIC

END_TEXT
Context()->normalStrings;

ANS( $line->cmp( checker=>sub {
		    my ($c, $s, $ans) = @_;
		    my $sd = $s->D('t');
		    return $sd->type == "Number" &&
		        Compute(0) == $s->eval(t=>0) &&
		        Compute($a) <= Compute($sd->string);
		 } ) );
ANS( $quad->cmp( checker=>sub {
		    my ($c, $s, $ans) = @_;
		    my $sd = $s->D('t');
		    my $sdd = $sd->D('t');
		    return $sdd->type == "Number" &&
		        Compute(0) == $s->eval(t=>0) &&
		        Compute(0) == $sd->eval(t=>0) &&
		        Compute($a) <= Compute($sdd->string);
		 } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

By the Racetrack Principle, comparing \(f'(t)\) and \(g(t)=$a t\),
we have \(f''(t) \leq g'(t)\) for all \(t > 0\) and \(f'(0)=0=g(0)\),
so that \(f'(t)\leq $a t\) for all \(t\geq 0\).
$PAR
Next apply the Racetrack Principle again to
\(f(t)\) and \(h(t)=\frac{$a}{2} t^2\). Since \(f'(t)\leq h'(t)\)
for all \(t\ge 0\) and \(f(0)=0=h(0)\), then
\(f(t) \leq h(t)=\frac{$a}{2} t^2\) for all \(t\geq 0\).

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
