# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 12/15/09

## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Differentiability, linearization and tangent planes)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('14.5')
## Section2('14.5')
## Section3('14.5')
## Problem1('25')
## Problem2('25')
## Problem3('25')
## KEYWORDS('gradient', 'directional derivative', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Vector");
Context()->flags->set(ijk=>1);
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;

$a = random(1,4,1);
$b = random(1,4,1);
$c = random(1,4,1);
$eqn = Compute( "$a*x^2 - $b*y^2 + $c*z^2" )->reduce();
$x0 = random(-1,1,2);
$y0 = random(-1,1,2);
$z0 = random(1,3,1);
$e0 = $a*$x0*$x0 - $b*$y0*$y0 + $c*$z0*$z0;

$fx = Compute("2*$a*x");
$fy = Compute("-2*$b*y");
$fz = Compute("2*$c*z");
$grad = Vector( "$fx i + $fy i + $fz k" );
$fx0 = $fx->substitute( x=>$x0, y=>$y0, z=>$z0 );
$fy0 = $fy->substitute( x=>$x0, y=>$y0, z=>$z0 );
$fz0 = $fz->substitute( x=>$x0, y=>$y0, z=>$z0 );
$normal = Vector( "$fx0 i + $fy0 j + $fz0 k" );
$plane = Compute("-1*($fx0*(x - $x0) + $fy0*(y - $y0))/$fz0 + $z0")->reduce();

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Check that the point \( ($x0,$y0,$z0) \) lies on the
given surface.  Then, viewing the surface as a level
surface for a function \( f(x,y,z) \), find a vector
normal to the surface and an equation for the tangent
plane to the surface at \( ($x0,$y0,$z0) \).

\[ $eqn = $e0 \]

$PAR
vector normal = \{ ans_rule(45) \}
$PAR
tangent plane:$BR
\( z = \) \{ ans_rule(45) \}

END_TEXT
Context()->normalStrings;

ANS($normal->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $correct->isParallel($student);
  }
) );
ANS($plane->cmp() );

$ad = ( $a == 1 ) ? "" : $a;
$bd = ( $b == 1 ) ? "" : $b;
$cd = ( $c == 1 ) ? "" : $c;

$eq2 = Compute( "$fx0 (x - $x0) + $fy0 (y - $y0) + $fz0 (z - $z0)" )->reduce();

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

First, we check that \($ad($x0)^2 - $bd($y0)^2 + $cd($z0)^2 = $e0\).
Then let
\( f(x,y,z) = $eqn\) so that the given surface is the level surface
\( f(x,y,z) = $e0\). Since \(f_x = $fx\), \(f_y = $fy\), and \(f_z = $fz\),
we have \(\mbox{grad} f($x0,$y0,$z0) = $normal\)  Since gradients are
perpendicular to level surfaces, a vector normal to the surface at
\(($x0,$y0,$z0)\) is \(\vec n=$fx0\vec i + $fy0\vec j + $fz0\vec k\).
Thus an equation for the tangent plane is
\[
$eq2 = 0,
\]
or
\[
z = $plane.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
