# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 12/15/09

## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Directional derivatives and the gradient)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('14.5')
## Section2('14.5')
## Section3('14.5')
## Problem1('51')
## Problem2('52')
## Problem3('63')
## KEYWORDS('gradient', 'directional derivative', 'multivariable', 'functions')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

Context("Vector2D");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);
$showPartialCorrectAnswers = 1;

$dummy = random(1,5,1);
$a = random(1,5,1);
$trig = ("cos", "sin")[random(0,1,1)];
$trig2 = ($trig eq "cos") ? "-sin" : "cos";

$zfunc = Compute( "(e^x - $a*x)*$trig(y)" )->reduce();

$x0 = random(1,5,1);
$y0 = random(2,6,1);

$fx0 = Compute( "(e^($x0) - $a)*$trig($y0)" );
$fy0 = Compute( "(e^($x0) - $a*$x0)*($trig2($y0))" );
$grad0 = Vector( "-$fx0 i - $fy0 j" );

$vx0 = random(1,8,1);
$vy0 = random(1,8,1);
$aval = Compute( "$fx0*$vx0 + $fy0*$vy0" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Consider the function \( f(x,y) = $zfunc \). Suppose \( S \) is the
surface \( z=f(x,y). \)
$PAR
${BBOLD}(a)$EBOLD
Find a vector which is perpendicular to the level curve of \( f \)
through the point \( ($x0,$y0) \) in the direction in which \( f \) 
decreases most rapidly.
$BR
vector = \{ ans_rule(55) \}

$PAR
${BBOLD}(b)$EBOLD
Suppose \( \vec v = $vx0\vec i + $vy0 \vec j + a \vec k \)
is a vector in 3-space which is tangent to the surface \( S \)
at the point \( P \) lying on the surface above \( ($x0,$y0) \).
What is \( a \)?
$BR
\(a = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($grad0->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $correct->isParallel($student);
  }
) );
# ANS($grad0->cmp() );
ANS($aval->cmp() );

$trigtex = '\' . $trig;
if ( $trig2 eq 'cos' ) {
    $trig2tex = '\cos';
    $trig2sgn = '+';
    $ntrig2sgn = '-';
} else {
    $trig2tex = '\sin';
    $trig2sgn = '-';
    $ntrig2sgn = '+';
}
$trig2sgnd = $trig2sgn eq '+' ? '' : '-';
$ad = $a == 1 ? '' : $a;
$ax0 = $a*$x0;

$ex0 = ( $x0 == 1 ) ? 'e' : "e^{$x0}";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The vector \(\nabla f(x,y)\) is perpendicular to the level curve of
\(f\) through \((x,y)\):
\[
\nabla f(x,y) = (e^x-$a)$trigtex y\vec i $trig2sgn (e^x- $ad x)$trig2tex y\vec j.
\]
Thus, at the point \(($x0,$y0)\),
\[
\nabla f($x0,$y0)=($ex0-$a)$trigtex $y0\vec i $trig2sgn
    ($ex0-$ax0)$trig2tex $y0\vec j
\]
The vector \(\nabla f\) points in the direction of greatest increase in
\(f\), so the vector we want is
\[
-\nabla f($x0,$y0) = -($ex0-$a)$trigtex $y0\vec i $ntrig2sgn
    ($ex0-$ax0)$trig2tex $y0\vec j.
\]

$PAR
${BBOLD}(b)$EBOLD
To find a vector normal to the surface, we write the surface in
the form
\[
F(x,y,z)=(e^x- $ad x)$trigtex y-z=0.
\]
Then
\[
\nabla F = (e^x - $a)$trigtex y\vec i $trig2sgn (e^x - $ad x)$trig2tex y\vec j
   - \vec k.
\]
So, at \(P\),
\[
\nabla F = ($ex0 - $a)$trigtex $y0\vec i $trig2sgn
    ($ex0 - $ax0)$trig2tex y\vec j - \vec k.
\]
The vector \(\vec v\) is perpendicular to \(\nabla F\), so
\[
\vec v\cdot\nabla F =
($vx0\vec i + $vy0\vec j + a\vec k)\cdot \nabla F($x0,$y0) =0.
\]
Thus
\[ a = $vx0 (($ex0-$a)$trigtex $y0) + $vy0 ($trig2sgnd
    ($ex0-$ax0)$trig2tex $y0). \]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
