## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Limit of a sequence)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('APEX Calculus')
## AuthorText1('Hartman')
## EditionText1('3.0')
## Section1('8.1')
## Problem1('20')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());

Context("Numeric")->variables->are(n=>'Real');
Context()->strings->add('diverges' => {caseSensitive => 0});
Context()->noreduce('(-x)-y');

$b = random(2,9,1);
$bb = $b * $b;
$b2 = 2 * $b;

$a = Compute("($b n + 1)/($b n) - ($b n) / ($b n + 1)");

$c3 = Compute("$b2 n + 1");
$c4 = Compute("$bb * n^2 + $b n");

$ans1num = Compute("2*$b + 1/n");
$ans1den = Compute("$b^2 * n + $b");


$ans2 = Real("0");

$multians1 = MultiAnswer($ans1num,$ans1den)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  format => "[ %s ] / [ %s ]",
  tex_format => "\frac{%s}{%s}",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ($c1,$c2) = @{$correct};
      my ($s1,$s2) = @{$student};
      if ($c1 == $s1 && $c2 == $s2) { return [1,1]; }
      $self->setMessage(1,"Factor and cancel") if ($c3 == $s1 && $c4 == $s2);
      return [0,0];
  }
);

sub displayfrac {
    my $ma = shift; # input is a multianswer object
    ANS($ma->cmp);
    if ($displayMode eq 'TeX') {
      "\(\displaystyle \frac{".$ma->ans_rule(10)."}{".$ma->ans_rule(10)."} \)";
    } else {
      "<table style='display:inline-block; vertical-align:middle; border:none;'><tr><td>".$ma->ans_rule(15).$BR.$HR.$ma->ans_rule(15)."</td></tr></table>" 
    }
}



BEGIN_PGML
Use algebra to simplify the expression before evaluating the limit.  In particular, combine the fractions into one fraction, factor the highest power of [`n`] from the numerator and denominator, then cancel as many factors of [`n`] as possible.  If the sequence does not converge, enter _diverges_ in the final answer box.

[``\lim_{n \to \infty} \left( [$a] \right) = \lim_{n \to \infty} \bigg(``] 
[@ displayfrac($multians1) @]* 
[``\bigg) = ``] [_____________]{$ans2}.
END_PGML


ENDDOCUMENT();
