##DESCRIPTION
## Word problem
## 
##ENDDESCRIPTION
## DBsubject('Intermediate Algebra')
## DBchapter('Ch 01: Algebra, Mathematical Models and Problem Solving')
## DBsection('Algebra, Mathematical Models, and Problem Solving')
## KEYWORDS('mathematical models','problem solving','formulas','algebra')
## TitleText1('Essentials of Intermediate Algebra')
## EditionText1('1')
## AuthorText1('Blitzer')
## Section1('1.5')
## Problem1('')
## Author('RA Cruz')
## Institution('The College of Idaho')
## Date: 2007/09

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "AnswerFormatHelp.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

######################################
#  Setup

Context()->variables->are(C=>'Real',
                          D=>'Real',
                          M=>'Real',
                          m=>'Real',
                          P=>'Real',
                          p=>'Real',
                          T=>'Real',
c=>'Real',
                          d=>'Real',
                          P=>'Real',
                          p=>'Real',
                          T=>'Real',
t=>"Real",);

for ($i = 0; $i < random(10,50); $i++) {
$n = random(0,1,1);
$a = random(2,15);
do {$b = random(2,15)} until ($a != $b);
do {$c = random(2,15)} until (gcd($a,$c) == 1);
}
@var = ("M", "p");
@exp1 = ("$a P"," $a T");
@exp2 = ("$b C+$c MC","$b D+ $c pm");
@nums = ("$a P- $b C", "$a T- $b D");
@dens = ("$c C", " $c m");
$num = Formula($nums[$n])->reduce;
$den = Formula($dens[$n])->reduce;
Context()->texStrings;
$equation = "$exp1[$n] = $exp2[$n]";

$frac = $var[$n];
$dispans = "\frac{$num}{$den}";

$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) || 
           (-$f1==$f1stu && -$f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu || -$f1==$f1stu) {
          return [1,0];
      } elsif ( ($numbogus==$f1stu || -$numbogus==$f1stu) ||
                ($denbogus==$f2stu || -$denbogus==$f2stu) ) {
          $self->setMessage(1,"Find a common denominator first");
          $self->setMessage(2,"Find a common denominator first");
          return [0,0];
      } elsif ( $f2==$f2stu || -$f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);

if ($displayMode eq 'TeX') {
#  $displayfrac = "\( $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \)";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Solve the equation for \($frac \):
$PAR
\[ $equation \]
$PAR
$displayfrac
END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;

ANS($multians->cmp(studentsMustReduceFractions=>1));

Context()-texStrings;
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \($frac = \displaystyle $dispans\)
  $PAR
  END_TEXT
}
else  {
    BEGIN_TEXT
    $PAR
    $BBOLD Note: $EBOLD
    Write your answer as a single fraction, and type the numerator and denominator separately in the answer blanks provided.
    $PAR
    END_TEXT
} 

Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
