# DESCRIPTION
# Solve for the given variable
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Equations and Inequalities')
## DBsection('Other Equations')
## KEYWORDS('equations','rational')
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "MUHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x", "y", "z", "a", "b", "c", "p", "q", "r", "s", "t",  "v", "w", "m", "n", "k");
$size = scalar(@vars);
for ($i = 1; $i < $size; $i++) {
  Context()->variables->add($vars[$i]=>"Real", uc($vars[$i])=>"Real");
}
Context()->variables->add(uc("x")=>"Real");
$ch[0] = random(0,$size-1,1);
do {$ch[1] = random(0,$size-1,1)} until ($ch[1] != $ch[0]);
do {$ch[2] = random(0,$size-1,1)} until ($ch[2] != $ch[1] && $ch[2] != $ch[0]);
do {$ch[3] = random(0,$size-1,1)} until ($ch[3] != $ch[2] && $ch[3] != $ch[1] && $ch[3] != $ch[0]);
$c = uc($vars[$ch[0]]);
$v[0] = $vars[$ch[1]];
$v[1] = $vars[$ch[2]];
$v[2] = $vars[$ch[3]];
$a = random(2,9);

$eqn = "$c = \frac{$v[0] ($v[1]+$v[2])}{$a}";

$choice = random(0,2);
$vc = $v[$choice];

if ($choice == 0) {
  $dans = Formula("$a*$c/($v[1]+$v[2])")->reduce;
  $num = Formula("$a*$c")->reduce;
  $den = Formula("$v[1]+$v[2]")->reduce;
} elsif ($choice == 1) { 
  $dans = Formula("($a*$c - $v[0]*$v[2])/$v[0]")->reduce;
  $num = Formula("$a*$c - $v[0]*$v[2]")->reduce;
  $den = Formula("$v[0]")->reduce;
} else {
  $dans = Formula("($a*$c - $v[0]*$v[1])/$v[0]")->reduce;
  $num = Formula("$a*$c - $v[0]*$v[1]")->reduce;
  $den = Formula("$v[0]")->reduce;
}

$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 0,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1 $var1^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01 $var1^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $vc = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $vc = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Solve the following equation for \($vc\):
$PAR \[ $eqn \] $PAR
$displayfrac
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle $vc = $dans\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $BR
  ${BBOLD}Help:${EBOLD} \{ MUHelp("algfrac") \}
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
