# DESCRIPTION
# Solve equation of the form A = n*B/d + a for B, where gcd(n,d) = 1, and 
# a is an arbitrary non-zero integer.
# ENDDESCRIPTION

## DBsubject('Algebra') 
## DBchapter('Equations and Inequalities');
## DBsection('Linear Equations') 
## KEYWORDS('linear', 'equations')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "MUHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
PeriodicRerandomization("3"); 


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->strings->add("no solutions"=>{}, "infinitely many"=>{}, "no solution"=>{alias=>"no solutions"}, "none"=>{alias=>"no solutions"}, "all reals"=>{alias=>"infinitely many"});
@v = ("F","G","H","C","K","B","M","N");
$c1 = random(0,scalar(@v)-1,1);
do {$c2 = random(0,scalar(@v)-1,1)} until ($c2 != $c1);
$v1 = $v[$c1]; $v2 = $v[$c2];
Context()->variables->are($v1=>"Real", $v2=>"Real", lc($v1)=>"Real", lc($v2)=>"Real");

$n = random(1,9);
do {$d = random(2,9)} until (gcd($n,$d) == 1);
$a = non_zero_random(-15,15);

$eqn = "$v1 = \frac{$n}{$d}$v2 + $a";

$danswer = Formula("$d ($v1 - $a)/$n")->reduce;
$num = Formula("$d($v1-$a)")->reduce;
$den = Formula("$n")->reduce;
$answer = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 0,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
    my ( $r1, $r2 ) = (0,0);
    $s1 = Formula($s1) unless ($s1->type eq 'Formula' || $s1->type eq 'String');
    $s2 = Formula($s2) unless ($s2->type eq 'Formula' || $s2->type eq 'String');
    $s1 = Formula("0.1 $var1^{1000}") if ($s1->type eq 'String'); #bogus answer
    $s2 = Formula("0.01 $var1^{500}") if ($s2->type eq 'String'); #bogus answer
    $r1 = 1 if ($c1 == $s1);
    $r2 = 1 if ($c2 == $s2);
    ($r1,$r2) = (1,1) if ($c1 == -$s1 && $c2 == -$s2);
    if ($c1*$s2 == $c2*$s1 && $c1 != $s1 && $c1 != -$s1) {
      $self->setMessage(1,"Your answer can be simplified further");
      $self->setMessage(2,"Your answer can be simplified further");
    }
    return [$r1,$r2];
  }
);

if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $v2 = ".$answer->ans_rule(10).$answer->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $v2 = \)",
  $answer->ans_rule(20).$BR.$HR.$answer->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Solve the equation for \($v2\):
$PAR \[ $eqn \] $PAR
$displayfrac
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());


##########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
Context()->texStrings;
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \($v2 = $danswer\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $BR
  ${BBOLD}Help:${EBOLD} \{ MUHelp("algfrac") \}
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus();

COMMENT('Features Periodic Rerandomization. Edited and updated in 2012/2013.
$BR
Desc: Solve equation of the form A = n*B/d + a for B, where gcd(n,d) = 1, and a is an arbitrary non-zero integer.');
ENDDOCUMENT();
