## DESCRIPTION
## Calculus
## ENDDESCRIPTION


## KEYWORDS('Multivariable','change of variables')
## Tagged by tda2d

## DBsubject('Calculus')
## DBchapter('Multiple Integrals')
## DBsection('Change of Variables in Multiple Integrals')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGmatrixmacros.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


###########################
#  Setup

Context("Numeric")->variables->are(u=>'Real',v=>'Real');

$a = random(2,4,1);
$b = random(5,7,1);

$aa = $a * $a;
$bb = $b * $b;

$x = Compute("u^2 / v");
$y = Compute("v / u");

$ulo = Real(1);
$uhi = Real($aa);
$vlo = Real(1);
$vhi = Real($bb);


$multians0 = MultiAnswer( $x, $y )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>  </table>",
  tex_format =>  "\begin{array}{r} %s \\ %s \end{array}",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians1 = MultiAnswer( Compute("2 u / v"), Compute("-u^2 / v^2"), Compute("-v / u^2"), Compute("1 / u"), Compute("1 / v") )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table> %s",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack; %s",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..4) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$area = Compute("($aa - 1) * ln($bb)");







###########################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Compute the area of the region \( D \) bounded by 
$PAR
\[
xy = 1, \
xy = $aa, \ \ \
xy^2 = 1, \ 
xy^2 = $bb
\]
$BR
in the first quadrant of the \( xy \)-plane.
$BR
$BR
(a) Graph the region \( D \).
$BR
$BR
(b) Using the non-linear change of variables 
\( u = xy \) and \( v = xy^2 \), find \( x \)
and \( y \) as functions of \( u \) and \( v \).
$BR
$BR
\( x = x(u,v) = \) \{ $multians0->ans_rule(20) \}
$BR
\( y = y(u,v) = \) \{ $multians0->ans_rule(20) \}
$BR
$BR
(c) Find the determinant of the 
Jacobian for this change of variables.
$PAR
$BCENTER
\{ 
mbox(
  "\( \displaystyle \left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  " = ".$SPACE.$multians1->ans_rule(20)
);
\}
$ECENTER
$BR
(d) Using the change of variables, set up a double integral
for calculating the area of the region \( D \).
$BR
$BR
END_TEXT
Context()->normalStrings;

############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

#ANS( $x->cmp() );
#ANS( $y->cmp() );

ANS( $multians0->cmp() );

ANS( $multians1->cmp() );








$multians2 = MultiAnswer( Compute("$bb"), Compute("1"), Compute("$aa"), Compute("1"), Compute("1/v") )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='0'><tr><td>
  <table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table></td> 
  <td><table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table></td>
  <td>%s</td></tr></table>",
  tex_format =>  "\int^{%s}_{%s} \int^{%s}_{%s} %s du dv",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..4) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);




#$integraltex = "\int_a^b \int_c^d f(x(u,v),y(u,v)) \left| \frac{\partial(x,y)}{\partial(u,v)} \right| du \, dv";

$integraltex = "\int\!\int_D dx \, dy = \int_a^b \int_c^d \left| \frac{\partial(x,y)}{\partial(u,v)} \right| du \, dv";

#
#  Display the answer blanks properly in different modes
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
   $integral =
   "\( \displaystyle $integraltex = \) \\\\ \( \displaystyle ".
   '\int_{'.$multians2->ans_rule(4).'}^{'.$multians2->ans_rule(4).'}'. 
   '\int_{'.$multians2->ans_rule(4).'}^{'.$multians2->ans_rule(4).'}'. 
   $multians2->ans_rule(20).'du \, dv \)';
  } else {
   $integral =
   BeginTable(center=>0).
     Row([
       "\( \displaystyle $integraltex = \int \)",
       $multians2->ans_rule(4).$BR.$BR.$multians2->ans_rule(4),
       '\( \displaystyle \int \)',
       $multians2->ans_rule(4).$BR.$BR.$multians2->ans_rule(4),
       $multians2->ans_rule(20), '\( du \, dv \)'],separation=>2).
   EndTable();
}
Context()->normalStrings;



Context()->texStrings;
BEGIN_TEXT
$integral
$BR
$BR
(e) Evaluate the double integral and compute the area of the region \( D \).
$BR
$BR
Area = \{ ans_rule(20) \} 
END_TEXT
Context()->normalStrings;


#ANS( Compute("1")->cmp() ); # A 
#ANS( Compute("$aa")->cmp() ); # B
#ANS( Compute("1")->cmp() ); # C
#ANS( Compute("$bb")->cmp() ); # D
#ANS( Compute("1 / v")->cmp() ); # Integrand * Jacobian

ANS( $multians2->cmp() );

ANS( $area->cmp() );

COMMENT('MathObject version.');

ENDDOCUMENT();
















\( $integraltex \)
$BR
\( a = \) \{ ans_rule(20) \}
$BR
\( b = \) \{ ans_rule(20) \}
$BR
\( c = \) \{ ans_rule(20) \}
$BR
\( d = \) \{ ans_rule(20) \}
$BR
\(  f(x(u,v),y(u,v)) 
\left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \) \{ ans_rule(20) \}
