## DESCRIPTION
## Calculus
## ENDDESCRIPTION


## KEYWORDS('Multivariable','change of variables')
## Tagged by tda2d

## DBsubject('Calculus')
## DBchapter('Multiple Integrals')
## DBsection('Change of Variables in Multiple Integrals')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


###########################
#  Setup

Context("Numeric")->variables->are(u=>'Real',v=>'Real');

$a = -2;
$b = random(3,6,1);

$x = Compute("e^($a u) cos($b v)");
$y = Compute("e^($a u) sin($b v)");

$ab = abs($a * $b);
$a2 = 2 * $a;

$det = Compute("$ab e^($a2 u)");

$A11 = $x->D('u');  $A12 = $x->D('v');
$A21 = $y->D('u');  $A22 = $y->D('v');

$multians1 = MultiAnswer($A11, $A12, $A21, $A22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);





###########################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose a change of coordinates 
\( T : \mathbb{R}^2 \to \mathbb{R}^2 \) 
from the \( uv \)-plane to the \( xy \)-plane is given by
$PAR
\[
x = $x, \ \ \ y = $y.
\]
$BR
Find the absolute value of the determinant of the
Jacobian for this change of coordinates.
$PAR
$BCENTER
\{ 
mbox(
  "\( \displaystyle \left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \bigg| \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(20),$multians1->ans_rule(20)], 
  [$multians1->ans_rule(20),$multians1->ans_rule(20)]], align=>'cc'),
  "\( \bigg| = \)".$SPACE.ans_rule(20)
);
\}
$ECENTER
END_TEXT
Context()->normalStrings;


############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );

ANS( $det->cmp() );


COMMENT('MathObject version.');

ENDDOCUMENT();
