## DESCRIPTION
## Calculus
## ENDDESCRIPTION


## KEYWORDS('Multivariable','change of variables')
## Tagged by tda2d

## DBsubject('Calculus')
## DBchapter('Multiple Integrals')
## DBsection('Change of Variables in Multiple Integrals')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGmatrixmacros.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


###########################
#  Setup

Context("Point")->variables->are(u=>'Real',v=>'Real');

$a = list_random(2,3,5,6,7,8);
do { $b = random(2,9,1); } until ($b != $a);
do { $c = random(2,9,1); } until ($c != $b && $c != $a);

$bb = $b * $b;
$cc = $c * $c;

$x = Compute("($c / sqrt($a)) u");
$y = Compute("($c / $b) v");

$multians1 = MultiAnswer( Compute("$c / sqrt($a)"), 0, 0, Compute("$c / $b") )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..3) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$det = Compute("$c^2 / (sqrt($a) * $b)");

$area = Compute("pi $c^2 / (sqrt($a) * $b)");




###########################
#  Main text

Context()->texStrings;
BEGIN_TEXT
(a) Find a linear change of coordinates that takes the unit 
circle \( u^2 + v^2 = 1 \) in the \( uv \)-plane to the 
ellipse \( \displaystyle $a x^2 + $bb y^2 = $cc \) in the \( xy \)-plane.
$BR
$BR
\( x = x(u,v) = \) \{ ans_rule(20) \}
$BR
\( y = y(u,v) = \) \{ ans_rule(20) \}
$BR
$BR
(b) Find the absolute value of the determinant of the 
Jacobian for this change of coordinates.
$PAR
$BCENTER
\{ 
mbox(
  "\( \displaystyle \left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \bigg| \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  "\( \bigg| = \)".$SPACE.ans_rule(20)
);
\}
$ECENTER
$BR
(c) Use your answer to part (b) to find the area enclosed by the ellipse \( $a x^2 + $bb y^2 = $cc \).
$BR
$BR
Area = \{ ans_rule(20) \} 
END_TEXT
Context()->normalStrings;


############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $x->cmp() );
ANS( $y->cmp() );

ANS( $multians1->cmp() );
ANS( $det->cmp() );

ANS( $area->cmp() );

COMMENT('MathObject version.');

ENDDOCUMENT();
