# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus','differential equations','trigonometric functions')
## Tagged by glr 02/06/09

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Higher Order Differential Equations')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('11.10')
## Section2('11.10')
## Problem1('6')
## Problem2('6')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')
## HHSection1('Second-Order Differential Equations: Oscillations')
## HHSection2('Second-Order Differential Equations: Oscillations')


DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( t=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(1,5,1);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

What values of \( \alpha \) and \( A \) make \( y = A\cos \alpha t \)
a solution
to \( y''+ $a y=0 \) such that \( y'(1)=$b \)?

$PAR
\(\alpha =\) \{ ans_rule(35) \}
$BR
\(A =\) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( Compute("sqrt($a)")->cmp( checker=>sub {
    my ( $correct, $student, $ansHash ) = @_;
    return $correct == $student || -1*$correct == $student;
} ) );

ANS( Compute("$b/(-sqrt($a)*sin(sqrt($a)))")->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

If
\(y=A\cos \alpha t\), then
\(y'=-\alpha A\sin \alpha t\) and
\(y''=-\alpha^2 A\cos \alpha t\).  Thus,
if \(y''+ $a y=0\), then \(-\alpha^2 A\cos \alpha t+ $a A\cos \alpha t=0\),
so \(A($a -\alpha^2)\cos \alpha t=0\). This is true for all \(t\) if
\(A=0\), or if \(\alpha =\pm \sqrt{$a}\).
We also have the initial condition: \(y'(1)=-\alpha A\sin \alpha = $b\).
Notice that this equation will not work if \(A=0\). If
\(\alpha = \sqrt{$a}\), then
\(A=-\frac{$b}{\sqrt{$a} \sin \sqrt{$a}}\).
$BR
Similarly, if \(\alpha =-\sqrt{$a}\), we find the same value for \(A\).
$PAR
Thus, the possible values are
\(A=-\frac{$b}{\sqrt{$a} \sin \sqrt{$a}}\) and
\(\alpha =\pm \sqrt{$a}\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
