##DESCRIPTION
##  Combinatorics: Recurrence relations.
##ENDDESCRIPTION

## Author1('Monika Keindl')
## Author2('Nandor Sieben')
## Author3('Terry Blows')
## DBsubject(Combinatorics)
## DBchapter(Recurrence relations)
## DBsection(Solving homogeneous)
## Date(December 2018)
## Institution(Northern Arizona University)

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl",
  "parserMultiAnswer.pl"
);

TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Numeric");

{
$x = non_zero_random(-7,7,1);
$y = non_zero_random(-7,7,1);
$z = non_zero_random(-7,7,1);
redo if ($x==$y or $x==$z or $y==$z);
}

$answer1 = Compute("$x");
$answer2 = Compute("$y");
$answer3 = Compute("$z");
$answer4 = Compute("$y");
$answer5 = Compute("$z");
$answer6 = Compute("$z");

$multians1 = MultiAnswer($answer1, $answer2, $answer3)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2, $stu3 ) = @{$student};
      my ( $a1, $a2, $a3 ) = @{$correct};
      my ($stu, $cor, $r1, $r2, $r3);

      $r1 = 0;
      $r2 = 0;
      $r3 = 0;
      if ($stu1==$a1 or $stu1==$a2 or $stu1==$a3) {
        $r1 = 1;
      }
      if ($stu2==$a1 or $stu2==$a2 or $stu2==$a3) {
        $r2 = 1;
      }
      if ( $stu2 == $stu1 ) {
        $r2 = 0;
      } 
      if ($stu3==$a1 or $stu3==$a2 or $stu3==$a3) {
        $r3 = 1;
      }
      if ( $stu3 == $stu1 or $stu3 == $stu2) {
        $r3 = 0;
      } 
      return [$r1, $r2, $r3];
      }
);

$multians2 = MultiAnswer($answer2, $answer3)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $stu1, $stu2 ) = @{$student};
      my ( $a1, $a2 ) = @{$correct};
      my ($stu, $cor, $r1, $r2);

      $r1 = 0;
      $r2 = 0;
      if ($stu1==$a1 or $stu1==$a2) {
        $r1 = 1;
      }
      if ($stu2==$a1 or $stu2==$a2) {
        $r2 = 1;
      }
      if ( $stu2 == $stu1 ) {
        $self->setMessage(2,"");
        $r2 = 0;
      } 
      return [$r1, $r2];
      }
);


##############################################################
#
#  Text
#

Context()->texStrings;
BEGIN_TEXT
A linear, constant-coefficient homogeneous recurrence relation has characteristic roots $x, $y, $y, $z, $z, $z.
What is the general solution of this recurrence relation?

$PAR
\( a_n = \alpha_1( \)\{$multians1->ans_rule(2)\}\()^n+
         \alpha_2( \)\{$multians1->ans_rule(2)\}\()^n+
         \alpha_3( \)\{$multians1->ans_rule(2)\}\()^n+
         \alpha_4n( \)\{$multians2->ans_rule(2)\}\()^n+
         \alpha_5n( \)\{$multians2->ans_rule(2)\}\()^n+
         \alpha_6n^2( \)\{ans_rule(2)\}\()^n
\)


END_TEXT
Context()->normalStrings;


##############################################################
#
#  Answers
#

ANS( $multians1->cmp() ); 
ANS( $multians2->cmp() ); 
ANS(Compute($answer6)->cmp);

ENDDOCUMENT();
