## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Matrix notation for systems)
## Date(01/30/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2006')
## Section1('4.1')
## Problem1('14')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

install_weighted_grader();

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(t=>"Real");

$a = random(2,6,1);
do { $b = random(2,6,1); } until ($b != $a);
do { $n = random(2,6,1); } until ($n != $a && $n != $b);

$P11 = Formula("t^$n");   $P12 = Real("$a");
$P21 = Formula("sin(t)"); $P22 = Formula("t");

$g1 = Formula("sec(t)");
$g2 = Real("-$b");

$multians1 = MultiAnswer($P11, $P12, $P21, $P22)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "[ %s, &nbsp;&nbsp; %s ] <br /> [ %s, &nbsp;&nbsp; %s ]",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c11, $c12, $c21, $c22 ) = @{$correct};
      my ( $s11, $s12, $s21, $s22 ) = @{$student};
      my @score = (0,0,0,0);
      if (ref($c11) eq ref($s11) && $c11 == $s11) { $score[0] = 1; }
      if (ref($c12) eq ref($s12) && $c12 == $s12) { $score[1] = 1; }
      if (ref($c21) eq ref($s21) && $c21 == $s21) { $score[2] = 1; }
      if (ref($c22) eq ref($s22) && $c22 == $s22) { $score[3] = 1; }
      return [ @score ];
  }

);

$multians2 = MultiAnswer($g1, $g2)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  format =>  "[ %s ] <br /> [ %s ]",
  tex_format =>  "\left\lbrack\begin{array}{r} %s \\ %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my ( $c1, $c2 ) = @{$correct};
      my ( $s1, $s2 ) = @{$student};
      my @score = (0,0);
      if (ref($c1) eq ref($s1) && $c1 == $s1) { $score[0] = 1; }
      if (ref($c2) eq ref($s2) && $c2 == $s2) { $score[1] = 1; }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Suppose
$PAR
\[
\begin{array}{rcl}
y_1' & = & $P11 y_1 + $P12 y_2 + $g1, \\
y_2' & = & $P21 y_1 + $P22 y_2 + $g2. 
\end{array}
\]
$BR
This system of linear differential equations can
be put in the form \( \mathbf{y}' = P(t) \mathbf{y} + \mathbf{g}(t) \).
Determine \( P(t) \) and \( \mathbf{g}(t) \).

$PAR
$BCENTER
\{ 
mbox(
  "\( P(t) = \)",
  display_matrix([
  [$multians1->ans_rule(15),$multians1->ans_rule(15)], 
  [$multians1->ans_rule(15),$multians1->ans_rule(15)]
  ], align=>'cc')
);
\}
$ECENTER

$PAR
$BCENTER
\{ 
mbox(
  "\( \mathbf{g}(t) = \)",
  display_matrix([
  [$multians2->ans_rule(15)], 
  [$multians2->ans_rule(15)]
  ], align=>'cc')
);
\}
$ECENTER

END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

WEIGHTED_ANS( $multians1->cmp(), 80 );
WEIGHTED_ANS( $multians2->cmp(), 20 );

COMMENT("Uses parserMultiAnswer.pl and has fancy formatting of student answers.");

ENDDOCUMENT();
