## DESCRIPTION
## Systems of ODEs: introduction to systems of ODEs
## ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Systems of differential equations)
## DBsection(Distinct real eigenvalues)
## Date(01/30/2011)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(SDE-DRE-IVP)
## Level(3)
## MO(1)
## TitleText1('Notes on Diffy Qs')
## AuthorText1('Jiri Lebl')
## EditionText1('December 9, 2010')
## Section1('3.1')
## Problem1('3')
## KEYWORDS('differential equations','systems of ODEs')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "AnswerFormatHelp.pl",
  "parserAssignment.pl",
  "unionLists.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->are(
t=>"Real"
);

$a = random(1,3,1);
$b = random(4,8,1);
$apb = $a + $b;
$bma = $b - $a;
$ab = $a * $b;

# char poly (r+a)(r-b)

$x0 = 0;
$y0 = random(1,9,1);

$y0a = $y0 * $a;
$y0b = $y0 * $b;

$x = Compute("-$y0/$apb e^(-$a t) + $y0/$apb e^($b t)");
$y = Compute("$y0a/$apb e^(-$a t) + $y0b/$apb e^($b t)");


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Find the solution to
$PAR
\[
\begin{array}{lcl}
x^{\,\prime} & = & y \\
y^{\,\prime} & = & $ab x + $bma y 
\end{array}
\]
$BR
$BR
if \( x(0) = $x0 \) and \( y(0) = $y0 \).
$BR
$BR
\( x(t) = \) 
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}
$BR
\( y(t) = \) 
\{ ans_rule(40) \}
\{ AnswerFormatHelp("formulas") \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $x->cmp() );
ANS( $y->cmp() );


;

ENDDOCUMENT();



$ITEMSEP
$ITEM Find the most general solution to the 
associated homogeneous differential equation.  
Use \( A \) and \( B \) in your answer to denote 
arbitrary constants.
$BR
$BR
\( y_h = \)
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}


$ITEMSEP
$ITEM Find the solution to the original 
nonhomogeneous differential equation satisfying
the initial conditions \( y(0) = $y0 \) and 
\( y'(0) = $yprime0 \).
$BR
$BR
\( y = \)
\{ ans_rule(60) \}
\{ AnswerFormatHelp("formulas") \}





ANS( $yhomogeneous->cmp( 
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my $stu   = Formula($student);

    ################################
    #  Check for arbitrary constants
    #
    Value->Error("Is your answer the most general solution?") 
    if ( 
      Formula($stu->D('A'))==Formula(0) || 
      Formula($stu->D('B'))==Formula(0) 
    );

    ############################################
    #  Check for linear independence (Wronskian)
    #
    my $x = Real(1.24);

    my $a11 = $stu->eval(A=>1,B=>0,x=>$x,y=>0);
    my $a12 = $stu->eval(A=>0,B=>1,x=>$x,y=>0);
 
    my $a21 = $stu->D('x')->eval(A=>1,B=>0,x=>$x,y=>0);
    my $a22 = $stu->D('x')->eval(A=>0,B=>1,x=>$x,y=>0);

    my $wronskian = $a11*$a22 - $a21*$a12;

    Value->Error("Your functions are not linearly independent or your answer is not complete") 
    if ($wronskian==Real(0));

    ########################################################
    #  Check that the student answer is a solution to the DE
    #
    my $stux  = Formula($stu->D('x'));
    my $stuxx = Formula($stu->D('x','x'));
    my $stuDE = Formula($stuxx + $kk*$stu)
    ->with(test_points=>[[1,1,1,0],[1,1,0,0],[2,1,1,0],[1,2,0,0]]);

    return ($stuDE==Formula(0));

}));



ANS( $yinitvalprob->cmp() );

