# DESCRIPTION
# Basic exponent laws
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## KEYWORDS('exponents')
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "PGunion.pl",
  "contextFraction.pl",
  "contextLimitedPowers.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
@vars = ("x","y","z","a","b","c","d","s","t");
$c1 = random(0,scalar(@vars)-1,1);
do {$c2 = random(0,scalar(@vars)-1,1)} until ($c2 != $c1);
do {$c3 = random(0,scalar(@vars)-1,1)} until ($c3 != $c1 && $c3 != $c2);
$var1 = $vars[$c1];
$var2 = $vars[$c2];
$var3 = $vars[$c3];
Context()->variables->are($var1=>"Real", $var2=>"Real", $var3=>"Real");
$p1 = random(2,19);
$p2 = random(2,7);
$p3 = random(2,7);
$p4 = random(2,7);
$pp1 = random(2,7);

$frac = "$var1^{$p1} \left($var1^{$p2} $var2^{$p3} $var3^{$p4}\right)^{-$pp1}";
if ($p1-$pp1*$p2 <= 0) {
  $num = Formula("1")->reduce;
  $den = Formula("$var1^{$p2*$pp1 - $p1} $var2^{$pp1*$p3} $var3^{$pp1*$p4}")->reduce;
} else {
  $num = Formula("$var1^{$p1-$p2*$pp1}")->reduce;
  $den = Formula("$var2^{$pp1*$p3} $var3^{$pp1*$p4}")->reduce;
}
LimitedPowers::OnlyPositiveIntegers();
$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);
Context()->operators->undefine("+","-","/");


if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###########################################################################
# problem statement 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write the following using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
Context()->operators->undefine("+","-");
ANS($multians->cmp() );


Context()->texStrings;
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle \frac{$num}{$den}\right\)
  $PAR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
