##DESCRIPTION
##
## Created by the UBC School of Biomedical Engineering and Department of Mechanical Engineering. Project led by Robyn Newell and Agnes d'Entremont.
##
## Contact: robyn.newell@ubc.ca and agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this 
## project provided by UBC Open Education Resources Fund. We also gratefully acknowledge
## additional funding support from BCcampus and the School of Biomedical
## Engineering.  
## Contributors to this project include Thomas Oxland, Peter Cripton, Dave Wilson, Lyndia Wu,
## Pranav Shrestha, Jackson Robinson, Jacqueline Kho, Laura Vargas, Mona Behrouzian
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people. 
##
##ENDDESCRIPTION

## DBsubject(Engineering Biomechanics Dynamics)
## DBchapter(Applications: Injury Biomechanics)
## DBsection(Accident Reconstruction)
## Institution('University of British Columbia')
## Author('UBC-SBME')
## Date(July 2021)
## Level()
## KEYWORDS()
## RESOURCES('')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
## Beginner

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PG.pl",
   "PGbasicmacros.pl",
   "PGchoicemacros.pl",
   "PGanswermacros.pl",
   "unionTables.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$m = 1830;

$age = 53;
$height = 187;
$driverMass = 167;

$dV = 80;
$BEVint = 53;

$crushWidth = random(0.8, 1.2, 0.001);
$C1 = random(32, 42, 1);
$C2 = random (32, 42, 1);
$A = 511;
$B = 68;

$G = $A**2/(2*$B);
$W = $crushWidth;
$E = $W*($G+(($A/2)*($C1+$C2))+(($B/6)*($C1**2 + $C2**2 + ($C1*$C2))));

$BEV1 = Real(sqrt(2*$E/$m)*3.6);
$BEV2 = Real(sqrt(2*$E/($m+$driverMass))*3.6);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
Consider the following CIREN case: $BR
\(\circ\) Frontal impact, Kia Sorento (mass = $m kg)
$BR \(\circ\) $age-year-old restrained male driver, $height cm and $driverMass kg
$BR \(\circ\) Delta-V = $dV km/hr; BEV = $BEVint km/hr
$BR \(\circ\) Crush width = $crushWidth m;
$BR \(\circ\) Crush profile: C1 = $C1 cm, C2 = $C2 cm (simplified from original 6-point crush profile) 
$BR \(\circ\) Generic crush coefficients: A = $A N/cm, B = $B N/(cm<sup>2</sup>)
$PAR
Calculate an estimate of BEV in km/hr using the generic crush coefficients. Assume perfectly frontal collision. $BR

\(BEV\) = \{ans_rule(10)\}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS(Real($BEV1)->cmp(
     checker => sub {
     my ($correct,$student,$ansHash) = @_;  # get correct and student MathObjects
     my ($sx) = $student->value;        # get student answer
     if ($sx == $BEV1 || $sx == $BEV2){
     $studentVal = $sx;
     return 1;
     }
     else {return 0;      # return 1 if correct, 0 otherwise
     }
    }));
ENDDOCUMENT();        
