##DESCRIPTION
## DBsubject('Statistics')
## DBchapter('Hypothesis tests')
## DBsection('Two sample proportion')
## Level(3)
##ENDDESCRIPTION
##KEYWORDS('Two-sample proportions', 'hypothesis tests', 'significance level', 'confidence interval')
##Tagged by Nelson Chen. Jul 8, 2016. 
##Attempts: Four.

########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGnumericalmacros.pl",
  "PGstatisticsmacros.pl",
  "parserRadioButtons.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 0;

##############################################################
#
#  Setup
#
#
Context("Numeric");

@m1 = rserve_eval('
m1 <- round(runif(1, min=75, max=79))
');

@p1hat = rserve_eval('
p1hat <- m1/100
');

@m2 = rserve_eval('
m2 <- round(runif(1, min=81, max=85))
');

@p2hat = rserve_eval('
p2hat <- m2/100
');



######################################################
# Part 1
$mc1_answer      = "\(H_0: p_C - p_N = 0\)  and  \(H_A: p_C - p_N \ne 0\)";

$mc1_distractor1 = "\(H_0: p_C - p_N = 0\)  and  \(H_A: p_C - p_N > 0\)";
$mc1_distractor2 = "\(H_0: p_C - p_N = 0\)  and  \(H_A: p_C - p_N < 0\)";
$mc1_distractor3 = "\(H_0: p_N = @p1hat[0] \)  and  \(H_A: p_N \ne @p1hat[0] \)";

$mc1 = new_multiple_choice ();
$mc1->qa(
    "What are the appropriate null and alternative hypotheses?",
    $mc1_answer
);

$mc1->extra(
    $mc1_distractor1,
    $mc1_distractor2,
    $mc1_distractor3,
);


######################################################
# Part 2
$mc2_answer = "The probability of committing the Type I error is 10$PERCENT.";

$mc2_distractor1 = "There is 10$PERCENT chance to reject the null hypothesis.";

$mc2_distractor2 = "The Type II error is 90$PERCENT.";

$mc2_distractor3 = "There is 10$PERCENT chance that the null hypothesis is true.";

$mc2_distractor4 = "The Type I error is 10$PERCENT.";

$mc2_distractor5 = "The probability of committing the Type II error is 90$PERCENT.";

$mc2 = new_checkbox_multiple_choice();

$mc2 -> qa (
"The committee decides to perform the hypothesis test at 10$PERCENT significance level. Which of the following is/are true? (Check all that apply.)", 
   $mc2_answer,
);
$mc2 -> extra(
   $mc2_distractor1,
   $mc2_distractor2,
   $mc2_distractor3,
   $mc2_distractor4,
   $mc2_distractor5,
);
$mc2->makeLast("None of the above.");

######################################################
#Part 3

$mc3_answer1 = "The two random samples are no larger than 10$PERCENT of their respective population sizes.";

$mc3_answer2 = "The two random samples of students are independent.";

$mc3_distractor1 = "The variable whether a student passes or not follows the Normal model.";

$mc3_distractor2 = "The population sizes are sufficiently large.";

$mc3 = new_checkbox_multiple_choice();

$mc3 -> qa (
"Which of the following assumptions and conditions are needed to carry out the hypothesis test? CHECK ALL THAT APPLY.", 
   $mc3_answer1,
   $mc3_answer2,
);
$mc3 -> extra(
   $mc3_distractor1,
   $mc3_distractor2,
);
$mc3->makeLast("None of the above.");

######################################################
# Part 4

$c1 = @p1hat[0] - @p2hat[0];

$c2 =  -$c1;
$s =  sqrt( @p1hat[0]*(1- @p1hat[0])/100 + @p2hat[0]*(1-@p2hat[0])/100 );
$r1 = 1.96 * $s;
$r2 = 2.00 * $s;

$a1 = $c1 - $r1;
$b1 = $c1 + $r1;

$a1_dis = sprintf "%.3f", $a1;
$b1_dis = sprintf "%.3f", $b1;



$a2 = $c1 - $r2;
$b2 = $c1 + $r2;

$a2_dis = sprintf "%.3f", $a2;
$b2_dis = sprintf "%.3f", $b2;


$a3 = $c2 - $r1;
$b3 = $c2 + $r1;


$a4 = $c2 - $r2;
$b4 = $c2 + $r2;


$multians = MultiAnswer($a2, $b2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $a_stu, $b_stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if (
           (abs($a_stu - $a1) < 0.02 && abs($b_stu - $b1) < 0.02 )  ||
           (abs($a_stu - $a2) < 0.02 && abs($b_stu - $b2) < 0.02 ) ||
           (abs($a_stu - $a3) < 0.02 && abs($b_stu - $b3) < 0.02 ) ||
           (abs($a_stu - $a4) < 0.02 && abs($b_stu - $b4) < 0.02 )
          )
      {
         return [1,1];
      }
      else {
         return [0,0];  
      }

  }
);



##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT
Smart Driver Driving School has many branches across provinces. In a province different from the one in problem 1, the teaching committee at the provincial branch randomly assigned half of their 5000 students enrolled this year to receive the conventional teaching method and the remaining half to receive the new teaching method. In a random sample of 100 students who received the conventional teaching method, @m1[0] passed the road test. In another random sample of 100 students who received the new teaching method, @m2[0] passed the road test. The committee would like to use the sample data to test if the two teaching methods have different passing rates. Let \(p_C\) and \(p_N\) be the true passing rates of the conventional and the new teaching methods, respectively.

$PAR

$BBOLD Part i) $EBOLD
\{ $mc1->print_q() \}
$BR
\{ $mc1->print_a() \}

$PAR
$BBOLD Part ii) $EBOLD
\{ $mc2->print_q() \}
$BR
\{ $mc2->print_a() \}

$PAR
$BBOLD Part iii) $EBOLD
\{ $mc3->print_q() \}
$BR
\{ $mc3->print_a() \}

$PAR
$BBOLD Part iv) $EBOLD
Construct a 95$PERCENT confidence interval for the difference in the true passing rates (expressed as proportions) between the conventional and new teaching methods: 

$PAR
95$PERCENT confidence interval: ( \{ $multians->ans_rule(6) \},\{ $multians->ans_rule(6)\}) [Your two numbers must be rounded to 3 decimal places. Keep at least 6 significant figures in intermediate steps.]


END_TEXT

Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 1;


ANS( radio_cmp( $mc1->correct_ans() ) );
ANS( checkbox_cmp( $mc2->correct_ans() ) );
ANS( checkbox_cmp( $mc3->correct_ans() ) );
ANS( $multians->cmp() ); 



Context()->normalStrings;

COMMENT('MathObject version.');


## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
 (a) We are interested in testing whether the two teaching
methods give the same passing rates, against the alternative that the two
rates differ.
$BR
$BR
(b) The significance level gives the probability of a type I error. 
$BR
$BR
(c) We require conditions that ensure both independent samples within
the two groups and independence of the two samples.  
$BR
$BR
(d) We are asked to find a 95% confidence interval for the difference
in the two proportions, that is, a confidence interval for \(p_{C}-p_{N}.\) 
The interval will be centred on the estimate 
\[
\hat{p}_{C}-\hat{p}_{N}=@p1hat[0]-@p2hat[0].
\]
The estimate of the standard deviation of this estimate is 
\[
\sqrt{\frac{@p1hat[0] \left( 1-@p1hat[0]\right) }{100}+\frac{
@p2hat[0]\left( 1-@p2hat[0]\right) }{100}}.
\]
The 95% confidence interval is then approximately
\[
\left( @p1hat[0]-@p2hat[0] \right) \pm 1.96 \sqrt{\frac{@p1hat[0] \left( 1-@p1hat[0] \right) }{100}+\frac{@p2hat[0] \left( 1-@p2hat[0]\right) }{100}}=\left( $a1_dis, $b1_dis \right) .
\]

END_SOLUTION
ENDDOCUMENT(); 
