##DESCRIPTION
##  Introduction to Probability: Unions, Intersections, and Complements
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Jason Aubrey')
## DBsubject(Probability)
## DBchapter(Sample Space)
## DBsection(Odds)
## Date(Spring 2012)
## Institution(University of Missouri)
## Level(2)
## KEYWORDS('probability', 'union', 'intersection', 'complement')

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextLimitedNumeric.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Numeric");
$case = random(1,4,1);

if($case == 1){
   $string = 'in favor of rolling a total of seven when two dice are tossed';
   $ans1 = 1;
   $ans2 = 5;}
elsif($case ==2){
   $string = 'against rolling an \(8\) or higher when two dice are tossed';
   $ans1 = 21;
   $ans2 = 15;}
elsif($case == 3){
   $string = 'in favor of getting at least two tails when five fair coins are flipped';
   $ans1 = 13;
   $ans2 = 3;}
else{
   $string = 'in favor of selecting an ace or a diamond when one card is selected';
   $ans1 = 4;
   $ans2 = 9;}

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( $f1/($f1+$f2) == $f1stu/($f1stu+$f2stu)) {
          return [1,1];
      } else {
          return [0,0];
      }
  }
);
##############################################################
#
#  Text
#

Context()->texStrings;
BEGIN_TEXT
Find the odds $string.
$PAR
\{$multians->ans_rule(10)\} \(:\) \{$multians->ans_rule(10)\}
END_TEXT
Context()->normalStrings;


##############################################################
#
#  Answers
#

ANS($multians->cmp);

ENDDOCUMENT();
