##DESCRIPTION
##  
##ENDDESCRIPTION


## DBsubject(Algebra)
## DBchapter(Operations on polynomial and rational expressions)
## DBsection(Polynomials: divide)
## Date(6/23/2018)
## Institution(City Tech)
## Author(K. Andrew Parker)
## MO(1)
## KEYWORDS('algebra', 'inequality', 'fraction')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "PGML.pl",
   "PGgraders.pl",
   "niceTables.pl",
   "scaffold.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

# Give full credit if the student is able to complete the
# quotient without doing the scaffold steps...
install_problem_grader(~~&full_partial_grader);

##############################################################
#
#  Setup
#
#
Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');

$q[0] = non_zero_random(-10,10,1);
$q[1] = non_zero_random(-7,7,1);
$q[2] = random(-1,1,2)*random(2,5,1);

$quot = Formula("$q[2] x^2 + $q[1] x + $q[0]")->reduce;
$finalAns = MultiAnswer($quot, Formula(0))->with(
    checkTypes=>0,
    singleResult=>1,
    separator=>' with remainder ',
    tex_separator=>'\text{ with remainder }',
    checker=> sub{
        my ($correct, $student, $self) = @_;
        my ($quotStu, $remStu) = @{$student};
        my ($quotCor, $remCor) = @{$correct};
        $self->setMessage(1,"Your quotient is incorrect.") if ($quotStu != $quotCor);
        $self->setMessage(2,"Your remainder is incorrect.") if ($remStu != $remCor);
        $self->setMessage(2,"Your remainder should be a constant (no variables).") unless (Value::isReal($remStu));
        return [ $quotCor==$quotStu,$remCor==$remStu ];
    }
);

# divisor: (x - $rDiv)
$rDiv = random(-1,1,2)*random(2,8,1);
$divisor = Formula("x - $rDiv")->reduce;

$d[0] = -1*$rDiv*$q[0];
$int1 = -1*$rDiv*$q[1]; # for showing work
$d[1] = $int1 + $q[0];
$int2 = -1*$rDiv*$q[2]; # for showing work
$d[2] = $int2 + $q[1];
$d[3] = $q[2];

$dividend = Formula("$d[3] x^3 + $d[2] x^2 + $d[1] x + $d[0]")->reduce;

##############################################################
#
#  Text
#
#

BEGIN_PGML

    
        
            
            
>> ### Long Division ### <<

    
        
        
Let's divide [`2x^3-11x^2-41x+9`] by [`x-8`]. In this case, the _dividend_ is [`2x^3-11x^2-41x+9`] and the _divisor_ is [`x-8`]. Make sure that the polynomials are written in _descending order_. This means that we want the degrees of the terms to appear in decreasing order. For example, [`-4x^2+2x^7+x^3`] should be written as [`2x^7+x^3-4x^2`]. The initial setup should look like this:

[@
LayoutTable(
    [
        ["x-8",["2 x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["-11 x^2",cellcss=>'border-top:2px solid blue;'],["-41 x",cellcss=>'border-top:2px solid blue;'],["+9",cellcss=>'border-top:2px solid blue;']],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***

with the divisor to the left and the dividend to the right.

*Eliminating the first term*

The first step is to eliminate the first term of the dividend, [`\color{magenta}{2x^3}`].  To do this, divide the highest term of  [`\color{magenta}{2x^3}-11x^2-41x+9`] by the highest term of [`\color{grey}{x}-8`].  In other words, divide [`\color{magenta} {2 x^3}`] by [`\color{gray} {x}`]. 

>>[`\dfrac{\color{magenta} {2 x^3}}{\color{gray} {x}} = \color{magenta}{2x^2}`]<<

The result, [`\color{magenta}{2x^2}`], is the highest term of the quotient. 
Then multiply [`\color{magenta}{2x^2}`] by the divisor, [`\color{gray}{x}-8`]. 

>>[`\color{magenta} {2 x^2}(\color{gray}{x}-8) = \color{magenta}{2x^3-16x^2}`]<<


The next step is to subtract [`\color{magenta}{2x^3-16x^2}`] from the first two terms of the dividend to eliminate its highest term, [`\color{magenta}{2x^3}`]. 

>>[`\color{magenta}{2x^3}-11x^2 - (\color{magenta}{2x^3-16x^2}) =
2x^3-11x^2 -2x^3+16x^2 = \color{brown}{5x^2} `]<<

This is how our setup should look like at the end of the first step:

[@
LayoutTable(
    [
        ['','',["2 x^2",cellcss=>'color:magenta; border-bottom:2px solid blue;'],'',''],
        ["x-8",["2 x^3",cellcss=>'color:magenta; border-top:2px solid blue; border-left:2px solid blue;'],["-11 x^2",cellcss=>'border-top:2px solid blue;'],["-41 x",cellcss=>'border-top:2px solid blue;'],["+9",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( 2 x^3",cellcss=>'color:magenta; border-bottom:2px solid black;'],["-16 x^2 )",cellcss=>'color:magenta; border-bottom:2px solid black;'],'\downarrow',''],
        ['','0',["5 x^2",cellcss=>'color:brown; '],"-41 x",''],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***

*Eliminating the second term*

The second step is to eliminate [`\color{brown} {5x^2}`].  To do this, divide the highest term of  [`\color{brown} {5x^2}-41x`] by the highest term of [`\color{gray} {x}-8`].  In other words, divide [`\color{brown} {5x^2}`] by [`\color{gray} {x}`]. 

>>[`\dfrac{\color{brown} {5 x^2}}{\color{gray} {x}} = \color{brown}{5x}`]<<

The result, [`\color{brown}{5x}`], is the second highest term of the quotient. Multiply [`\color{brown}{5x}`] by the divisor, [`\color{gray}{x}-8`]. 

>>[`\color{brown} {5x}(\color{gray}{x}-8) = \color{brown}{5x^2-40x}`]<<


The next step is to subtract [`\color{brown}{5x^2-40x}`] from [`\color{brown} {5x^2}-41x`] to eliminate [`\color{brown} {5x^2}`]. 

>>[`\color{brown} {5x^2}-41x-(\color{brown}{5x^2-40x}) =  5x^2-41x-5x^2+40x = \color{green}{-x}`] <<

This is how our setup should look like at the end of the second step:

[@
LayoutTable(
    [
        ['','',["2 x^2",cellcss=>'color:magenta; border-bottom:2px solid blue;'],["+ 5 x",cellcss=>'color:brown; border-bottom:2px solid blue;'],''],
        ["x-8",["2 x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["-11 x^2",cellcss=>'border-top:2px solid blue;'],["-41 x",cellcss=>'border-top:2px solid blue;'],["+9",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( 2 x^3",cellcss=>'color:magenta; border-bottom:2px solid black;'],["-16 x^2 )",cellcss=>'color:magenta; border-bottom:2px solid black;'],'\downarrow','\downarrow'],
        ['','0',"5 x^2","-41 x",'\downarrow'],
        ['','',["- ( 5 x^2",cellcss=>'color:brown; border-bottom:2px solid black;'],["- 40 x )",cellcss=>'color:brown; border-bottom:2px solid black;'],'\downarrow'],
        ['','','0',["- x",cellcss=>'color:green;'],"+9"],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***

*Eliminating the third term*

The last step is to eliminate [`\color{green} {-x}`].  To do this, divide the highest term of  [`\color{green} {-x}+9`] by the highest term of [`\color{gray} {x}-8`].  In other words, divide [`\color{green} {-x}`] by [`\color{gray} {x}`]. 

>>[`\dfrac{\color{green} {-x}}{\color{gray} {x}} = \color{green}{-1}`]<<

The result, [`\color{green}{-1}`], is the third term of the quotient. 
Then multiply [`\color{green}{-1}`] by the divisor, [`\color{gray}{x}-8`]. 

>>[`\color{green} {-1}(\color{gray}{x}-8) = \color{green}{-x+8}`]<<


The next step is to subtract [`\color{green}{-x+8}`] from [`\color{green}{-x}+9`] to eliminate [`\color{green} {-x}`]. 

>>[`\color{green}{-x}+9-(\color{green}{-x+8}) =  -x+9+x-8 = \color{red}{1}`] <<

The result, [`\color{red}{1}`], is the remainder.
The quotient is [`\color{magenta}{2 x^2}\color{brown}{+5x}\color{green}{-1}`].

[@
LayoutTable(
    [
        ['','',["2 x^2",cellcss=>'color:magenta; border-bottom:2px solid blue;'],["+ 5 x",cellcss=>'color:brown; border-bottom:2px solid blue;'],["-1",cellcss=>'color:green; border-bottom:2px solid blue;']],
        ["x-8",["2 x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["-11 x^2",cellcss=>'border-top:2px solid blue;'],["-41 x",cellcss=>'border-top:2px solid blue;'],["+9",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( 2 x^3",cellcss=>'color:magenta; border-bottom:2px solid black;'],["-16 x^2 )",cellcss=>'color:magenta; border-bottom:2px solid black;'],'\downarrow','\downarrow'],
        ['','0',"5 x^2","-41 x",'\downarrow'],
        ['','',["- ( 5 x^2",cellcss=>'color:brown; border-bottom:2px solid black;'],["- 40 x )",cellcss=>'color:brown; border-bottom:2px solid black;'],'\downarrow'],
        ['','','0',"- x","+9"],
        ['','','',["- ( -x",cellcss=>'color:green; border-bottom:2px solid black;'],["+ 8)",cellcss=>'color:green; border-bottom:2px solid black;']],
        ['','','',["0",cellcss=>'color:red;'],["1",cellcss=>'color:red;']]
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***

The dividend, the divisor, the quotient and the remainder are related in the following way:

>>[`dividend = divisor * quotient + remainder`]<<

or
>>[`\dfrac{dividend}{divisor} = quotient + \dfrac{remainder}{divisor}.`]<<

So we have:

>>[`2x^3-11x^2-41x+8 = (x-8) * (\color{magenta}{2 x^2}\color{brown}{+5x}\color{green}{-1}) + \color{red}{1}`]<<

or 

>>[`\dfrac{2x^3-11x^2-41x+8}{x-8} =  (\color{magenta}{2 x^2}\color{brown}{+5x}\color{green}{-1}) + \dfrac{\color{red}{1}}{x-8}.`]<<
    
        
            
            
>> ### Practice ### <<

>> ### Compute the quotient and the remainder of [` \left([$dividend]\right) \div \left([$divisor]\right) `] ### <<

END_PGML

Scaffold::Begin();
Section::Begin("Long Division: Eliminating the first term");
$part1Ans = MultiAnswer(Formula("$q[2] x^2"),Formula("$d[3] x^3"),Formula("$int2 x^2"),Formula("$q[1] x^2"))->with(singleResult=>1, checker=>sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q2stu, $int3stu, $int2stu, $q1stu ) = @{$student};
      my ( $q2, $int3p, $int2p, $q1 ) = @{$correct};
      $int3stuTeX = $int3stu->TeX;
      # make sure the student's quotient * $divisor = intermediate binomial
      $self->setMessage(2,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q2stu + Formula("$rDiv * $q2stu") != $int3stu);
      $self->setMessage(3,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q2stu - Formula("x * $q2stu") != $int2stu);
      # because the product with the divisor is correct, if there's an error
      # it's in the student's selection for the quotient
      $self->setMessage(1,"\( $d[3] x^3 - $int3stuTeX \) will not be zero. <BR> You may need to try a different quotient.") if ($int3stu != $int3p && $q2stu != $q2);
      # if the remainder is incorrect, by this point, it's because they
      # subtracted incorrectly...
      $self->setMessage(4,"You have subtracted incorrectly.") if ($q1stu != Formula("$d[2] x^2 - $int2stu"));
      return [$q2==$q2stu,$int3p==$int3stu,$int2p==$int2stu,$q1==$q1stu];
});

BEGIN_PGML



[@
LayoutTable(
    [
        ['','',[ $part1Ans->ans_rule(5), noencase=>1],'',''],
        ["$divisor",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',['\\( \\large - ( \\)'.$part1Ans->ans_rule(5),noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;'],[ $part1Ans->ans_rule(5).'\\( \\large ) \\)',noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;'],'',''],
        ['','0',[ $part1Ans->ans_rule(5),noencase=>1],'',''],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
);
@]***
END_PGML

ANS($part1Ans->cmp);

Section::End();

Section::Begin("Long Division: Eliminating the second term");

$part2Ans = MultiAnswer(Formula("$q[1] x"),Formula("$q[1] x^2"),Formula("$int1 x"),Formula("$q[0] x"))->with(singleResult=>1, checker=>sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q1stu, $int2stu, $int1stu, $q0stu ) = @{$student};
      my ( $q1, $int2p, $int1p, $q0 ) = @{$correct};
      $int2stuTeX = $int2stu->TeX;
      # make sure the student's quotient * $divisor = intermediate binomial
      $self->setMessage(2,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q1stu + Formula("$rDiv * $q1stu") != $int2stu);
      $self->setMessage(3,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q1stu - Formula("x * $q1stu") != $int1stu);
      # because the product with the divisor is correct, if there's an error
      # it's in the student's selection for the quotient
      $self->setMessage(1,"\( $q[1] x^2 - $int2stuTeX \) will not be zero. <BR> You may need to try a different quotient.") if ($int2stu != $int2p && $q1stu != $q1);
      # if the remainder is incorrect, by this point, it's because they
      # subtracted incorrectly...
      $self->setMessage(4,"You have subtracted incorrectly.") if ($q0stu != Formula("$d[1] x - $int1stu"));
      return [$q1==$q1stu,$int2p==$int2stu,$int1p==$int1stu,$q0==$q0stu];
});

BEGIN_PGML



[@
LayoutTable(
    [
        ['','',"$q[2] x^2",[ $part2Ans->ans_rule(5), noencase=>1],''],
        ["$divisor",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow',''],
        ['','0',"$q[1] x^2","+$d[1] x",''],
        ['','',['\\( \\large - ( \\)'.$part2Ans->ans_rule(5),noencase=>1,cellcss=>'color:red; border-bottom:2px solid black;'],[ $part2Ans->ans_rule(5).'\\( \\large ) \\)', noencase=>1,cellcss=>'color:red; border-bottom:2px solid black;'],''],
        ['','','0',[ $part2Ans->ans_rule(5),noencase=>1],''],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
);
@]***
END_PGML

ANS($part2Ans->cmp);

Section::End();

Section::Begin("Long Division: Eliminating the third term");

$part3Ans = MultiAnswer(Formula("$q[0]"),Formula("$q[0] x"), Formula("$d[0]"),Formula("0"))->with(singleResult=>1, checker=>sub {
      my ( $correct, $student, $self ) = @_;
      my ( $q0stu, $int1stu, $int0stu, $r0stu ) = @{$student};
      my ( $q0, $int1p, $int0p, $r0 ) = @{$correct};
      $int1stuTeX = $int1stu->TeX;
      # make sure the student's quotient * $divisor = intermediate binomial
      $self->setMessage(2,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q0stu + Formula("$rDiv * $q0stu") != $int1stu);
      $self->setMessage(3,"You have incorrectly multiplied your quotient with \($divisor\).") if ($divisor*$q0stu - Formula("x * $q0stu") != $int0stu);
      # because the product with the divisor is correct, if there's an error
      # it's in the student's selection for the quotient
      $self->setMessage(1,"\( $q[0] x - $int1stuTeX \) will not be zero. <BR> You may need to try a different quotient.") if ($int1stu != $int1p && $q0stu != $q0);
      # if the remainder is incorrect, by this point, it's because they
      # subtracted incorrectly...
      $self->setMessage(4,"You have subtracted incorrectly.") if ($r0stu != Formula("$d[0] - $int0stu"));
      return [$q0==$q0stu,$int1p==$int1stu,$int0p==$int0stu,$r0==$r0stu];
});

BEGIN_PGML


[@
LayoutTable(
    [
        ['','',"$q[2] x^2","+ $q[1] x",[ $part3Ans->ans_rule(5),noencase=>1]],
        ["$divisor",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow','\downarrow'],
        ['','0',"$q[1] x^2","+$d[1] x",'\downarrow'],
        ['','',["- ( $q[1] x^2",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int1 x )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow'],
        ['','','0',"$q[0] x","+$d[0]"],
        ['','','',['\\( \\large - ( \\)'.$part3Ans->ans_rule(5), noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;'],[ $part3Ans->ans_rule(5).'\\( \\large ) \\)', noencase=>1, cellcss=>'color:red; border-bottom:2px solid black;']],
        ['','','','0',[ $part3Ans->ans_rule(5),noencase=>1]]
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
);
@]***

END_PGML

ANS($part3Ans->cmp);

Section::End();

Scaffold::End();

BEGIN_PGML

Use the steps above to compute the quotient and remainder. If you can find the quotient and remainder without using the steps above, input your answer here for full credit:

The quotient of [` \left([$dividend]\right) \div \left([$divisor]\right) `] is [_______________]{$finalAns} with remainder [_____]{$finalAns}

END_PGML

##############################################################
#
#  Answers
#
#

##############################################################
#
#  Hints
#
#



BEGIN_PGML_HINT

*Eliminating the first term*

Divide the highest term of the dividend, [`[$dividend]`], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$d[3]] x^3`] by [`x`]. The result is the highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. Enter the answer in the two middle boxes.  Next, subtract this product from [`[$d[3]] x^3 + [$d[2]] x^2`]. Enter the answer in the last box. The first term, [`[$d[3]] x^3`], must be eliminated at the end of this process.

*Eliminating the second term*

Divide the highest term of the polynomial you found in the last subtraction by the highest term of the divisor, [`[$divisor]`].   The result is the second highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. Enter the answer in the two middle boxes.  Next, subtract this product from [`[$q[1]] x^2+[$d[1]] x`]. Enter the answer in the last box. The second term, [`[$q[1]] x^2`], must be eliminated at the end of this process.

*Eliminating the third term*

Divide the highest term of the polynomial you found in the last subtraction  by the highest term of the divisor, [`[$divisor]`]. The result is the last term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. Enter the answer in the two middle boxes.  Next, subtract this product from [`[$q[0]] x+[$d[0]] `]. Enter the answer, which is the remainder of the division, in the last box. The third term, [`[$q[0]] x`], must be eliminated at the end of this process. 

END_PGML_HINT

##############################################################
#
#  Solutions
#
#





BEGIN_PGML_SOLUTION
    
*Eliminating the first term*

Divide the highest term of the dividend, [`[$dividend]`], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$d[3]] x^3`] by [`x`]. 

>>[`\dfrac{[$d[3]] x^3}{x} = [$d[3]] x^2`]<<

The result is the highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. 

>>[`[$d[3]] x^2 ([$divisor]) = [$d[3]] x^3 + [$int2] x^2`]<<

Enter the answer in the two middle boxes.  Next, subtract this product from [`[$d[3]] x^3 + [$d[2]] x^2`]. Enter the answer in the last box. 

>>[`[$d[3]] x^3 + [$d[2]] x^2-([$d[3]] x^3 + [$int2] x^2) = [$q[1]] x^2 `]<<

Note that the first (highest) term of the dividend, [`[$d[3]] x^3`], was eliminated in  this process.

[@
LayoutTable(
    [
        ['','',"$q[2] x^2",'',''],
        ["$divisor",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'',''],
        ['','0',"$q[1] x^2","",''],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***



*Eliminating the second term*

Divide the highest term of the polynomial you found in the last subtraction, [`[$q[1]] x^2`], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$q[1]] x^2`] by [`x`]. 

>>[`\dfrac{[$q[1]] x^2}{x} = [$q[1]] x`]<<

The result is the second highest term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. 

>>[`[$q[1]] x ([$divisor]) = [$q[1]] x^2 + [$int1] x`]<<

Enter the answer in the two middle boxes.  Next, subtract this product from [`[$q[1]] x^2+[$d[1]] x`]. Enter the answer in the last box. 

>>[`[$q[1]] x^2+[$d[1]] x-([$q[1]] x^2 + [$int1] x) = [$q[0]] x `]<<

Note that the second term, [`[$q[1]] x^2`], was eliminated in this process.

[@
LayoutTable(
    [
        ['','',"$q[2] x^2","+ $q[1] x",''],
        ["$divisor",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow',''],
        ['','0',"$q[1] x^2","+$d[1] x",''],
        ['','',["- ( $q[1] x^2",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int1 x )",cellcss=>'color:red; border-bottom:2px solid black;'],''],
        ['','','0',"$q[0] x",''],
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***



*Eliminating the third term*

Divide the highest term of the polynomial you found in the last subtraction, [`[$q[0]] x `], by the highest term of the divisor, [`[$divisor]`].  In other words, divide [`[$q[0]] x`] by [`x`]. 

>>[`\dfrac{[$q[0]] x}{x} = [$q[0]]`]<<

The result is the last term of the quotient. Enter it in the top box. Then multiply this result by the divisor, [`[$divisor]`]. 

>>[`[$q[0]] ([$divisor]) = [$q[0]] x + [$d[0]] `]<<

Enter the answer in the two middle boxes.  Next, subtract this product from [`[$q[0]] x+[$d[0]] `]. Enter the answer, which is the remainder of the division, in the last box. 

>>[` [$q[0]] x + [$d[0]]-([$q[0]] x + [$d[0]]) = 0 `]<<

Note that the third term, [`[$q[0]] x`], was eliminated in this process. 



[@
LayoutTable(
    [
        ['','',"$q[2] x^2","+ $q[1] x","+ $q[0]"],
        ["$divisor",["$d[3] x^3",cellcss=>'border-top:2px solid blue; border-left:2px solid blue;'],["+ $d[2] x^2",cellcss=>'border-top:2px solid blue;'],["+ $d[1] x",cellcss=>'border-top:2px solid blue;'],["+ $d[0]",cellcss=>'border-top:2px solid blue;']],
        ['',["- ( $d[3] x^3",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int2 x^2 )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow','\downarrow'],
        ['','0',"$q[1] x^2","+$d[1] x",'\downarrow'],
        ['','',["- ( $q[1] x^2",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $int1 x )",cellcss=>'color:red; border-bottom:2px solid black;'],'\downarrow'],
        ['','','0',"$q[0] x","+$d[0]"],
        ['','','',["- ( $q[0] x",cellcss=>'color:red; border-bottom:2px solid black;'],["+ $d[0] )",cellcss=>'color:red; border-bottom:2px solid black;']],
        ['','','','0','0']
    ],
    midrules=>0,
    align=>' c c c c c ',
    encase=>['\(','\)'],
)@]***


The quotient of [` \left([$dividend]\right) \div \left([$divisor]\right) `] is [`[$quot]`] with remainder [`0`].

END_PGML_SOLUTION

ENDDOCUMENT();        
