# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus', 'integral', 'differential equations', 'modeling')
## Tagged by glr 02/08/09

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Modeling with Differential Equations')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('11.6')
## Section2('11.6')
## Problem1('19')
## Problem2('23')

## Textbook tags
## HHChapter1('Differential Equations')
## HHChapter2('Differential Equations')
## HHSection1('Applications and Modeling')
## HHSection2('Applications and Modeling')

DOCUMENT();
loadMacros(
"PG.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGauxiliaryFunctions.pl",
"MathObjects.pl",
);

Context("Numeric");

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$fval = -1;
$niter = 0;
## this is not a good way of doing this.
while ( $fval < 0 && $niter < 3 ) {
    $y1 = random(0.75,0.95,0.05);
    $y1p = 100*$y1;
    $y2 = $y1 - random(0.05,0.15,0.01);
    $y2 = $y2 - 0.01 if ( 2*$y1 - $y2 == 1 );

    ## check to be sure our exponential goes the right direction
    if ( $y1 - $y2 > 1 - $y1 ) {
        $oldy2 = $y2;
        $y2 = $y1 - 0.4*(1 - $y1);
    }
    $y2p = 100*$y2;

    $k = log((1 - $y1)/($y1 - $y2));

    # $ae = -1*($y2 - $y1*$y1)/(2*$y1 - $y2 - 1);
    $a = ($y2 - $y1*$y1)/(1 + $y2 - 2*$y1);

    ## make sure that we don't get a negative final value
    $fval = $a + (1 - $a)*exp(-($k)*14);

    $niter++;
}
## if this didn't give a value that's nice, then just pick some values 
##    that work
if ( $fval < 0 ) {
    $y1 = 0.85;  $y1p = 85;
    $y2 = 0.80;  $y2p = 80;
    
    $k = log((1 - $y1)/($y1 - $y2));

    # $ae = -1*($y2 - $y1*$y1)/(2*$y1 - $y2 - 1);
    $a = ($y2 - $y1*$y1)/(1 + $y2 - 2*$y1);

    $fval = $a + (1 - $a)*exp(-($k)*14);
}

Context()->texStrings;
BEGIN_TEXT

As you know, when a course ends, students start to forget the material
they have learned.  One model (called the Ebbinghaus model) assumes
that the rate at which a student forgets material is proportional to
the difference between the material currently remembered and some
positive constant, \( a \).
$PAR
${BBOLD}A.$EBOLD
Let \( y=f(t) \) be the fraction of the original material remembered
\( t \) weeks after the course has ended.  Set up a differential
equation for \( y\), using \(k\) as any constant of proportionality
you may need (let \(k>0\)).  Your equation will contain two constants;
the constant \( a \) (also positive) is less than \( y \) for all 
\( t \). 
$BR
\(\frac{dy}{dt} = \) \{ ans_rule(35) \}
$BR
What is the initial condition for your equation?
$BR
\( y(0) = \) \{ ans_rule(15) \}

$PAR
${BBOLD}B.$EBOLD
Solve the differential equation. 
$BR
\( y = \) \{ ans_rule(35) \}

$PAR
${BBOLD}C.$EBOLD
What are the practical meaning (in terms of the amount remembered) 
of the constants in the solution \( y=f(t)\)?  If after one week the
student remembers $y1p percent of the material learned in the semester,
and after two weeks remembers $y2p percent, how much will she or he
remember after summer vacation (about 14 weeks)?
$BR
percent = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( fun_cmp( "-k*(y - a)", 'var'=>['k','y','a','t'] ) );
ANS( num_cmp( 1 ) );
ANS( fun_cmp( "(1-a)*e^(-k*t) + a", 'var'=>['k','y','a','t'] ));
ANS( Compute("$a + (1 - $a)*e^(-($k)*14)")->cmp( checker=> sub {
     my ( $correct, $student, $ansHash ) = @_;
     return $correct == $student || 100*$correct == $student; } ) );

# $aval = round(-1000*($y2 - $y1*$y1)/(2*$y1 - $y2 - 1))/1000;
# $kval = round(1000*(log(($y1 - $aval)/(1 - $aval))))/1000;
$aval = round(1000*$a)/1000;
$kval = round(1000*$k)/1000;
$twoy1 = 2*$y1;
$y2p1 = 1 + $y2;
$y1sq = $y1*$y1;
$ysum = round(1000*($a + (1 - $a)*exp(-$k*14)))/1000;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Our differential equation is 
\[ 
  \frac{dy}{dt} = -k(y-a),
\]
where \(k >0\) and \(a\) are constants.
We can solve this using separation of variables, getting 
\( \int \frac{dy}{y-a} = \int -k\,dt\), so 
\(\ln |y-a| = \ln(y-a) = -kt+C\). Thus,
\(y-a=Ae^{-kt}\) where \(A=e^C\).  

Initially nothing has been forgotten, so \(y(0)=1\). 
Therefore,  \( 1-a = Ae^0 = A\), so 
\(y - a = (1 - a) e^{-kt}\), or
\(y = ( 1 - a ) e^{-kt} + a\).

As \(t \to \infty\),  \(e^{-kt} \to 0\), so \(y\to a\).
Thus, \(a\) represents the fraction of material which is remembered in
the long run.  The constant \(k\) tells us about the rate at which
material is forgotten.  

$PAR
Then, after one week we have 
\[ $y1 = (1 - a) e^{-k} + a, \]
and after two,
\[ $y2 = (1 - a) e^{-2k} + a. \]
Solving the first of these for \(e^{-k}\), we get
\[ e^{-k} = \frac{$y1 - a}{1 - a}, \]
so that, plugging this in to the second,
\[ $y2 = (1 - a)\left(\frac{$y1 - a}{1 - a}\right)^2 + a, \]
or
\[ $y2 - a = \frac{($y1 - a)^2}{1-a}. \]
Cross-multiplying and expanding, we get
\[ a^2 - $y2p1 a + $y2 = a^2 - $twoy1 a + $y1sq, \]
so that 
\[ a = $aval. \]

$PAR
Plugging this back into the expression for \(e^{-k}\) and taking the
natural log of both sides, we get \(k = $kval\).  Thus after the
summer we have 
\[ y(14) = (1 - $aval) e^{-14 ($kval)} + $aval \approx $ysum. \]


END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
