#DESCRIPTION
##  Calculus: Second-order homogeneous linear equations
##ENDDESCRIPTION


## DBsubject(Differential equations)
## DBchapter(Higher order differential equations)
## DBsection(Linear, constant coefficients, homogeneous (distinct real roots))
## Date(06/16/2011)
## Institution(University of Minnesota)
## Author(Justin Sukiennik)
## Level(3)
## MO(1)
## TitleText1('Calculus: Concepts and Contexts')
## AuthorText1('Stewart')
## EditionText1('4 Custom UMTYMP Ed.')
## Section1('17.1')
## Problem1('1')
## KEYWORDS('calculus', 'differential equations', 'linear equations', 'second-order', 'homogeneous')

#####################################################################
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

#####################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

#####################################################################

Context("Numeric");

$a = random(2,5,1);
$b = random(6,10,1);

$c = Compute("$b-$a");
$d = Compute("$b*$a");

$ans1 = Compute("$a");
$ans2 = Compute("-$b");

#####################################################################

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);

#####################################################################

Context()->texStrings;
BEGIN_TEXT
The general solution to the second-order differential equation \(\displaystyle y'' + $c y' - $d y = 0\) is in the form \(\displaystyle y(x) = c_1 e^{r_1 x} + c_2 e^{r_2 x}.\)  Find the values of \(r_1\) and \(r_2.\) $PAR
Answer: \(r_1 = \) \{$multians->ans_rule(15) \} and \(r_2 = \) \{$multians->ans_rule(15) \}
END_TEXT
Context()->normalStrings;

#####################################################################

ANS( $multians->cmp() ); 

ENDDOCUMENT();        # This should be the last executable line in the problem.
