## DESCRIPTION
## First order ODEs: slope fields and the uniqueness theorem
## ENDDESCRIPTION

## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Equilibrium points and phase lines)
## Date(01/01/2012)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(eq_pt_06)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Elementary Differential Equations')
## AuthorText1('Kohler and Johnson')
## EditionText1('2')
## Section1('1.3')
## Problem1('10')
## KEYWORDS('differential equations','first order','slope fields','direction fields')

#  Initialization

DOCUMENT(); 

loadMacros('PGstandard.pl', 'PGML.pl', 'unionLists.pl', 'PGcourse.pl');

#  Setup

Context("Numeric")->variables->are(x=>"Real",y=>"Real");

$a = random(3,9,1);

$ans = Formula("y*($a-y)")->cmp(checker=>sub{
    my ( $correct, $student, $ansHash ) = @_;
    my $corrval = $correct->eval(x=>0,y=>1);
    my $studval = $student->eval(x=>0,y=>1);
    my $mult = Real($studval/$corrval);
    if ($mult<0){
        Value->Error("Check the sign of y'(t)");
        return 0;
    } 
    my $multStud = $mult*$correct;
    return $student == $multStud;
});



#  Main text

BEGIN_PGML
Find an autonomous differential equation [``\frac{dy}{dt}=f(x,y)``] where 

    [`f(x,y)`] is a polynomial of degree 2, and 

    the differential equation has equilibrium solutions where [`y=0`] and [`y=[$a]`], and 

    [` y^{\,\prime}(t) > 0 `] when [` 0 < y < [$a] `], and

    [` y^{\,\prime}(t) < 0 `] when [` -\infty < y < 0 `]
and when [` [$a] < y < \infty `]. 


Answer: 

    [`` \displaystyle \frac{dy}{dt} = ``]
[_]{$ans}{20}
 [@  helpLink('formulas')  @]*
END_PGML

#  Answers

$showPartialCorrectAnswers = 1;

BEGIN_PGML_SOLUTION
Equilibrium solutions occur where [`y'(t)=f(x,y)=0`].  Thus  [`f(x,y)=0`] wherever [`y=0`] or [`y=8`].  Since [`f(x,y)`] is a polynomial of degree 2 it follows that [`f(x,y)`] is a multiple of the product [`y(y-8)`]:

    [`f(x,y)=Cy(y-8)`] where [`C`] is a constant.

The other two conditions (in fact, either one of them), say that [`C`] is negative.  Thus any polynomial of the form

    [`f(x,y)=Cy(y-8)`], where [`C<0`]

is a constant, will work.
END_PGML_SOLUTION



ENDDOCUMENT();