## DESCRIPTION
##   Equation for a sphere
## ENDDESCRIPTION

## KEYWORDS('Implicit', 'Sphere')
## Tagged by nhamblet

## DBsubject('Calculus')
## DBchapter('Vectors and the Geometry of Space')
## DBsection('Three-Dimensional Coordinate Systems')
## Date('10/23/09')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('12.1')
## Section2('12.1')
## Problem1('7')
## Problem2('11')

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "PGcourse.pl",
  "parserMultiAnswer.pl",
);


TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric");
Context()->variables->add(y=>"Real",z=>"Real");

# control the sign
$s = random(-1,1,2);

# center (a,b,c)
$a = $s*random(1,5,1);
$b = $s*random(-5,-1,1);
$c = non_zero_random(-5,5,1);
# radius r
$r = random(3,6,1);

# ImplicitEquation sometimes fails to find a solution, in which case it fails outright
#$answer = ImplicitEquation("$r**2 = (x - $a)**2 + (y - $b)**2 + (z - $c)**2");


$ans1 = Formula("(x - $a)**2 + (y - $b)**2 + (z - $c)**2");
$ans2 = Formula("$r**2");

$multians = MultiAnswer($ans1, $ans2)->with(
  singleResult => 1,
  checker => sub {
     my ( $correct, $student, $self ) = @_;
     my ( $f1stu, $f2stu ) = @{$student};
     my ( $f1, $f2 ) = @{$correct};
     my $g1 = $f1 - $f2;
     my $g2 = $f2 - $f1;
     my $gstu = $f1stu - $f2stu;
     if ( ($g1 == $gstu ) || ($g2 == $gstu) ) {
          return 1;
     } else {
          return 0;
     }
  }
);

##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT

Find an equation for the sphere centered at \( ($a, $b, $c) \) with radius \( $r \).
$BR$BR
\{$multians->ans_rule(50)\} = \{$multians->ans_rule(20)\}

END_TEXT
Context()->normalStrings;

##################################################



ANS($multians->cmp());



$showPartialCorrectAnswers = 1;


COMMENT('MathObject version');
ENDDOCUMENT();
