## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Row, column, and null spaces)
## Date(July 2013)
## Institution(TCNJ and Hope College)
## Author(Paul Pearson)
## MO(1)
## Level(2)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('4.3')
## Problem1('')
## KEYWORDS('matrix' 'null space' 'column space')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixReduce.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Numeric')->variables->are(
	r => 'Real',
	s => 'Real',
	t => 'Real'
);

$r1 = non_zero_random(-5,5,1);
$r2 = non_zero_random(-5,5,1);
$r3 = non_zero_random(-5,5,1);
$r4 = non_zero_random(-5,5,1);

$s1 = non_zero_random(-5,5,1);
$s2 = non_zero_random(-5,5,1);
$s3 = random(-5,5,1);
$s4 = non_zero_random(-5,5,1);

$t1 = non_zero_random(-5,5,1);
$t2 = random(-5,5,1);
$t3 = non_zero_random(-5,5,1);
$t4 = random(-5,5,1);

$str1 = Compute(
	"$r1*r + $s1*s + $t1*t"
)->reduce;
$str2 = Compute(
	"$r2*r + $s2*s + $t2*t"
)->reduce;
$str3 = Compute(
	"$r3*r + $s3*s + $t3*t"
)->reduce;
$str4 = Compute(
	"$r4*r + $s4*s + $t4*t"
)->reduce;

Context('Matrix');

$A = Matrix([
[$r1,$s1,$t1],
[$r2,$s2,$t2],
[$r3,$s3,$t3],
[$r4,$s4,$t4],
]);

Context()->texStrings;
BEGIN_TEXT
Let \(W\) be the set of all vectors of the form
\[
\begin{bmatrix}
$str1 \\
$str2 \\
$str3 \\
$str4
\end{bmatrix}
\]
with \(r\), \(s\) and \(t\) real.
Find a matrix \(A\) such that \(W=\text{Col}(A)\).
$BR $BR  
\( A= \) \{ $A->ans_array \}.
END_TEXT
Context()->normalStrings;

ANS($A->cmp( checker => sub {
	my ($correct, $student, $ansHash) = @_;
	Context('Fraction');
	my $Ac = rcef(apply_fraction_to_matrix_entries($correct));
	my $Bc = rcef(apply_fraction_to_matrix_entries($student));
	return $Ac == $Bc;
    }
));

ENDDOCUMENT();
