## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by tda2d

## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Row, column, and null spaces)
## Date(July 2013)
## Institution(TCNJ and Hope College)
## Author(Paul Pearson)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('4.1')
## Problem1('')
## KEYWORDS('matrix' 'null space')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixReduce.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Fraction')->parens->set("[" => {formMatrix => 1});
Context()->flags->set(
  tolerance => 0.001,
  tolType => "absolute",
);

do {

$A = Matrix([
[non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
]);

$A = apply_fraction_to_matrix_entries($A);

$B = rref($A);

} until (($B->element(1,1) == 1) and ($B->element(2,2) == 1) and ($B->element(3,3) == 1));

$Zin = Matrix([0,0,0,0,0])->transpose;
$Zout = Matrix([0,0,0])->transpose;

$c1 = Fraction(-1*($B->element(1,4) + $B->element(1,5)));
$c2 = Fraction(-1*($B->element(2,4) + $B->element(2,5)));
$c3 = Fraction(-1*($B->element(3,4) + $B->element(3,5)));

$answer = Matrix([$c1,$c2,$c3,1,1])->transpose;

Context()->texStrings;
BEGIN_TEXT
Let \[ A = $A.\]
Give a non-zero vector \( \vec{x} \) in the null space of \( A \).
$BR
$BR
\( \vec{x} = \) \{ $answer->ans_array(10)\}.
END_TEXT
Context()->normalStrings;

ANS($answer->cmp(checker=>sub {
my ( $correct, $student, $ansHash ) = @_;
my $s = Matrix($student);
return 0 if $s == $Zin;
return $A * $s == $Zout;
}));
;
ENDDOCUMENT();
