# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## Tagged by glr 12/02/09

## DBsubject(Calculus - multivariable)
## DBchapter(Differentiation of multivariable functions)
## DBsection(Directional derivatives and the gradient)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(MaxRateOfChange)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('14.4')
## Problem1('69')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('6')
## Section2('14.4')
## Problem2('54')
## KEYWORDS('gradient', 'directional derivative', 'derivatives')

## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "PGcourse.pl"
);

Context("Vector2D");
Context()->flags->set( ijk=>1 );
$showPartialCorrectAnswers = 1;

$a = random(2,5,1);
$x0 = random(1,4,1);
$y0 = random(1,4,1);
$vx = non_zero_random(-3,3,1);
$vy = non_zero_random(-3,3,1);
$vec = Vector( ( $vx, $vy ) );

$f = Compute( "$a*x*y + y^2" )->reduce();

$fx = Compute( "$a*y" );
$fy = Compute( "$a*x + 2*y" );
$gr = Vector( ( $fx, $fy ) );

$gr0 = Vector( ( Compute("$a*$y0"), Compute("$a*$x0 + 2*$y0") ) );
$gr0mag = norm( $gr0 );

$fsubv = Compute( $gr0.$vec/norm($vec) );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

${BBOLD}(a)$EBOLD
What is the rate of change of \( f(x,y) = $f \) at the point
\( ($x0,$y0) \)
in the direction \( \vec v = $vec \)?
$BR
\( f_{\vec v} = \) \{ ans_rule(35) \}
$PAR

${BBOLD}(b)$EBOLD
What is the direction of maximum rate of change of
\( f \) at \( ($x0,$y0) \)?
$BR
direction = \{ ans_rule(35) \}
$BR
${BITALIC}(Give your answer as a vector.)$EITALIC
$PAR

${BBOLD}(c)$EBOLD
What is the maximum rate of change?
$BR
maximum rate of change = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($fsubv->cmp() );
ANS($gr0->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    my @c = $correct->value;
    my @s = $student->value;
    return $student->isParallel($correct) && $c[0]*$s[0] > 0;
}, showCoordinateHints => 0 ) );
ANS($gr0mag->cmp() );

$normvec = "\sqrt{" . ($vx*$vx + $vy*$vy) . "}";
$fsubvd = "\frac{" . ($a*$y0*$vx + ($a*$x0+2*$y0)*$vy) . "}{$normvec}";
$gr0magd = "\sqrt{" . ($a*$a*$y0*$y0 + ($a*$x0 + 2*$y0)**2) . "}";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

We see that \[\nabla f = $gr,\] so
at the point \(($x0,$y0)\), we have 
\[ \nabla f = $gr0.\]
$PAR

${BBOLD}(a)$EBOLD
The directional derivative is
\[ \nabla f($x0,$y0)\cdot \frac{\vec v}{|\vec v|} =
($gr0)\cdot \frac{$vec}{$normvec} = $fsubvd. \]
$PAR

${BBOLD}(b)$EBOLD
The direction of maximum rate of change is
\[\nabla f($x0,$y0) = $gr0. \]
$PAR

${BBOLD}(c)$EBOLD
The maximum rate of change is 
\[ |{\nabla f($x0,$y0)}| = $gr0magd. \]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
