# WeBWorK problem written by Alex Jordan 
# Portland Community College
# ENDDESCRIPTION



##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "pccTables.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl",
);


TEXT(beginproblem());

##############################################


Context("Numeric");
Context()->flags->set(reduceConstants=>0);

$a[0] = random(-3,-1,1);
$a[1] = $a[0] + random(1,6,1);
$a[2] = $a[1] + random(1,6,1);
$a[3] = $a[2] + random(1,6,1);
$a[4] = $a[3] - random(1,6,1);


$multiansg = MultiAnswer(String(""), Compute("$a[0]"), Compute("$a[1]"), Compute("$a[2]"), Compute("$a[3]"))->with(
  singleResult => 1,checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( @stu ) = @{$student};
      my ( @cor ) = @{$correct};
      my $return = 1;
      for my $i (0..4) {
           if ($stu[$i] != $cor[$i]) {$return = 0;}
           };
      return $return;
      }
);


$multiansh = MultiAnswer(Compute("$a[1]"), Compute("$a[2]"), Compute("$a[3]"), Compute("$a[4]"),String(""))->with(
  singleResult => 1,checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( @stu ) = @{$student};
      my ( @cor ) = @{$correct};
      my $return = 1;
      for my $i (0..4) {
           if ($stu[$i] != $cor[$i]) {$return = 0;}
           };
      return $return;
      }
);

$multiansk = MultiAnswer(Compute("$a[0]")+3, Compute("$a[1]")+3, Compute("$a[2]")+3, Compute("$a[3]")+3, Compute("$a[4]")+3)->with(
  singleResult => 1,checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( @stu ) = @{$student};
      my ( @cor ) = @{$correct};
      my $return = 1;
      for my $i (0..4) {
           if ($stu[$i] != $cor[$i]) {$return = 0;}
           };
      return $return;
      }
);

$multiansm = MultiAnswer(String(""), Compute("$a[0]")+3, Compute("$a[1]")+3, Compute("$a[2]")+3, Compute("$a[3]")+3)->with(
  singleResult => 1,checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( @stu ) = @{$student};
      my ( @cor ) = @{$correct};
      my $return = 1;
      for my $i (0..4) {
           if ($stu[$i] != $cor[$i]) {$return = 0;}
           };
      return $return;
      }
);

##############################################

BEGIN_PGML

The first table gives values for the function [`f`].


END_PGML
BEGIN_TEXT
$PAR
\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(f(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},"\($a[0]\)","\($a[1]\)","\($a[2]\)","\($a[3]\)","\($a[4]\)"]
],
allcellcss=>"padding:6pt;"
)
\}
$PAR
END_TEXT
BEGIN_PGML


Let [`g(x)=f(x-1)`], [`h(x)=f(x+1)`], [`k(x)=f(x)+3`], and [`m(x)=f(x-1)+3`].  Fill in as many values for these functions in the tables below as possible.  Some spaces are meant to be left blank, since you do not have enough information to fill in everything.  This problem will not tell you which parts you have correct, so to not give away which spaces should be left blank.

END_PGML
BEGIN_TEXT
$PAR

\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(g(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},$multiansg->ans_rule(5), $multiansg->ans_rule(5), $multiansg->ans_rule(5), $multiansg->ans_rule(5), $multiansg->ans_rule(5)]
],
allcellcss=>"padding:6pt;"
)
\}

$PAR

\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(h(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},$multiansh->ans_rule(5), $multiansh->ans_rule(5), $multiansh->ans_rule(5), $multiansh->ans_rule(5), $multiansh->ans_rule(5)]
],
allcellcss=>"padding:6pt;"
)
\}

$PAR

\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(k(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},$multiansk->ans_rule(5), $multiansk->ans_rule(5), $multiansk->ans_rule(5), $multiansk->ans_rule(5), $multiansk->ans_rule(5)]
],
allcellcss=>"padding:6pt;"
)
\}

$PAR


\{DataTable(
[
[{data=>"\(x\)",header=>'RH',rowcss=>"border-top:solid 1px;"},"\( -2 \)","\(-1\)","\(0\)","\(1\)","\(2\)"],
[{data=>"\(m(x)\)",header=>'RH',rowcss=>"border-bottom:solid 1px;"},$multiansm->ans_rule(5), $multiansm->ans_rule(5), $multiansm->ans_rule(5), $multiansm->ans_rule(5), $multiansm->ans_rule(5)]
],
allcellcss=>"padding:6pt;"
)
\}

END_TEXT


##############################################

$showPartialCorrectAnswers = 1;


ANS( $multiansg->cmp() ); 
ANS( $multiansh->cmp() ); 
ANS( $multiansk->cmp() ); 
ANS( $multiansm->cmp() );


##############################################

for my $i (0..4) {$k[$i] = $a[$i]+3;};

BEGIN_SOLUTION

\(g\) is the function \(f\) shifted \(1\) unit to the right. So the points that are represented in the table for \(f\) need to be shifted \(1\) unit right to get points that are on \(g\)'s graph:

\[\begin{aligned}
(-2,$a[0])&\to(-1,$a[0])\\
(-1,$a[1])&\to(0,$a[1])\\
(0,$a[2])&\to(1,$a[2])\\
(1,$a[3])&\to(2,$a[3])\\
(2,$a[4])&\to(3,$a[4])
\end{aligned}\]

And so we find \(g(-1)=$a[0]\), \(g(0)=$a[1]\), \(g(1)=$a[2]\), and \(g(2)=$a[3]\), giving us four values for the table for \(g\). We have no information for \(g(-2)\), so we leave that entry blank.
$PAR

\(h\) is the function \(f\) shifted \(1\) unit to the left. So the points that are represented in the table for \(f\) need to be shifted \(1\) unit left to get points that are on \(h\)'s graph:

\[\begin{aligned}
(-2,$a[0])&\to(-3,$a[0])\\
(-1,$a[1])&\to(-2,$a[1])\\
(0,$a[2])&\to(-1,$a[2])\\
(1,$a[3])&\to(0,$a[3])\\
(2,$a[4])&\to(1,$a[4])
\end{aligned}\]

And so we find \(h(-2)=$a[1]\), \(h(-1)=$a[2]\), \(h(0)=$a[3]\), and \(h(1)=$a[4]\), giving us four values for the table for \(h\). We have no information for \(h(2)\), so we leave that entry blank.
$PAR

\(k\) is the function \(f\) shifted \(3\) units up. So the points that are represented in the table for \(f\) need to be shifted \(4\) units up to get points that are on \(k\)'s graph:

\[\begin{aligned}
(-2,$a[0])&\to(-2,$k[0])\\
(-1,$a[1])&\to(-1,$k[1])\\
(0,$a[2])&\to(0,$k[2])\\
(1,$a[3])&\to(1,$k[3])\\
(2,$a[4])&\to(2,$k[4])
\end{aligned}\]

And so we find \(k(-2)=$k[0]\), \(k(-1)=$k[1]\), \(k(0)=$k[2]\), \(k(1)=$k[3]\), and \(k(2)=$k[3]\), giving us all five values for the table for \(k\). 
$PAR

\(m\) is the function \(f\) shifted \(1\) unit to the right and \(3\) units up. So the points that are represented in the table for \(f\) need to be shifted \(1\) unit right and \(3\) up to get points that are on \(m\)'s graph:

\[\begin{aligned}
(-2,$a[0])&\to(-1,$k[0])\\
(-1,$a[1])&\to(0,$k[1])\\
(0,$a[2])&\to(1,$k[2])\\
(1,$a[3])&\to(2,$k[3])\\
(2,$a[4])&\to(3,$k[4])
\end{aligned}\]

And so we find \(m(-1)=$k[0]\), \(m(0)=$k[1]\), \(m(1)=$k[2]\), and \(m(2)=$k[3]\), giving us four values for the table for \(m\). We have no information for \(m(-2)\), so we leave that entry blank.

END_SOLUTION
 

ENDDOCUMENT();