# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus','derivative')
## Tagged by glr 12/19/08

## DBsubject('Calculus')
## DBchapter('Limits and Derivatives')
## DBsection('Derivatives')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('2.4')
## Section2('2.4')
## Problem1('11')
## Problem2('13')

## Textbook tags
## HHChapter1('Key Concept: The Derivative')
## HHChapter2('Key Concept: The Derivative')
## HHSection1('Interpretations of the Derivative')
## HHSection2('Interpretations of the Derivative')


DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
"parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->strings->add( cal=>{}, lb=>{}, "lb/cal"=>{}, "cal/lb"=>{},
		         "cal/day"=>{}, "lb/day"=>{} );
$showPartialCorrectAnswers = 1;

# include javascript for pop-up explaining units
HEADER_TEXT(<<EOF);
<script language="javascript" type="text/javascript">
<!-- //
function unitspopup() {
    var url = "/webwork2_files/units.html";
    var opt = "height=625,width=600,location=no,menubar=no,status=no" +
              "resizable=yes,scrollbars=yes,toolbar=no,";
    window.open(url,'examdata_info',opt).focus();
}
// -->
</script>
EOF

$c1 = random(1500,1900,100);
$w1 = random(145,165,5);
$c2 = 2000;
$c3 = random(2100,2500,100);
$w2 = $w1 + random(5,10,1);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

A laboratory study investigating the relationship between
diet and weight in adult humans found that the weight of a subject,
\( W \), in pounds, was a function, \( W=f(c) \), of the
average number of Calories per day, \( c \), consumed by the subject.
$PAR
${BBOLD}(a)$EBOLD
In the statement \( f($c1) = $w1 \)
$BR
what are the units of $c1? \{ ans_rule(25) \}
$BR
what are the units of $w1? \{ ans_rule(25) \}
$BR
(see the \{ htmlLink("javascript:unitspopup()","units") \} cribsheet
for correct units notation)
$BR
${BITALIC}(Think about what this statement means in terms of the
weight of the subject and the number of calories that the subject
consumes.)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
In the statement \( f'($c2)=0 \),
$BR
what are the units of $c2? \{ ans_rule(25) \}
$BR
what are the units of 0? \{ ans_rule(25) \}
$BR
(see the \{ htmlLink("javascript:unitspopup()","units") \} cribsheet
for correct units notation)
$BR
${BITALIC}(Think about what this statement means in terms of the
weight of the subject and the number of calories that the subject
consumes.)$EITALIC

$PAR
${BBOLD}(c)$EBOLD
In the statement \( f^{-1}($w2) = $c3 \),
$BR
what are the units of $w2? \{ ans_rule(25) \}
$BR
what are the units of $c3? \{ ans_rule(25) \}
$BR
(see the \{ htmlLink("javascript:unitspopup()","units") \} cribsheet
for correct units notation)
$BR
${BITALIC}(Think about what this statement means in terms of the
weight of the subject and the number of calories that the subject
consumes.)$EITALIC

$PAR
${BBOLD}(d)$EBOLD
What are the units of \( f'(c)=dW/dc \)?  \{ ans_rule(25) \}
$BR
(see the \{ htmlLink("javascript:unitspopup()","units") \} cribsheet
for correct units notation)

END_TEXT
Context()->normalStrings;

ANS( String("cal")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );
ANS( String("lb")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );
ANS( String("cal")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );
ANS( String("lb/cal")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );
ANS( String("lb")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );
ANS( String("cal")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );
ANS( String("lb/cal")->cmp( checker=>sub {
    my ( $cor, $student, $ansHash ) = @_;
    return NumberWithUnits( "1", $student ) ==
	NumberWithUnits( "1", $cor );
} ) );


Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR


Since \(W=f(c)\) where \(W\) is weight in pounds and \(c\)
is the number of Calories consumed per day:
$PAR
${BBOLD}(a)$EBOLD
\( f($c1) = $w1 \) means that consuming $c1 cal each day results in a
weight of $w1 lb (the units of $c1 are cal, and $w1 lb).
$PAR
${BBOLD}(b)$EBOLD
\( f'(2000)=0 \) means that consuming 2000 cal each day causes neither
weight gain nor loss (the units of the 2000 are cal, and 0 lb/cal).
$PAR
${BBOLD}(c)$EBOLD
\( f^{-1}($w2)=$c3 \) means that a weight of $w2 lb is caused by a
consumption of $c3 cal each day (the units of $w2 are lb, and $c3 cal).
$PAR
${BBOLD}(d)$EBOLD
The units of \(dW/dc\) are lb/cal

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
