# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Polynomial and rational functions)
## DBsection(Fitting Exponentials And Polynomials To Data)
## Institution(Loyola University Chicago)
## Author(Adam Spiegler)
## TitleText1('Functions Modeling Change')
## AuthorText1('Connally')
## EditionText1('3')
## Section1('9.7')
## Problem1('6')
## TitleText2('Functions Modeling Change')
## AuthorText2('Connally')
## EditionText2('4')
## Section2('11.7')
## Problem2('6')
## KEYWORDS('best fit,'exponential regression','polynomial','power regression')

DOCUMENT();

loadMacros("PG.pl",
           "PGbasicmacros.pl",
           "PGanswermacros.pl",
           "PGstandard.pl",
           "extraAnswerEvaluators.pl",
           "PGauxiliaryFunctions.pl",
#           "PGgraphmacros.pl",
           "MathObjects.pl",
#           "PGchoicemacros.pl"
           "parserNumberWithUnits.pl",
          # "parserFormulaWithUnits.pl",
          # "parserFormulaUpToConstant.pl",
           "PGcourse.pl",
	   "AnswerFormatHelp.pl",
          );

TEXT(beginproblem());

Context("Numeric");
Context()->flags->set(useBaseTenLog=>0);        ####    log == ln
Context()->variables->are( b=>'Real' );
Context()->flags->set( tolerance=>1000000, tolType=>'absolute' );

$showPartialCorrectAnswers = 1;

#$actual_a = random(3,7,1);
#$actual_p = random(1.3,1.9,.1);


@x = (6667,960,6800,9500,1088,2733,3000,6300,1500,665);
@y = (56567,18200,110525,120100,20700,78250,58200,96400,31700,25050);
$n = 10;

for ($i = 0; $i < @y; $i++){
    $y_th[$i] = floor($y[$i]/1000);
    $y_dec[$i] = $y[$i] - $y_th[$i]*1000;
    $y_disp[$i] = "$y_th[$i],$y_dec[$i]";
         if ($y_dec[$i] < 100) { $y_disp[$i] = "$y_th[$i],0$y_dec[$i]"};
         if ($y_dec[$i] < 10) { $y_disp[$i] = "$y_th[$i],00$y_dec[$i]"};
    $brow1[$i] = $x[$i];
    $brow2[$i] = $x[$i+4];
    $Qrow1[$i] = $y_disp[$i];
    $Qrow2[$i] = $y_disp[$i+4];
}

$sum_x = 0;
$sum_y = 0;
$sum_prod = 0;
$sum_xsq = 0;
$sum_ysq = 0;

for ($i = 0; $i <@x; $i++ ) {
    $sum_x = $sum_x + log($x[$i]);
    $sum_y = $sum_y + log($y[$i]);
    $sum_prod = $sum_prod + log($x[$i])*log($y[$i]);
    $sum_xsq = $sum_xsq + log($x[$i])**2;
    $sum_ysq = $sum_ysq + log($y[$i])**2;
}

$b = ($n*$sum_prod - $sum_x * $sum_y)/($n*$sum_xsq - $sum_x**2);
$a = ($sum_y - $sum_x*$b)/$n;
$rd_b = round($b*10000)/10000;
$rd_a = round($a*10000)/10000;

#$b_prime = ($n*$sum_prod - $sum_x * $sum_y)/($n*$sum_ysq - $sum_y**2);
#if ($b >= 0) { $r = ($b*$b_prime)**(1/2)}
#        else { $r = -($b*$b_prime)**(1/2)};
#$rd_r = round($r*10000)/10000;

$initial = round(exp($rd_a)*1000)/1000;

$power_fit = "$initial b^($rd_b)";
$given = 7800;
$brain = round($initial*($given)**($rd_b));
$brain_th = floor($brain/1000);
$brain_dec = $brain - $brain_th*1000;
$brain_disp = "$brain_th,$brain_dec";
if ($brain_dec < 100) { $brain_disp = "$brain_th,0$brain_dec"; }
if ($brain_dec < 10) { $brain_disp = "$brain_th,00$brain_dec"; }

## interesting fact: generating a least-squares fit to the given data
## using Mathematica produces the fit 374.301 b^0.6264, while fitting
## the log of the data gives the result here, 221.429 b^0.685.  we might
## allow either solution here...
$power_fit_ans = Compute("$power_fit");
$power_fit_ans2 = Compute("374.3012*b^(0.626400)");
$value = NumberWithUnits( $power_fit_ans->eval( b=>$given ), 'mg' );
$value2 = NumberWithUnits( $power_fit_ans2->eval( b=>$given ), 'mg' );

#$power_coef = $rd_r;

#http://mac-huwis.lut.ac.uk/ wis/lectures/primate-adaptation/10PrimateBrains.pdf

TEXT(MODES(TeX=>"$BEGIN_ONE_COLUMN$HR",HTML=>''));

Context()->texStrings;
BEGIN_TEXT

Anthropologists suggest that the relationship between the body weight and
brain weight of primates can be modeled with a power function.  The table
below lists various body weights and the corresponding brain weights of
different primates.
$BR
$BCENTER
\{ begintable(11) \}
\{ row("\( b \), g", @x ) \}
\{ row("\( Q \), mg", @y_disp ) \}
\{ endtable() \}
$ECENTER
$PAR
(a) Using the table, find a power function that gives the brain weight,
\( Q \) (in mg), as a function of the body weight, \( b \) (in gm).
$BR
\( Q(b) = \) \{ ans_rule(20) \}

$PAR
(b) The Erythrocebus (Patas monkey) has a body weight of \( 7800 \) gm.
Estimate its brain weight.
$BR
Its brain weighs approximately \{ ans_rule(20) \}.
(include units:
\{ AnswerFormatHelp("units") \})

END_TEXT
Context()->normalStrings;

ANS( $power_fit_ans->cmp( checker=>sub { my ( $c, $s, $a ) = @_;
					 return ( $c == $s ||
						  $power_fit_ans2 == $s );
					} ) );
ANS( $value->cmp( checker=>sub { my ( $c, $s, $a ) = @_;
			        return ( $c == $s || $value2 == $s ); } ) );

#$base = $x[0]**$p;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
(a) Performing regression with a power function
\( Q(b) = $initial b^{$rd_b} \).
$PAR

(b) The estimated brain weight of the Erythrocebus is
\( Q($given)=$initial($given)^{($rd_b)} \approx $brain_disp \) mg.

END_SOLUTION
Context()->normalStrings;

TEXT(MODES(TeX=>"$HR$END_ONE_COLUMN",HTML=>''));

COMMENT('MathObject version');
ENDDOCUMENT();
