# DESCRIPTION
# Problem from Functions Modeling Change, Connally et al., 3rd ed.
# WeBWorK problem written by Adam Spiegler, <aspiegler@luc.edu>
# ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Polynomial and rational functions)
## DBsection(Fitting Exponentials And Polynomials To Data)
## Institution(Loyola University Chicago)
## Author(Adam Spiegler)
## TitleText1('Functions Modeling Change')
## AuthorText1('Connally')
## EditionText1('3')
## Section1('9.7')
## Problem1('4')
## TitleText2('Functions Modeling Change')
## AuthorText2('Connally')
## EditionText2('4')
## Section2('11.7')
## Problem2('4')
## KEYWORDS('best fit,'exponential regression','polynomial','power regression')

DOCUMENT();

loadMacros("PG.pl",
           "PGbasicmacros.pl",
           "PGanswermacros.pl",
           "PGstandard.pl",
           "extraAnswerEvaluators.pl",
           "PGauxiliaryFunctions.pl",
           "MathObjects.pl",
           "parserPopUp.pl",
	   "PGcourse.pl",
);

TEXT(beginproblem());

Context("Numeric");
Context()->flags->set(useBaseTenLog=>0);        ####    log == ln
Context()->variables->are( t=>'Real' );
Context()->flags->set( tolerance=>0.01 );

$showPartialCorrectAnswers = 1;

$n = 7;
@yr = (1970, 1975, 1980, 1985, 1990, 1995, 2000);
@t = (5, 10, 15, 20, 25, 30, 35);
#for ($i = 0; $i <= n-1; $i++) {$t[$i] = $yr[$i] - 1965};
@lob =(17, 19, 22, 20, 27, 36, 56);

$sum_x = 0;
$sum_y = 0;
$sum_prod = 0;
$sum_xsq = 0;
$sum_ysq = 0;

for ($i = 0; $i <= $n-1; $i++){
                             $sum_x = $sum_x + log($t[$i]);
                             $sum_y = $sum_y + log($lob[$i]);
                             $sum_prod = $sum_prod + log($t[$i])*log($lob[$i]);
                             $sum_xsq = $sum_xsq + log($t[$i])**2;
                             $sum_ysq = $sum_ysq + log($lob[$i])**2};

$p = ($n*$sum_prod - $sum_x * $sum_y)/($n*$sum_xsq - $sum_x**2);
$unlog_init = ($sum_y - $sum_x*$p)/$n;
$rd_p = round($p*10000)/10000;
$rd_unlog_init = round($unlog_init*10000)/10000;

$initial = round(exp($rd_unlog_init)*100)/100;
$power_fit = "$initial t^($rd_p)";

$a = 0.066;
$b = -1.515;
$c = 25.571;

$quad_fit = "$a t^2 + $b t + $c";

$popup = PopUp(["Pick better fit",
        "The power function in (a).", "The quadratic function in (b)."],
        "The quadratic function in (b).");

## again, we get a very different response when fitting against the
## actual data instead of the log data with Mathematica; for the sake
## of completeness we include the second as well
$power_fit_ans = Compute( "$power_fit" );
$power_fit_ans2 = Compute( "1.916*t^(0.8929)" );

#\footnote{Adapted from {\em The New York Times}, p.16, May 31, 2001.}
Context()->texStrings;
BEGIN_TEXT

According to the National Marine Fisheries Service, the Maine lobster
catch (in millions of pounds) has greatly increased in the past 30 years.
See the table below, where \( t \) is years since 1965
$PAR
$BCENTER
\{ begintable(8) \}
\{ row ( "Year", @yr ) \}
\{ row ( "\( t \)", @t ) \}
\{ row ( "Lobster", @lob ) \}
\{ endtable() \}
$ECENTER
$BR
Use a calculator or computer to fit the data with:
$PAR
(a) A power function of the form \( y=at^b \):
$BR
\( y = \) \{ ans_rule(20) \}
$PAR
(b) A quadratic function of the form \( y=at^2+bt+c \):
$BR
\( y = \) \{ ans_rule(20) \}
$PAR
(c) Which is a better fit?  Be sure you can explain why.
\{ $popup->menu() \}

END_TEXT
Context()->normalStrings;

ANS( $power_fit_ans->cmp( checker=>sub{ my($c,$s,$a) = @_;
				        return ( $c==$s || $power_fit_ans2 == $s); } ) );
ANS( Compute( $quad_fit )->cmp() );
ANS( $popup->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR
$BBOLD  SOLUTION $EBOLD
$PAR
(a) Calculator result: \( y=6.222t^{0.504} \). Answers may vary.
$PAR
(b) Calculator result: \( y=0.066t^2-1.515t+25.571 \).  Answers may vary.
$PAR
(c)The quadratic is the better fit. The power function is approximately the square root function so it is concave down, but the catch values are increasing rapidly in recent years. The power function goes through \( (0,0) \),
but the 1965 catch was certainly not zero. The quadratic function is shifted and stretched so is the better fit.

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
