## DESCRIPTION
## Eigenvalues and systems of differential equations
## ENDDESCRIPTION

## KEYWORDS ('linear algebra','matrix','eigenvalue','eigenvector','systems','systems of differential equations')

## DBsubject('Calculus')
## DBchapter('Differential Equations')
## DBsection('Systems of Differential Equations')
## Date('03/19/2011')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')


##################################
#  Initialization


DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGmatrixmacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"MatrixUnimodular.pl",
"unionLists.pl",
"parserRadioButtons.pl",
);

TEXT(beginproblem());


##################################
#  Setup

Context("Numeric");
Context()->variables->add(
t=>"Real",
A=>"Real",B=>"Real",
);

SRAND($psvn);


$lambda1 = random(1,5,1);
do { $lambda2 = random(4,9,1); } until ($lambda1 < $lambda2);

$radio = RadioButtons(
  ["source / unstable node","sink / stable node","saddle","center point / ellipses","spiral source","spiral sink","none of these"],
  "source / unstable node", # correct answer
  last => ["none of these"], # can be a list
);



@M = unimodular_diagonalization_GL2Z($lambda1,$lambda2);

$A = new Matrix(2,2);
$A->assign(1,1, $M[0]);
$A->assign(2,1, $M[1]);
$A->assign(1,2, $M[2]);
$A->assign(2,2, $M[3]);


$evects = MultiAnswer($M[4], $M[5], $M[6], $M[7])->with(
    singleResult => 1,
    checkTypes => 0,
    allowBlankAnswers => 1,
    format =>  "Eigenvectors = (%s, %s), (%s, %s)",
    tex_format =>  "\text{Eigenvectors }= (%s, %s), (%s, %s)" ,
    checker => sub {
        my ($correct, $student, $self ) = @_;
	my ($stuv11, $stuv12, $stuv21, $stuv22) = @{$student}[0..3];
        my $flag  = 0;
        my $flag1 = 0;
        my $flag2 = 0;
        my $counter = 0;
        if ($stuv11 && $stuv12) {
           my $v1 = new Matrix(2, 1);
           if($stuv11 != 0 || $stuv12 !=0){
             $v1-> assign(1,1, $stuv11);
             $v1-> assign(2,1, $stuv12);
             if( $A * $v1 == $lambda1 * $v1){
                $flag1=1;
             }
           }
        }
        if($stuv21 && $stuv22){
           my $v2 = new Matrix(2, 1);
           if($stuv21 != 0 || $stuv22 !=0){
             $v2-> assign(1,1, $stuv21);
             $v2-> assign(2,1, $stuv22);
             if( $A * $v2 == $lambda2 * $v2){
               $flag2=1;
             }
           }
        }
        if($flag1==1 && $flag2 ==1){
           $flag=1;
        }else{
           if($flag1 + $flag2 == 1){
              Value::Error("One of the eigenvectors is correct");
           }else{
              Value::Error("Neither of the eigenvectors is correct");
           }
        }
        return $flag; 
       
    }
);


$x1a = Compute("$M[4] A e^($lambda1 t) + $M[6] B e^($lambda2 t)");
$x2a = Compute("$M[5] A e^($lambda1 t) + $M[7] B e^($lambda2 t)");

$x1b = Compute("$M[4] B e^($lambda1 t) + $M[6] A e^($lambda2 t)");
$x2b = Compute("$M[5] B e^($lambda1 t) + $M[7] A e^($lambda2 t)");


$gensolution = MultiAnswer($x1a, $x2a)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $x1stu, $x2stu ) = @{$student};
      if ( ($x1a == $x1stu && $x2a == $x2stu) ||
           ($x1b == $x1stu && $x2b == $x2stu) ) {
          return [1,1];
      } elsif (($x1a == $x1stu || $x1b == $x1stu)) {
          return [1,0];
      } elsif (($x2a == $x2stu || $x2b == $x2stu)) {
          return [0,1];
      } else {
          return [0,0];
      }
  }
);




#######################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

$ITEM
Find the eigenvalues and corresponding eigenvectors 
for the matrix 
\(
\displaystyle
P =
\left\lbrack
\begin{array}{rr}
$M[0] & $M[2] \\
$M[1] & $M[3]
\end{array} 
\right\rbrack.
\)
$BR
$BR
The smaller eigenvalue is \( \lambda_1 = \) \{ans_rule(10) \},
and the larger eigenvalue is \( \lambda_2 = \) \{ ans_rule(10) \}.
$BR
The corresponding eigenvectors are
\{ 
mbox(
[
'\( \vec{v}_1 =  \) ', 
display_matrix([[$evects->ans_rule(5)],[$evects->ans_rule(5)]]),
' and \( \vec{v}_2 =  \) ', 
display_matrix([[$evects->ans_rule(5)],[$evects->ans_rule(5)]]),
],
allowbreaks=>'yes'
) 
\}

$ITEMSEP
$ITEM
Find the most general solution to the linear system of 
differential equations \( \vec{x}^{\, \prime} = P \vec{x} \). 
Use \( A \) and \( B \) to denote arbitrary constants.
$BR
$BR
\( x_1(t) = \) \{ $gensolution->ans_rule(40) \}
$BR
\( x_2(t) = \) \{ $gensolution->ans_rule(40) \}

$ITEMSEP
$ITEM The behavior of this system is best described as
$BR
\{ $radio->buttons() \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


###################################
#  Answer evaluation

ANS( Real($lambda1)->cmp() );
ANS( Real($lambda2)->cmp() );
ANS( $evects->cmp() );
ANS( $gensolution->cmp() );
ANS( $radio->cmp() );

COMMENT('MathObject version.  Distinct positive integer eigenvalues.');

ENDDOCUMENT();