# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/12/09


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Planes)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(planes_n_given_eqn)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('13.3')
## Problem1('13')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('6')
## Section2('13.Review')
## Problem2('21')
## KEYWORDS('vectors', 'dot product')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')

Context("Vector");
Context()->flags->set( ijk=>1 );

Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(2,8,1);
$c = random(2,8,1);
$d = random(2,8,1);

$n = Vector( (-$a, $c, 1) );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a normal vector \(\vec n\) to the plane
\( z - $a (x - $b) = $c($d - y) \)

$PAR
\(\vec n = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($n->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student->isParallel($correct);
  } ) );

$ab = $a * $b;
$cd = $c * $d;
$cdmab = $cd - $ab;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The equation can be rewritten as
\[ z - $a x + $ab = $cd - $c y, \]
or
\[ -$a x + $c y + z = $cdmab. \]
so \(\vec n=-$a\vec i + $c\vec j + \vec k\) (or any
vector parallel to that, of course).

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
