# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/12/09


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Planes)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(planes_n_given_eqn)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('5')
## EditionText2('6')
## EditionText3('7')
## Section1('13.3')
## Section2('13.3')
## Section3('13.3')
## Problem1('17')
## Problem2('31')
## Problem3('39')
## KEYWORDS('vectors', 'dot product')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')

Context("Vector");
Context()->flags->set( ijk=>1 );

Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
  k => {TeX => "\,\mathit{\vec k}"},
);

$showPartialCorrectAnswers = 1;

$a = random(2,8,1);
$b = random(2,8,1);
$b++ if ( $a == $b );

$n = Vector( "<$a, $b, -1>" );
$v = Vector( "<$b, -$a, 0>" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

${BBOLD}(a)$EBOLD
Find a vector \(\vec n\) perpendicular to the plane
\[ z = $a x + $b y. \]
\( \vec n = \) \{ ans_rule(35) \}

$PAR
${BBOLD}(b)$EBOLD
Find a vector \(\vec v\) parallel to the plane.
$BR
\( \vec v = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($n->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student->isParallel($correct);
  } ) );

ANS($v->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $student.$n == 0;
  } ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
Writing the plane in the form
\( $a x + $b y - z = 0\) shows that a normal vector is
\[ \vec n = $a \vec i + $b \vec j - \vec k.\]
Any multiple of this vector is also a correct answer.

$PAR
${BBOLD}(b)$EBOLD
Any vector perpendicular to \(\vec n\) is parallel to the plane, so one
possible answer is
\[\vec v= $b\vec i - $a\vec j.\]
Many other answers are possible.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
