# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

## Tagged by glr 06/12/09


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Dot product, length, and unit vectors)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('13.3')
## Problem1('53')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('6')
## Section2('13.3')
## Problem2('56')
## KEYWORDS('vectors', 'dot product')

## Textbook tags
## HHChapter1('A Fundamental Tool: Vectors')

Context("Vector");
Context()->flags->set( ijk=>1 );

Context()->constants->set(
  i => {TeX => "\mathit{\vec i}"},
  j => {TeX => "\mathit{\vec j}"},
  k => {TeX => "\mathit{\vec k}"},
);
$showPartialCorrectAnswers = 1;

# some nice pythagorean triples
@ptrips = ( [ 3, 4, 5 ],   [ 5, 12, 13 ], [ 8, 15, 17 ], [ 9, 40, 41 ] );

$uInd = random(0,2,1);
$vInd = $uInd + random(1,@ptrips-$uInd-1,1);

$uOrd = list_random(0,1);
$vOrd = list_random(0,1);

@u = $uOrd ? ( $ptrips[$uInd]->[1], $ptrips[$uInd]->[0] ) :
    ( $ptrips[$uInd]->[0], $ptrips[$uInd]->[1] );
push( @u, $ptrips[$uInd]->[2] );
@v = $vOrd ? ( $ptrips[$vInd]->[1], $ptrips[$vInd]->[0] ) :
    ( $ptrips[$vInd]->[0], $ptrips[$vInd]->[1] );
push( @v, $ptrips[$vInd]->[2] );

## when we solve, we get the vector
##    a <(u2 v1 - u1 v2), (u0 v2 - u2 v0)>,
## where a is positive or negative depending on whether this dotted
##    with u is greater than or less than 0.
# if ( $u[0]*(-1*$u[2]*$v[1]-$u[1]*$v[2]) + $u[1]*($u[0]*$v[2]-$u[2]*$v[0]) > 0 ) {
#     $vec = Vector(( (-1*$u[2]*$v[1]-$u[1]*$v[2]), ($u[0]*$v[2]-$u[2]*$v[0]) ));
# } else {
#     $vec = Vector(( ($u[2]*$v[1]+$u[1]*$v[2]), -1*($u[0]*$v[2]-$u[2]*$v[0]) ));
# }
## simpler solution: take u/|u| + v/|v|
$umag = $ptrips[$uInd]->[2];
$vmag = $ptrips[$vInd]->[2];
$vec = Vector(( ($u[0]/$umag + $v[0]/$vmag), ($u[1]/$umag - $v[1]/$vmag) ));

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT

Find a vector \(\vec w\) that bisects the smaller of the two angles
formed by \( $u[0]\,\vec i + $u[1]\,\vec j \) and
\( $v[0]\,\vec i - $v[1]\,\vec j \).

$PAR
\( \vec w = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( $vec->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    # return $student->isParallel($correct);
    return unit($correct) == unit($student);
} ) );

## this is kludgy.  there should be a routine to do this
##    for us.
sub redfrac {
    my ( $n, $d ) = @_;
    ( $n, $d ) = reduce( $n, $d );

    if ( $n == 1 && $d == 1 ) {
        return 1;
    } elsif ( $n == 1 ) {
	return "\frac{1}{$d}";
    } elsif ( $d == 1 ) {
        return "$n";
    } else {
	return "\frac{$n}{$d}";
    }
}

$xyc = redfrac( -1*($u[2]*$v[1] + $u[1]*$v[2]), ($u[0]*$v[2] - $u[2]*$v[0]) );
$axc = redfrac( ($u[2]*$v[1] + $u[1]*$v[2]), ($u[0]*$v[2] - $u[2]*$v[0]) );
$udw = "($u[0])($xyc) + $u[1]";
$udwr = round(100*(-1*$u[0]*($u[2]*$v[1] + $u[1]*$v[2])/($u[0]*$v[2] - $u[2]*$v[0])+$u[1]))/100;

if ( -1*$u[0]*($u[2]*$v[1] + $u[1]*$v[2])/($u[0]*$v[2] - $u[2]*$v[0]) + $u[1] > 0 ) {
    $s = ">";
    $ycond = "y > 0";
    $c2 = $xyc;
    $s3 = "+";
} else {
    $s = "<";
    $ycond = "y < 0";
    $c2 = $axc;
    $s3 = "-";
}

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Let
\(\vec u = $u[0]\,\vec i + $u[1]\,\vec j\) and
\(\vec v = $v[0]\,\vec i - $v[1]\,\vec j \).
We seek a vector \(\vec w = x\,\vec i + y\,\vec j\) such
that the cosine of the angle between \(\vec u\) and \(\vec w\)
equals the cosine of the angle between \(\vec v\) and
\(\vec w\).  Thus
\[ {\vec u \cdot \vec w\over \Vert \vec u \Vert \Vert \vec w \Vert} =
   {\vec v \cdot \vec w\over \Vert \vec v \Vert \Vert \vec w \Vert} \]
or
\[ {$u[0] x + $u[1] y\over $u[2]\sqrt{x^2+y^2}} =
   {$v[0] x - $v[1] y\over $v[2]\sqrt{x^2+y^2}}. \]
Simplifying, we have \( x = $xyc y\).

The vector we want is of the form
\(\vec w = $xyc\,y\,\vec i + y\,\vec j\),
but should we take \(y>0\) or \(y<0\)?
The smaller of the two angles formed by \(\vec u\) and \(\vec v\) is
between \(0^\circ\) and \(180^\circ\), and so \(\vec w\) must make an acute
angle with \(\vec u\) and \(\vec v\).  This will be the case when
\(\vec u\cdot \vec w > 0\), so we want to pick \(y\) so that this is
the case.  If \(y>0\) then
\(\vec u\cdot \vec w = ($udw) y = $udwr y $s 0\).
Therefore we want to take \($ycond\).

Thus \(\vec w\) can be any positive multiple of the vector
\($c2 \vec i $s3 \vec j\).

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
