# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('line integral', 'integrals', 'parametric curve', 'multivariable', 'calculus')
## Tagged by glr 07/07/10


## DBsubject('Calculus')
## DBchapter('Vector Calculus')
## DBsection('Line Integrals')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## EditionText1('5')
## AuthorText1('Hughes-Hallett')
## Section1('18.2')
## Problem1('31')

## Textbook tags
## HHChapter1('Line Integrals')
## HHSection1('Computing Line Integrals over Parameterized Curves')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->are( t=>'Real' );
$showPartialCorrectAnswers = 1;

$a = random(2,5,1);
$twoa = 2*$a;

$mag = Formula( "$a" );

$mc = MultiAnswer( Compute("cos(t)"), Compute("sin(t)"), Compute("0"),
		   Compute("2*pi"), $mag, Compute("-sin(t)"),
		   Compute("cos(t)"),
		   Compute("-$a*sin(t)"), Compute("$a*cos(t)") )->with(
    singleResult=>0,
    checker=>sub {
	my ( $c, $s, $ans ) = @_;

	## the parameterization is correct if its magnitude is one,
	##    the functions aren't constant, and we start and end at
	##    the same place.  we really only want to go around once,
	##    too.
	my ( $sx, $sy, $t0, $t1, $sm, $vx, $vy, $sfx, $sfy ) = @{$s};

	my $sx0 = $sx->eval( t=>$t0 );
	my $sx1 = $sx->eval( t=>$t1 );
	my $sy0 = $sy->eval( t=>$t0 );
	my $sy1 = $sy->eval( t=>$t1 );

	my @res = ( 0, 0, 0, 0, 0, 0, 0, 0, 0 );
	$res[4] = $mag == $sm ? 1 : 0;

	if ( Formula("1") == $sx*$sx + $sy*$sy && $t0 != $t1 &&
	     atan2( $sy0, $sx0 ) == atan2( $sy1, $sx1 ) ) {

	    $res[0] = 1;
	    $res[1] = 1;

	    my $dt = ($t1 - $t0)/20;
	    my $theta0 = atan2( $sy0, $sx0 );
	    my $dtheta = abs( $theta0 -
			      atan2( $sy->eval( t=>($t0+$dt) ),
				     $sx->eval( t=>($t0+$dt) ) ) );
	    my $isErr = 0;
	    for ( my $i=2; $i<18; $i++ ) {
		## make sure that we're not past the starting point
		my $xtmp = $sx->eval( t=>($t0 + $i*$dt) );
		my $ytmp = $sy->eval( t=>($t0 + $i*$dt) );
		my $theta1 = atan2( $ytmp, $xtmp );
		if ( abs( $theta0 - $theta1 ) < $dtheta ) {
		    $ans->setMessage(3, 'Make sure that your ' .
				     'parameteriation traverses the ' .
				     'curve only once.');
		    $isErr = 1;
		    break;
		}
	    }
	    if ( ! $isErr ) {
		$res[2] = 1;
		$res[3] = 1;
	    }

	    $res[5] = $sx->D('t') == $vx ? 1 : 0;
	    $res[6] = $sy->D('t') == $vy ? 1 : 0;

	    $res[7] = Compute("-$a*$sy") == $sfx ? 1 : 0;
	    $res[8] = Compute("$a*$sx") == $sfy ? 1 : 0;
	}
	return [ @res ];
    }
);

$int = Compute( "2*$a*pi" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT
$BBOLD Note: In this problem you will not see if partial answers are correct until you have entered all parts except the line integral's value. $EBOLD
$PAR
Let \( \vec F = -$a y\vec i + $a x\vec j \) and let \( C \) be the
unit circle oriented counterclockwise, parameterized with the parameter
\(t\).
$PAR
${BBOLD}(a)$EBOLD
Parameterize \(C\).
$BR
\( x(t) = \) \{ $mc->ans_rule(35) \}
$BR
\( y(t) = \) \{ $mc->ans_rule(35) \}
$BR
with \{ $mc->ans_rule(5) \} \(\le t\le\) \{ $mc->ans_rule(5) \}

$PAR
${BBOLD}(b)$EBOLD
Find \( \Vert\vec{F}\Vert \) on \(C\):
$BR
\( \Vert\vec{F}\Vert = \) \{ ans_rule(35) \}

$PAR
${BBOLD}(c)$EBOLD
With the parameterization in (a), what is a vector \(\vec v(t)\) that is
tangent to the circle \(C\)?
$BR
\(\vec v(t) = \) \{ $mc->ans_rule(25) \} \(\vec i + \)
\{ $mc->ans_rule(25) \} \(\vec j\)

$PAR
${BBOLD}(d)$EBOLD
With the same parameterization, what is \(\vec F\) on \(C\)?
$BR
\(\vec F = \)  \{ $mc->ans_rule(25) \} \(\vec i + \)
\{ $mc->ans_rule(25) \} \(\vec j\)
$PAR
(Note that your results from (c) and (d) show that \(\vec F\)
is tangent to \(C\) at all points.)

$PAR
${BBOLD}(e)$EBOLD
Find \(\int_C\vec F\cdot d\vec r\).
$BR
\(\int_C\vec F\cdot d\vec r = \) \{ ans_rule(35) \}
$BR
(Note how is this related to the length of the curve \(C\) (that
is, the circumference of the circle).)           

END_TEXT
Context()->normalStrings;

ANS($mc->cmp() );
ANS($int->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The unit circle centered at the origin has equation \(x^2+y^2=1\),
which we may parameterize with \(x = \cos(t)\), \(y = \sin(t)\), with
\(0\le t\le 2\pi\).

$PAR
${BBOLD}(b)$EBOLD
At any point in the plane, the magnitude of \(\vec{F}\) is given by
\(\Vert{\vec F}\Vert=\sqrt{(-$a y)^2+ ($a x)^2} = 5\sqrt{x^2 + y^2}\).
Thus, along the unit circle,
\(\Vert{\vec F}\Vert=$a\).  We can also see this by using the parameterization
in ${BBOLD}(a)$EBOLD: on \(C\),
\(\vec F = -$a\sin(t)\,\vec i + $a\cos(t)\,\vec j\), so that
\(\Vert \vec F \Vert = $a\).

$PAR
${BBOLD}(c)$EBOLD
Suppose \(\vec r=x\,\vec i + y\,\vec j\) is a radius vector to a
point \((x,y)\) on the unit circle centered
at the origin.  Then
\(\vec r = \cos(t)\,\vec i + \sin(t)\,\vec j\), so that a tangent to the
circle is
\[ \vec r\,'(t) = -\sin(t)\,\vec i + \cos(t)\,\vec j.\]

$PAR
${BBOLD}(d)$EBOLD
We have \(\vec F = -$a\sin(t)\,\vec i + $a\cos(t)\,\vec j\), so that
\(\vec F\) is tangent to the circle at every point.

$PAR
${BBOLD}(e)$EBOLD
We have
\[ \int_C\vec F\cdot d\vec r = \int_C\vec F\cdot \vec r\,'\,dt
  = \int_0^{2\pi} (-$a\sin(t)\,\vec i + $a\cos(t)\,\vec j)\cdot
    (-\sin(t)\,\vec i + \cos(t)\,\vec j)\,dt
\]
\[
  = \int_0^{2\pi} $a\,dt = $twoa\pi,
\]
which is \((\Vert \vec F\Vert)(\mbox{length of }C)\).

END_SOLUTION
Context()->normalStrings;

ENDDOCUMENT();
