##DESCRIPTION
##  Polynomials: Multiplicity of zeroes
##ENDDESCRIPTION


## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Polynomial equations and functions)
## DBsection(Zeros and multiplicities)
## Date(Spring 2012)
## Institution(University of Missouri)
## MLT(ListZerosMultiplicitesOfFactoredPoly)
## Level(2)
## TitleText1('College Algebra')
## AuthorText1('Beecher, Penna, Bittinger')
## EditionText1('4')
## Section1('4.1')
## Problem1('')
## KEYWORDS('algebra', 'polynomials', 'zeroes', 'multiplicities')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

#######################################
#  Set-up

$b = random(2,10,1);
$b2 = ($b)**2;
do {$n = random(2,9,1)} until ($b != $n);
do {$m = random(2,12,1)} until ($m != $b && $m != $n);
$k = $m + $n;

$f = "(x - $b)^{$m} (x^2- $b2)^{$n}";

## Answers
$ans = List($b,-$b);
$multians = MultiAnswer($b, $k, -$b, $n)->with(
  singleResult => 0,
  allowBlankAnswers => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu, $f3stu, $f4stu) = @{$student};
      my ( $f1, $f2, $f3, $f4) = @{$correct};
      if ( 
(   ($f1stu == $f1 && $f2stu == $f2)  && ($f3stu == $f3 && $f4stu == $f4)       ||
(    $f1stu == $f3 && $f2stu == $f4)  && ($f3stu == $f1 && $f4stu == $f2 )    ) )
{return [1,1,1,1];}  

else { return [0,0,0,0];  } 

} );



################################################
#  Main text

BEGIN_TEXT
Find the zeroes and multiplicities of the polynomial \(f(x) = $f\).
$PAR
$SPACE The zeroes are \( x = \) \{ans_rule(20)\} $SPACE $SPACE (separate your answers by commas).
$PAR
$SPACE The zero \(x = \)  \{$multians->ans_rule(5)\}  has multiplicity \{ $multians->ans_rule(5)\}.
$BR
$SPACE The zero \(x = \)  \{$multians->ans_rule(5)\}  has multiplicity \{ $multians->ans_rule(5)\}.

END_TEXT
Context()->normalStrings;




################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( Compute($ans)->cmp() );
ANS( Compute($multians)->cmp() );


COMMENT('MathObject version');
ENDDOCUMENT();
