## This problem is a modification of the following problem
## Library/PCC/BasicAlgebra/Exponents/exponentsMultiplication10.pg
## Modifications were made by R. Noonan Heale and A. Schoening

## ORIGINAL PROBLEM DESCRIPTION:


# WeBWorK problem written by Chris Hughes, 2013
# Portland Community College
#
# Template:
# Simplify the following expression
#     x^m * x^n
#
# We use an INTEGER value for x on the interval [2,20]
#
# This problem chooes m and n to be POSITIVE
#
# Last edited: Carl Yao 6/28/13
#
# ENDDESCRIPTION

## DBCCSS('8.EE.1')
## DBsubject(Arithmetic)
## DBchapter(Integers)
## DBsection(Exponentiation)
## Institution(PCC)
## Author(Alex Jordan, Carl Yao, Chris Hughes)
## Level(2)
## MO(1)
## KEYWORDS('multiply','exponent','simplify')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "answerHints.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
Context()->flags->set(formatStudentAnswer=>'parsed',reduceConstants=>0); 
Context()->operators->undefine('*','* ', ' *', '/', '/ ', ' /' , '//','+','-');     
Context()->operators->set(
   '^' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef},
   '**' => {class => 'bizarro::BOP::power', isCommand => 1, perl=>undef});

# m and n are the exponents
$mybase=random(2,9,1);
$m=random(2,10,1);
$n=random(2,10,1);

$total = $m+$n;
$left = Formula("$mybase^($m)");
$right = Formula("$mybase^($n)");
$ans = Formula("$mybase^($total)");





##############################################
TEXT(beginproblem());
BEGIN_PGML
Use the properties of exponents to simplify the following
   
    [`[$left]\cdot[$right] = `]
    [____________]


Do not use a calculator.  Write answer in exponential notation, rather than multiplying it out.  

END_PGML

$wrong=$m*$n;
ANS($ans -> cmp(checker=>sub{
my ( $correct, $student, $ansHash ) = @_;
return 0 if $ansHash->{isPreview} || $correct != $student;
$student = $ansHash->{student_formula};
$correct = $correct->{original_formula} if defined $correct->{original_formula};
$student = Formula("$student"); $correct = Formula("$correct");
return 0 unless ($correct == $student);
Context()->flags->set(bizarroPow=>1);
delete $correct->{test_values}, $student->{test_values};
Value->Error("Simplify according to the rules of exponents, not in any other way") unless (($correct == $student) or ($student == $correct));
Context()->flags->set(bizarroPow=>0);
return 1;
}) ->
  withPostFilter(AnswerHints(
    [Compute("$mybase^($wrong)")] =>
    "When multiplying terms with the same base, you do not multiply the exponents.")));


BEGIN_PGML_SOLUTION
We _add_ the exponents as follows

    [`\begin{aligned}
           [$mybase]^{[$m]}\cdot [$mybase]^{[$n]}&=[$mybase]^{[$m]+[$n]}\\
                             &=[$mybase]^{[$total]}
       \end{aligned}`]

END_PGML_SOLUTION
##############################################

ENDDOCUMENT();
