## DESCRIPTION
## The geometry of vectors in R^n
## ENDDESCRIPTION


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Vectors and vector arithmetic)
## Date(01/01/2013)
## Institution(Hope College)
## Author(Paul Pearson)
## MLT(GiveTwoPointsForVec)
## Level(1)
## TitleText1('Multivariable Mathematics')
## AuthorText1('Darin Stephenson')
## EditionText1('Version 3.0')
## Section1('1.1')
## Problem1('13')
## KEYWORDS('vectors')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("Vector2D");
Context()->flags->set( ijk=>0 );

$A = non_zero_point3D(-9,9,1);
do { $B = non_zero_point3D(-9,9,1); } until ($A != $B);
$AB = Vector($B - $A);

$multians = MultiAnswer($A, $B)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $Astu, $Bstu ) = @{$student};
      my ( $Acor, $Bcor ) = @{$correct};
      if ( $Astu == Point(0,0,0) || $Bstu == Point(0,0,0) ) { 
           return 0;
      }
      if ( Vector($Bstu - $Astu) == $AB ) {
          return 1;
      } else {
           return 0;
      }
  }
);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Find a representation of the vector \( \overrightarrow{AB} = $AB \)
in \( \mathbb{R}^3 \) by giving appropriate values for the points
\( A \) and \( B \) such that neither \( A \) nor \( B \) is the origin.
$BR
$BR
\( A = \) \{ $multians->ans_rule(20) \} 
\{ AnswerFormatHelp("points") \}
$BR
\( B = \) \{ $multians->ans_rule(20) \} 
\{ AnswerFormatHelp("points") \}
END_TEXT
Context()->normalStrings;


########################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );



ENDDOCUMENT();
