## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Integration by parts (without trigonometric functions))
## MLT(notrig_04)
## MLTleader(1)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('APEX Calculus')
## AuthorText1('Hartman')
## EditionText1('3.0')
## Section1('6.2')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserWordCompletion.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());

Context("Numeric");
Context()->variables->are(
"x"   => "Real",
"dx"  => "Real",
);

$ans1u = Compute("ln(x)");
$ans1dv = Compute("dx"); 

$multians1 = MultiAnswer( $ans1u, $ans1dv )->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      if ( ($c[0] == $s[0]) && ($c[1] == $s[1]) ) { return 1; } else { return 0; }
  }
);


$ans2u = Compute("ln(x)");
$ans2dv = Compute("x dx"); 

$multians2 = MultiAnswer( $ans2u, $ans2dv )->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      if ( ($c[0] == $s[0]) && ($c[1] == $s[1]) ) { return 1; } else { return 0; }
  }
);


$ans3u = Compute("x^2");
$ans3dv = Compute("e^x dx"); 

$multians3 = MultiAnswer( $ans3u, $ans3dv )->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      if ( ($c[0] == $s[0]) && ($c[1] == $s[1]) ) { return 1; } else { return 0; }
  }
);


BEGIN_PGML
## Reading questions on integration by parts

a. When using integration by parts to evaluate [``\int \ln(x) \, dx``] it works best to choose [`u =`] [____]{$multians1} and [`dv = `] [____]{$multians1}.

b. When using integration by parts to evaluate [``\int x \ln(x) \, dx``] it works best to choose [`u =`] [____]{$multians2} and [`dv = `] [____]{$multians2}.

c. When using integration by parts to evaluate [``\int x^2 e^x \, dx``] it works best to choose [`u =`] [____]{$multians3} and [`dv = `] [____]{$multians3}.
END_PGML

ENDDOCUMENT();
