## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear independence)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix');

$A = Matrix([
[non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
[non_zero_random(-5,5,1),non_zero_random(-5,5,1)],
]);

$x = Matrix([
[random(-5,-1,1)],
[random(2,5,1)],
]);

$x1 = $x->element(1,1);
$x2 = $x->element(2,1);

$b = $A * $x;

$M = Matrix([
[$A->element(1,1), $A->element(1,2), $b->element(1,1)],
[$A->element(2,1), $A->element(2,2), $b->element(2,1)],
[$A->element(3,1), $A->element(3,2), $b->element(3,1)],
]);

$answer = Matrix([
[$x1],
[$x2],
[-1],
]);

Context()->texStrings;
BEGIN_TEXT
Let \[ A = $M. \]
Observe that the third column of \( A \) is \( $x1 \) times the first column 
plus \( $x2 \) times the second column.  Use this information to find a 
nontrivial solution to \( A \mathbf{x} = \mathbf{0} \).
$BR
$BR
\( \boldsymbol{x} = \) \{ $answer->ans_array(10) \}
END_TEXT
Context()->normalStrings;

ANS($answer->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $S = Matrix($student);
      my $Z = Matrix([0,0,0])->transpose;
      return 0 if $S == $Z;
      return $M * $S == $Z;
}));


ENDDOCUMENT();
