## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear independence)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix');

$A = Matrix([
[1,1,0],
[0,0,1],
[0,0,0],
[0,0,0],
]);

Context()->texStrings;
BEGIN_TEXT
Construct a \( 4 \times 3 \) matrix \( A = \lbrack \ \vec{a}_1 \ \ \vec{a}_2 \ \ \vec{a}_3 \ \rbrack \)
such that \( \lbrace \vec{a}_1, \vec{a}_2 \rbrace \) is a linearly dependent 
set and \( \vec{a}_3 \) is not in the span of \( \lbrace \vec{a}_1, \vec{a}_2 \rbrace \).
$BR
$BR
\( A = \) \{ $A->ans_array(10) \}
END_TEXT
Context()->normalStrings;

ANS($A->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $S = Matrix($student);
      my $S12 = Matrix([
      [$S->element(1,1),$S->element(1,2)],
      [$S->element(2,1),$S->element(2,2)],
      [$S->element(3,1),$S->element(3,2)],
      [$S->element(4,1),$S->element(4,2)],
      ]);
      my $S13 = Matrix([
      [$S->element(1,1),$S->element(1,3)],
      [$S->element(2,1),$S->element(2,3)],
      [$S->element(3,1),$S->element(3,3)],
      [$S->element(4,1),$S->element(4,3)],
      ]);
      my $S23 = Matrix([
      [$S->element(1,2),$S->element(1,3)],
      [$S->element(2,2),$S->element(2,3)],
      [$S->element(3,2),$S->element(3,3)],
      [$S->element(4,2),$S->element(4,3)],
      ]);
      return ( 
      ( (($S12->transpose)*$S12)->det == Real(0) ) and 
           ( ( (($S13->transpose)*$S13)->det != Real(0) ) or 
             ( (($S23->transpose)*$S23)->det != Real(0) ) 
           ) 
      );
}));


ENDDOCUMENT();
