## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Abstract vector spaces)
## DBsection(Linear independence)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "MatrixUnits.pl",
  "parserMultiAnswer.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);
install_problem_grader(~~&std_problem_grader);
$showPartialCorrectAnswers = 0;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix')->variables->are('x'=>'Real');
Context()->noreduce('(-x)-y');

$U1 = GL3Z();
$U2 = GL3Z();

$A = $U1**2 * $U2;

$A1 = $A->column(1);
$A2 = $A->column(2);
$A3 = $A->column(3);

foreach my $i (1..3) {
  foreach my $j (1..3) {
    $a[$i][$j] = $A->element($i,$j);
  }
}

$p1 = Formula("$a[1][1] + $a[2][1] * x + $a[3][1] * x^2")->reduce;
$p2 = Formula("$a[1][2] + $a[2][2] * x + $a[3][2] * x^2")->reduce;
$p3 = Formula("$a[1][3] + $a[2][3] * x + $a[3][3] * x^2")->reduce;

$popup = PopUp(['choose','linearly dependent','linearly independent'],'linearly independent');

$multians = MultiAnswer(0, 0, 0)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      return Matrix([$s[0],$s[1],$s[2]]) == Matrix([0,0,0]);
  }
);

Context()->texStrings;
BEGIN_TEXT
Are the vectors \( $p1, \ $p2 \) and \( $p3 \) linearly independent?
$BR
\{ $popup->menu \}
$BR
$BR
If the vectors are independent, enter zero in every answer blank since zeros are only the values that make the equation below true.  
If they are dependent, find numbers, not all zero, that make the equation below true.    You should be able to explain and justify your answer.
$BR
$BR
\( 0 = \) 
\{ $multians->ans_rule(5) \} \( ($p1) + \)
\{ $multians->ans_rule(5) \} \( ($p2) + \)
\{ $multians->ans_rule(5) \} \( ($p3). \)
END_TEXT
Context()->normalStrings;

ANS($popup->cmp);
ANS($multians->cmp);

COMMENT('Always independent.');
ENDDOCUMENT();
