## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION


## DBsubject(Linear algebra)
## DBchapter(Euclidean spaces)
## DBsection(Linear independence)
## Date(July 2013)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## KEYWORDS('linear algebra','vector space','subspaces')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 1;
$refreshCachedImages = 1;
TEXT(beginproblem()); 

Context('Matrix');

$A = Matrix([
[1,0,0],
[0,1,0],
[0,0,1],
[0,0,0],
]);

Context()->texStrings;
BEGIN_TEXT
Construct a \( 4 \times 3 \) matrix \( A = \lbrack \ \vec{a}_1 \ \ \vec{a}_2 \ \ \vec{a}_3 \ \rbrack \)
such that \( \lbrace \vec{a}_1, \vec{a}_2 \rbrace \) is a linearly independent 
set and \( \vec{a}_3 \) is not in the span of \( \lbrace \vec{a}_1, \vec{a}_2 \rbrace \).
$BR
$BR
\( A = \) \{ $A->ans_array(10) \}
END_TEXT
Context()->normalStrings;

ANS($A->cmp(checker => sub {
      my ( $correct, $student, $self ) = @_;
      my $S = Matrix($student);
      my $C = Matrix($correct);
      return ((($S->transpose) * $S)->det != Real(0));
}));


ENDDOCUMENT();
