# DESCRIPTION
# Gives an expression and asks to translate from english into math
# ENDDESCRIPTION

## DBsubject('Intermediate Algebra')
## DBchapter('Ch 01: Algebra, Mathematical Models, and Problem Solving')
## DBsection('Algebraic Expressions and Real Numbers')
## KEYWORDS('real numbers')
## Author('Rick Lynch')
## Institution('University of Missouri-Columbia')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserAssignment.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
Context("Numeric");
Context()->variables->are(m=>"Real",n=>"Real");
Context()->strings->add("="=>{}, "<="=>{}, "<"=>{}, ">"=>{}, ">="=>{});
parser::Assignment->Allow;

Context()->texStrings;
$m1 = random(5,15);
do {$m2 = random(5,15)} until ($m2 != $m1);
$n1 = random(5,15);
do {$n2 = random(5,15)} until ($n2 != -$n1);
$ans11 = Formula("m/$m1");
$ans12 = "=";
$ans13 = Formula("$m2 m");
$ans21 = Formula("$n1/n");
$ans22 = "=";
$ans23 = Formula("$n2+n");
$dispans1 = "\frac{m}{$m1} = $m2 m";
$dispans2 = "\frac{$n1}{n} = $n2+n";

$ma1 = MultiAnswer($ans11,$ans12,$ans13)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);

$ma2 = MultiAnswer($ans21,$ans22,$ans23)->with(
  singleResult => 0,
  checker => sub {
    my ($correct, $student, $self) = @_;
    my ($a1s,$a2s,$a3s) = @{$student};
    my ($a1,$a2,$a3) = @{$correct};
    if ((($a1s == $a1 && $a3s == $a3) || ($a1s == $a3 && $a3s == $a1)) && ($a2s == $a2)) {
      return [1,1,1];
    } else {
      return [0,0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Write each sentence using mathematical symbols:
$PAR
(a) The quotient of \(m\) and \($m1\) is \($m2\) times \(m\): \{$ma1->ans_rule(10)\} \{$ma1->ans_rule(5)\} \{$ma1->ans_rule(10)\}
$PAR
(b) The quotient of \($n1\) and \(n\) is \($n2\) more than \(n\): \{$ma2->ans_rule(10)\} \{$ma2->ans_rule(5)\} \{$ma2->ans_rule(10)\}
END_TEXT


###########################################################################
# check the answer  
###########################################################################
ANS($ma1->cmp());
ANS($ma2->cmp());


###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} (a) \($dispans1\), $SPACE (b) \($dispans2\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Note:${EBOLD} The middle answer blank in both parts is for you to enter "\(<\)=", "\(<\)", "=", "\(>\)", or "\(>\)=" for "less than or equal to", "less than", "equals", "greater than", or "greater than or equal to", respectively.
  $BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
