##DESCRIPTION
##  Exam problem related to coterminal angles
##ENDDESCRIPTION


## DBsubject(Trigonometry)
## DBchapter(Geometric and algebraic foundations for trigonometry)
## DBsection(Reference angles (using coterminal angles))
## Date(9/11/2012)
## Institution(Mississippi College)
## Author(John Travis)
## Level(2)
## MO(1)
## KEYWORDS('trigonometry', 'angle', 'coterminal')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",    
   "MathObjects.pl",
   "parserMultiAnswer.pl",
   "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$angle = random(95,175,10); 

#$ans = Real("$angle")->with(period=>360);

#$ans  = $angle;

##############################################################
#
#  Text
#
#

$multians = MultiAnswer($angle, $angle,$angle)->with(
  singleResult => 0,
  allowBlankAnswers=>1,
  # checkTypes => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @s = @{$student};
      my @c = @{$correct};
      @history = ();
      @ret = ();
      foreach (0..$#c){@ret = (@ret,0);}
      foreach $i (0..$#c){
	if (  ref($c[$j]) eq ref($s[$i]) && $s[$i] %360 == $c[$i]){
           $ret[$i] = 1;
           foreach $h  (0 .. $#history){
              if(ref($history[$h]) eq ref($s[$i]) && $s[$i] == $history[$h]){ $ret[$i] = 0};
           }
           @history = ($s[$i], @history);
        }
      }
      return [@ret];
  }
);


BEGIN_TEXT
Determine three $BBOLD different$EBOLD angles which are coterminal with \($angle^{\circ}\):
$PAR
$BCENTER
 \{$multians->ans_rule(5)\}\(^{\circ}\), \{$multians->ans_rule(5)\}\(^{\circ}\), \{$multians->ans_rule(5)\}\(^{\circ}\)
$ECENTER

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

ANS( $multians->cmp() ); 

ENDDOCUMENT();        
