## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Limit of a sequence)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('APEX Calculus')
## AuthorText1('Hartman')
## EditionText1('3.0')
## Section1('8.1')
## Problem1('22')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserMultiAnswer.pl",
  "unionTables.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());

Context("Numeric")->variables->are(n=>'Real');
Context()->strings->add('diverges' => {caseSensitive => 0});
Context()->noreduce('(-x)-y');

do {
  $a = random(-1,1,2) * random(2,9,1);
  $b = random(2,9,1);
  $c = random(2,9,1);
  $p = list_random(3,5,7);
} until ( 
  gcd($a,$b) == 1 && $a != $c && $b != $c
);

$ans1num = Compute("$a")->reduce;
$ans1den = Compute("($b + $c/n^$p)^(1/$p)")->reduce;
$ans2 = Real("$a / ($b^(1/$p))");

$multians1 = MultiAnswer($ans1num,$ans1den)->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  format => "[ %s ] / [ %s ]",
  tex_format => "\frac{%s}{%s}",
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ($c1,$c2) = @{$correct};
      my ($s1,$s2) = @{$student};
      if ($c1 == $s1 && $c2 == $s2) { return [1,1]; }
      if ($c1 == $s1 && $c2 != $s2) { return [1,0]; }
      if ($c1 != $s1 && $c2 == $s2) { return [0,1]; }

      if ($c1 == -1*$s1 && $c2 == -1*$s2) { return [1,1]; }
      if ($c1 == -1*$s1 && $c2 != -1*$s2) { return [1,0]; }
      if ($c1 != -1*$s1 && $c2 == -1*$s2) { return [0,1]; }

      return [0,0];
      #return [ $c1 == $s1, $c2 == $s2 ] ;
  }
);

sub displayfrac {
    my $ma = shift; # input is a multianswer object
    ANS($ma->cmp);
    if ($displayMode eq 'TeX') {
      "\(\displaystyle \frac{".$ma->ans_rule(10)."}{".$ma->ans_rule(10)."} \)";
    } else {
      "<table style='display:inline-block; vertical-align:middle; border:none;'><tr><td>".$ma->ans_rule(15).$BR.$HR.$ma->ans_rule(15)."</td></tr></table>" 
    }
}



BEGIN_PGML
Use algebra to simplify the expression before evaluating the limit.  In particular, factor the highest power of [`n`] from the numerator and denominator, then cancel as many factors of [`n`] as possible.  If the sequence does not converge, enter _diverges_ in the final answer box.

[``\lim_{n \to \infty} \frac{ [$a] n }{ ([$b] n^{[$p]} + [$c] )^{1/ [$p] } } = \lim_{n \to \infty} \bigg(``] 
[@ displayfrac($multians1) @]* 
[``\bigg) = ``] [_____________]{$ans2}.
END_PGML


ENDDOCUMENT();
