# DESCRIPTION
# Factor an expression of the form a x^2 + b x + c, where a > 1, b < 0, and c < 0.
# ENDDESCRIPTION

## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## DBsubject(Algebra)
## DBchapter(Factoring)
## DBsection(Factoring trinomials)
## Institution(University of Missouri-Columbia)
## Level(2)
## KEYWORDS('factoring', 'factor', 'factorization', 'substitution')


###########################################################################
# initialization 
###########################################################################
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
#  "contextLimitedPowers.pl",
#  "contextLimitedPolynomial.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;


###########################################################################
# setup contexts and variables 
###########################################################################
#Context("LimitedPolynomial-Strict");
Context("Numeric");
$var = list_random('a','b','m','n','x','y','z','w','p','q','r','s','t','u','v');
Context()->variables->are($var=>"Real", uc($var)=>"Real");

do {
  $a = random(1,4);
  do {$b = non_zero_random(-4,4)} until (gcd($a,$b) == 1);
  $c = random(1,4);
  do {$d = non_zero_random(-4,4)} until (gcd($c,$d) == 1);
  $e1 = $a*$c;
  $e2 = $a*$d+$b*$c;
  $e3 = $b*$d;
} until ($e1 > 1 && $e2 < 0 && $e3 < 0 && ($a != $c && $b != $d));

$fac1 = Formula("$a*$var + $b")->reduce;
$fac2 = Formula("$c*$var + $d")->reduce;

$expn = Formula("$e1*$var^2+$e2*$var+$e3")->reduce;

$answer = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $c1, $c2 ) = @{$correct};
    my ( $s1, $s2 ) = @{$student};
#    Check that the answers are not constants and their product is correct
    $score1 = 1; $score2 = 1;
    if ( $s1->D("$var") == Formula("0") ){
       $self -> setMessage(1,"Answers can't be constants.");
    }
    if ( $s2->D("$var") == Formula("0") ){
       $self -> setMessage(2,"Answers can't be constants.");
    }
    if ( $score1*$score2 == 0 ) {
       return [$Sscore1,$score2];
    } elsif ($s1*$s2 == $c1*$c2){
       return [1,1];
    } else {
      return [0,0];
    }
  }
);


###########################################################################
# state the problem 
###########################################################################
Context()->texStrings;
BEGIN_TEXT
Factor the expression and simplify your answer as much as possible:
$PAR $BCENTER 
\($expn = \big( \) \{ $answer->ans_rule(20) \} \( \big) \big( \) \{ $answer->ans_rule(20) \} \( \big) \) 
$ECENTER
END_TEXT
Context()->normalStrings;


###########################################################################
# check the answer  
###########################################################################
ANS($answer->cmp());

COMMENT('Desc: Factor an expression of the form a x^2 + b x + c, where a > 1, b < 0, and c < 0.');
ENDDOCUMENT();
