##DESCRIPTION
## Use integration by parts to evaluate an integral
##ENDDESCRIPTION

## DBsubject(Differential equations)
## DBchapter(Laplace transforms)
## DBsection(Inverse transformations)
## Date(11/13/2010)
## Institution(Alfred University)
## Author(Darwyn Cook)
## Level(3)
## TitleText1('')
## AuthorText1('Zill')
## EditionText1('')
## Section1('.')
## Problem1('')
## KEYWORDS('laplace', 'convolution')


DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "Alfredmacros.pl",
  "unionTables.pl",
  "weightedGrader.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;
Context("Numeric");
Context()->variables->are(t=>"Real",s=>"Real",a=>"Real",m=>"Real",tau=>"Real");
Context()->variables->set(tau => {TeX => "\tau"});
$n = random(1,9);
$fact = 1;
for($i=1;$i<=$n;$i++){$fact = $fact*$i;};

## table has the format f, numerator of F, denominator of F
@Table = ([Formula("t^$n"),Formula("$fact"),Formula("s^($n+1)")],
[Formula("e^(m*t)"),Formula("1"),Formula("s-m")],
[Formula("cos(a*t)"),Formula("s"),Formula("a^2+s^2")],
[Formula("sin(a*t)"),Formula("a"),Formula("a^2+s^2")],
[Formula("sinh(a*t)"),Formula("-a"),Formula("a^2-s^2")],
[Formula("cosh(a*t)"),Formula("s"),Formula("s^2-a^2")]) ;


$numformulas = @Table;
$minusone = $numformulas-1;
$plusone = $numformulas + 1;

$choose1 = random(0,$minusone);
$choose2 = random(0,$minusone);

if($choose2 == $choose1){$choose2=($choose2 + 1)%$numformulas;}


$ftau =  $Table[$choose1][0]->substitute(t=>"tau");
$gtau =  $Table[$choose2][0]->substitute(t=>"tau");
$ftranslate =  $Table[$choose1][0]->substitute(t=>"t-tau");
$gtranslate =  $Table[$choose2][0]->substitute(t=>"t-tau");

$convolution1 = $ftau*$gtranslate;
$convolution2 = $gtau*$ftranslate;


$num = Formula("$Table[$choose1][1]*$Table[$choose2][1]");
$den = Formula("$Table[$choose1][2]*$Table[$choose2][2]");

$F = Formula($num/$den)->reduce;

Context()->texStrings;
BEGIN_TEXT
If \(L\lbrace f\rbrace (s) = $F\) then 
$BR
$BR
$BCENTER
\{BeginTable(center=>0).
    Row(['\(\displaystyle f(t) =\)',tableintegral(width=>1),ans_rule(20),'\(d\tau\)'],separation=>2).
EndTable();
\}
$ECENTER
$BR
$BBOLD Note $EBOLD \(\tau\) is typed as tau.
END_TEXT
Context()->normalStrings;

WEIGHTED_ANS(Compute("0")->cmp, 5);
WEIGHTED_ANS(Formula("t")->cmp, 10);
WEIGHTED_ANS( $convolution1->cmp( checker=>sub {
     my ( $correct, $student, $ansHash ) = @_;
     return (($convolution1 == $student)||($convolution2 == $student));
} ), 85 );
ENDDOCUMENT();        # This should be the last executable line in the problem.
