##DESCRIPTION
#
# File Created: 3/27/2009
# Last Modified: 3/27/2009
# Problem Author: Juan Marin
# WeBWorK Entry: 
# Location: Alfred University
#
##ENDDESCRIPTION

##KEYWORDS()
##

## DBsubject('')
## DBchapter('')
## DBsection('')
## Date('3/27/2009')
## Author('Juan Marin')
## Institution('Alfred University')
## TitleText1('Differential Equations with Boundary-Value Peoblems')
## EditionText1('7')
## AuthorText1('Zill and Cullen')
## Section1('4.7-CauchyEulerEquations')
## Problem1('')


DOCUMENT();

# Load whatever macros you need for the problem
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "PGgraders.pl",
  "PGgraphmacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "parserImplicitEquation.pl",
  "PGcourse.pl"
);
 
$u = random(2,5,1);
$v = random(6,9,1);

$b = $v+$u+1;
$c = $v*$u;
$e = random(2,9,1);
$be = $b*$e;
$ce = $c*$e;
$p = 1-$u;
$q = 1-$v;

$x0=random(2,9);

$f = Compute("x^{$p}");
$g = Compute("x^{$q}");
$h = Compute("$e");
Context()->operators->remove('/','+');
$w = Compute("x^(-1-$u-$v)*($u-$v)*($v-1)*($u-1)*$e");
Context("Interval");
$int = Compute("(0,inf)");

Context()->texStrings;
BEGIN_TEXT
$PAR
It can be shown that \(y_1=$f,\ y_2=$g\) and \(y_3=$h\) are 
solutions to the differential equation 
\(x^2D^3y+$b x D^2y+ $c Dy =0\).
$PAR 
\(W(y_1, y_2, y_3 )\) =  \{$w->ans_rule\}.
$BR
For an IVP with initial conditions at \(x = $x0\), \(c_1y_1+c_2y_2+c_3y_3\) is the general solution for x on what interval? \{$int->ans_rule\}
END_TEXT
Context()->normalStrings;
ANS($w->cmp(checker=>sub {
                my ( $correct, $student, $self ) = @_;
                my $context = Context()->copy;
                $context->flags->set(no_parameters=>0);
                $context->variables->add('c0'=>'Parameter');
                my $c0 = Formula($context,'c0');
                $student = Formula($context,$student);
                $correct = Formula($context,"(x^(-1-$u-$v)*($u-$v)*($v-1)*($u-1)*$e)");
                return $correct == $student;}));
ANS($int->cmp());
ENDDOCUMENT();
