# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Partial fractions)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(PartFracPureAlg)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('7.4')
## Section2('7.4')
## Section3('7.4')
## Section4('7.4')
## Problem1('3')
## Problem2('3')
## Problem3('5')
## Problem4('5')
## KEYWORDS('calculus', 'integral', 'antiderivatives', 'indefinite integrals', 'partial fractions')

## Textbook tags
## HHChapter1('Integration')
## HHChapter2('Integration')



DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "hhAdditionalMacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->are(y=>"Real");
$showPartialCorrectAnswers = 1;
SRAND($psvn);

$a = random(2,8,1);
$b = random(2,5,1);
$bsq = $b*$b;
$tbsq = 2*$bsq;

$func = Compute("$a/(y^3 - $bsq y)");
$fac1 = Compute("-$a/($bsq y)");
$fac2 = Compute("$a/($tbsq (y - $b))");
$fac3 = Compute("$a/($tbsq (y + $b))");

$mp = MultiAnswer($fac1, $fac2, $fac3)->with(
    singleResult => 0,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	my ( $f1stu, $f2stu, $f3stu ) = @{$student};
	my ( $f1, $f2, $f3 ) = @{$correct};
	if ( ( $f1 == $f1stu && 
	       ( ( $f2 == $f2stu && $f3 == $f3stu ) ||
		 ( $f2 == $f3stu && $f3 == $f2stu ) ) ) ||
	     ( $f1 == $f2stu && 
	       ( ( $f2 == $f1stu && $f3 == $f3stu ) ||
		 ( $f2 == $f3stu && $f3 == $f1stu ) ) ) ||
	     ( $f1 == $f3stu && 
	       ( ( $f2 == $f1stu && $f3 == $f2stu ) ||
		 ( $f2 == $f2stu && $f3 == $f1stu ) ) ) ) {
	    return [1,1,1];
	} else {
	    @ret = (0,0,0);
	    if ( $f1stu == $f1 || $f1stu == $f2 || $f1stu == $f3 ) {
	        $ret[0] = 1;
		if ( ( $f2stu == $f1 || $f2stu == $f2 || $f2stu == $f3 ) &&
		     $f1stu != $f2stu ) {
		    $ret[1] = 1;
		} elsif ( ($f3stu == $f1 || $f3stu == $f2 || $f3stu == $f3) &&
			  $f1stu != $f3stu ) {
		    $ret[2] = 1;
		}
	    } elsif ( $f2stu == $f1 || $f2stu == $f2 || $f2stu == $f3 ) {
	        $ret[1] = 1;
		if ( ( $f3stu == $f1 || $f3stu == $f2 || $f3stu == $f3 ) &&
		     ( $f2stu != $f3stu ) ) {
		    $ret[2] = 1;
		}
	    } elsif ( $f3stu == $f1 || $f3stu == $f2 || $f3stu == $f3 ) {
		$ret[2] = 1;
	    }
	    return [ @ret ];
	}
    }
);

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Split into partial fractions:$BR
\( \displaystyle $func = \)
\{ $mp->ans_rule(15) \} + \{ $mp->ans_rule(15) \} + 
\{ $mp->ans_rule(15) \}

END_TEXT
Context()->normalStrings;

ANS($mp->cmp() );

$ca = reduced_frac($a,$bsq);
$cb = reduced_frac($a,$tbsq);

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Since \(y^3- $bsq y=y(y-$b)(y+$b)\), we take
\[\frac{$a}{y^3-$bsq y}= \frac{A}{y}  +\frac{B}{y-$b} +\frac{C}{y+$b}.\]
So,
\[
  $a = A(y-$b)(y+$b)+By(y+$b)+Cy(y-$b),
\]
or, on rearranging to collect powers of \(y\),
\[
  $a = (A+B+C)y^2 +($b B- $b C)y- $bsq A,
\]
giving
\[ A+B+C = 0, \]
\[ $b B - $b C = 0, \]
and
\[ -$bsq A = $a. \]
Thus \(A=-$ca\), \(B=C=$cb\) so
\[ \frac{$a}{y^3-$bsq y}=-$ca \left(\frac{1}{y}\right)  +
    $cb \left(\frac{1}{y-$b}\right) + $cb \left(\frac{1}{y+$b}\right).\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
