# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 02/08/09

## DBsubject(Calculus - single variable)
## DBchapter(Techniques of integration)
## DBsection(Partial fractions)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('7.4')
## Section2('7.4')
## Section3('7.4')
## Section4('7.4')
## Problem1('57')
## Problem2('60')
## Problem3('72')
## Problem4('73')
## KEYWORDS('calculus', 'integral', 'antiderivatives', 'indefinite integrals', 'partial fractions')

## Textbook tags
## HHChapter1('Integration')
## HHChapter2('Integration')



DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "parserFormulaUpToConstant.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->variables->add(w=>"Real");
Context()->flags->set(reduceConstants=>0);
$showPartialCorrectAnswers = 1;

$b = random(1,8,1);
$bsq = $b*$b;

$frac = Compute("2x/(x^2 - $bsq)");
$fac1 = Compute("1/(x-$b)");
$fac2 = Compute("1/(x+$b)");

$int1 = Compute("ln(|x-$b|)");
$int2 = Compute("ln(|x+$b|)");

$mp1 = MultiAnswer($fac1, $fac2)->with(
    singleResult => 0,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	my ( $f1stu, $f2stu ) = @{$student};
	my ( $f1, $f2 ) = @{$correct};
	if ( ( $f1 == $f1stu && $f2 == $f2stu ) ||
	     ( $f1 == $f2stu && $f2 == $f1stu ) ) {
	    return [1,1];
	} else {
	    if ( $f1 == $f1stu || $f2 == $f1stu ) {
# 		$self->setMessage(1,"correct");
# 		$self->setMessage(2,"incorrect");
		return [1,0];
	    } elsif ( $f1 == $f1stu || $f2 == $f2stu ) {
# 		$self->setMessage(1,"incorrect");
# 		$self->setMessage(2,"correct");
		return [0,1];
	    } else {
		return [0,0];
	    }
	}
    }
);

$mp2 = MultiAnswer($int1, $int2)->with(
    singleResult => 0,
    checker => sub {
	my ( $correct, $student, $self ) = @_;
	my ( $f1stu, $f2stu ) = @{$student};
	my ( $f1, $f2 ) = @{$correct};
	if ( ( $f1 == $f1stu && $f2 == $f2stu ) ||
	     ( $f1 == $f2stu && $f2 == $f1stu ) ) {
	    return [1,1];
	} else {
	    if ( $f1 == $f1stu || $f2 == $f1stu ) {
# 		$self->setMessage(1,"correct");
# 		$self->setMessage(2,"incorrect");
		return [1,0];
	    } elsif ( $f1 == $f1stu || $f2 == $f2stu ) {
# 		$self->setMessage(1,"incorrect");
# 		$self->setMessage(2,"correct");
		return [0,1];
	    } else {
		return [0,0];
	    }
	}
    }
);

$ans1 = Compute("ln(|x-$b|) + ln(|x+$b|)");
$ans2w = Compute("ln(|w|)");
$ans2 = Compute("ln(|x^2 - $bsq|)");

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Calculate the integral below by partial
fractions and by using the indicated substitution.
Be sure that you can show how the results you obtain are the same.
\[ \int $frac\,dx \]
$PAR
First, rewrite this with partial fractions:
$BR
\( \int $frac\,dx = \int \)
\{$mp1->ans_rule(15)\} \( dx + \int \)
\{$mp1->ans_rule(15)\} \( dx = \)
\{$mp2->ans_rule(15)\} + \{$mp2->ans_rule(15)\} \( + C\).
$BR
${BITALIC}(Note that you should not include the \(+C\) in your
entered answer, as it
has been provided at the end of the expression.)$EITALIC
$PAR

Next, use the substitution \(w = x^2 - $bsq\) to find the integral:
$BR
\( \int $frac\,dx = \int \)
\{ ans_rule(15) \} \( dw = \)
\{ ans_rule(15) \} \( + C = \)
\{ ans_rule(15) \} \( + C\).
$BR
${BITALIC}(For the second answer blank, give your antiderivative
in terms of the variable \(w\).  Again, note that you should not
include the \(+C\) in your answer.)$EITALIC

END_TEXT
Context()->normalStrings;

ANS($mp1->cmp() );
ANS($mp2->cmp() );
ANS(Compute("1/w")->cmp() );
ANS($ans2w->cmp() );
ANS($ans2->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Using partial fractions, we write
\[
\frac{2x}{x^2-$bsq}=\frac{A}{x+$b}+\frac{B}{x-$b}
\]
so that
\[ 2x = A(x-$b)+B(x+$b) = (A+B)x + $b (-A+B). \]
Solving, we get \(A=B=1\), so that
\[
\int \frac{2x}{x^2-$bsq}\,dx=
  \int \left(\frac{1}{x+$b}+\frac{1}{x-$b}\right)\,dx
  =\ln|x+$b|+\ln|x-$b|+C.
\]
$PAR
Using the substitution \(w=x^2-$bsq\), we get \(dw= 2x\,dx\),
so we have
\[
\int \frac{2x}{x^2-$bsq}\,dx=
  \int \frac{dw}{w}=\ln|w|+C = \ln|x^2-$bsq|+C. \]
$PAR
The properties of logarithms show that the two results are the
same:
\[
\ln|x+$b|+\ln|x-$b| = \ln|(x+$b)(x-$b)| =\ln|x^2-$bsq|.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
