##DESCRIPTION
##  Number theory: simple Diophantine equation (only one prime common factor)
##ENDDESCRIPTION

##KEYWORDS('number theory', 'diophantine', 'equation')

## DBsubject('Number Theory')
## DBchapter('Diophantine equations')
## DBsection('Solving Diophantine equations')
## Date('07/14/2018')
## Author('John Perry')
## Institution('University of Southern Mississippi')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "scaffold.pl",
   "parserMultiAnswer.pl",
   "parserPopUp.pl",
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

$has_solution = PopUp( [ "?", "yes", "no" ], "yes" );

@primes = ( 7, 11, 13, 17 );
$d = @primes[ random(0,3) ];
$a = $d * random(2,6);
$b = $a;
while ($a == $b || ($a % $b ==0) || ($b % $a == 0)) {
  $b = $d * random(2,6);
}
$x0 = random(-9,-6,1);
$y0 = random(6,9,1);
$c = $a * $x0 + $b * $y0;
$ma = MultiAnswer(0,0)->with(
  singleResult => 1,
  checker => sub {
    my ( $correct, $student, $self ) = @_;
    my ( $x, $y ) = @{ $student };
    my ( $z ) = $a*$x + $b*$y;
    if ( $a * $x + $b * $y == $c ) { return 1; }
    else { return 0; }
  }
);

$x_d = $b/$d;
$y_d = $a/$d;

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

<b>This is a three-part problem. Once you answer the first part correct, a second part will appear. Once you answer that correctly, a third part will appear.</b>
$PAR
Consider the Diophantine equation \[ $a x + $b y = $c . \]

END_TEXT
Context()->normalStrings;

Scaffold::Begin();

Section::Begin("Part 1");

Context()->texStrings;
BEGIN_TEXT

Does this equation have integer solutions?
\{ $has_solution->menu() \}

END_TEXT
Context()->normalStrings;

ANS( $has_solution->cmp() );

Section::End();

Section::Begin("Part 2");

Context()->texStrings;
BEGIN_TEXT

Good! Find two integer solutions to the equation.
$PAR
\(x=\) \{ $ma->ans_rule(5) \}
\(y=\) \{ $ma->ans_rule(5) \}

END_TEXT
Context()->normalStrings;

ANS( $ma->cmp() );

Section::End();

Section::Begin("Part 3");

BEGIN_TEXT

What are the <b>smallest</b> positive integers I should add or subtract to the solutions we've already found to obtain another solution?
$PAR
Add (or subtract) \{ ans_rule(5) \} to \(x\).
$BR
Subtract (or add) \{ ans_rule(5) \} to \(y\).

END_TEXT
Context()->normalStrings;

ANS( Compute($x_d)->cmp() );
ANS( Compute($y_d)->cmp() );

Section::End();

Scaffold::End();

ENDDOCUMENT();        

