## DESCRIPTION
## Linear systems of ODEs
## ENDDESCRIPTION

## KEYWORDS('differential equations','linear systems')

## DBsubject('Differential Equations')
## DBchapter('Linear Systems of Differential Equations')
## DBsection('Linear Systems')
## Date('01/30/2011')
## Author('Paul Pearson')
## Institution('Fort Lewis College')
## TitleText1('Notes on Diffy Qs')
## EditionText1('December 9, 2010')
## AuthorText1('Jiri Lebl')
## Section1('3.3')
## Problem1('2')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGmatrixmacros.pl",
"unionLists.pl",
"parserMultiAnswer.pl",
);

TEXT(beginproblem());


#############################
#  Setup

Context("Numeric");
Context()->variables->add(
t=>"Real",
a=>"Real",b=>"Real",k=>"Real",
A=>"Real",B=>"Real"
);

$answer1 = Compute("a k e^(k t)");
$answer2 = Compute("b k e^(k t)");

$i = random(0,2,1);

if ($i == 0) {

  @matrix = (1,3,3,1);
  @evect1 = (1,1);
  $eval1  = 4;
  @evect2 = (1,-1);
  $eval2  = -2;

} elsif ($i == 1) {

  @matrix = (7,2,-4,1);
  @evect1 = (1,-1);
  $eval1  = 5;
  @evect2 = (1,-2);
  $eval2  = 3;

} else {

  @matrix = (5,3,3,5);
  @evect1 = (1,-1);
  $eval1  = 2;
  @evect2 = (1,1);
  $eval2  = 8;
}

$x1a = Compute("$evect1[0] A e^($eval1 t) + $evect2[0] B e^($eval2 t)");
$x2a = Compute("$evect1[1] A e^($eval1 t) + $evect2[1] B e^($eval2 t)");

$x1b = Compute("$evect1[0] B e^($eval1 t) + $evect2[0] A e^($eval2 t)");
$x2b = Compute("$evect1[1] B e^($eval1 t) + $evect2[1] A e^($eval2 t)");


$multians = MultiAnswer($x1a, $x2a)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $x1stu, $x2stu ) = @{$student};
      if ( ($x1a == $x1stu && $x2a == $x2stu) ||
           ($x1b == $x1stu && $x2b == $x2stu) ) {
          return [1,1];
      } elsif (($x1a == $x1stu || $x1b == $x1stu)) {
          return [1,0];
      } elsif (($x2a == $x2stu || $x2b == $x2stu)) {
          return [0,1];
      } else {
          return [0,0];
      }
  }
);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{ BeginList('OL',type=>'a') \}

  

$ITEM Assuming \( a, b \) and \( k \) are constants, 
calculate the following derivative.
\{ 
mbox([
  "\( \displaystyle \frac{d}{dt} \left( \left\lbrack \begin{array}{c} a \\ b \end{array} \right\rbrack e^{k t} \right) = \)",
  display_matrix([[ans_rule(10)], [ans_rule(10)]], align=>'c')], 
  valign=>'middle'
);
\}

$ITEMSEP
$ITEM Find a value of \( k \) so that 
\( \displaystyle 
\left\lbrack 
\begin{array}{r} $evect1[0] \\ $evect1[1] \end{array} 
\right\rbrack 
e^{k t}
\) 
is a solution to
\( \displaystyle
\vec{x}\,'
= 
\left\lbrack
\begin{array}{rr}
$matrix[0] & $matrix[1] \\
$matrix[2] & $matrix[3]
\end{array}
\right\rbrack 
\vec{x}.
\)
$BR
\( k = \) 
\{ ans_rule(10) \}

$ITEMSEP
$ITEM Find a value of \( k \) so that 
\( \displaystyle 
\left\lbrack 
\begin{array}{r} $evect2[0] \\ $evect2[1] \end{array} 
\right\rbrack 
e^{k t}
\) 
is a solution to
\( \displaystyle
\vec{x}\,'
= 
\left\lbrack
\begin{array}{rr}
$matrix[0] & $matrix[1] \\
$matrix[2] & $matrix[3]
\end{array}
\right\rbrack 
\vec{x}.
\)
$BR
\( k = \) 
\{ ans_rule(10) \}

$ITEMSEP
$ITEM Write down the general solution in 
the form \( x_1 = ? \) and \( x_2 = ? \), 
i.e., write down a formula for each 
component of the solution.  Use \( A \)
and \( B \) to denote arbitrary constants.
$BR
\( x_1 = \) 
\{ $multians->ans_rule(40) \}
$BR
\( x_2 = \) 
\{ $multians->ans_rule(40) \}

\{ EndList('OL') \}
END_TEXT
Context()->normalStrings;


######################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( $answer1->cmp() );
ANS( $answer2->cmp() );

ANS( Compute("$eval1")->cmp() );
ANS( Compute("$eval2")->cmp() );

ANS( $multians->cmp() );

COMMENT("MathObject version.");

ENDDOCUMENT();


