# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Aaron Wooton, <wootton@up.edu>
# revised by Gavin LaRose, <glarose@umich.edu> 11/2008

## Tagged by glr 11/04/08

## Author('Gavin LaRose');
## DBsubject(Trigonometry)
## DBchapter(Trigonometric functions)
## DBsection(Combinations and compositions of functions)
## Institution(University of Michigan)
## MLT(Compose_find_components)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('1.5')
## Section2('1.5')
## Section3('1.5')
## Section4('1.5')
## Problem1('36')
## Problem2('33')
## Problem3('33')
## Problem4('59')
## KEYWORDS('calculus','function','trigonometric functions','inverse trigonometric functions','transformation of functions','translations')

## Textbook tags
## HHChapter1('A Library of Functions')
## HHChapter2('A Library of Functions')


DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

Context("Numeric");

TEXT(beginproblem());
$showPartialCorrectAnswers = 0;

$which = random(0,2,1);

if ( $which == 0 ) { $c="sin"; $ctex = "\sin"; $d="sine"; }
if ( $which == 1 ) { $c="cos"; $ctex = "\cos"; $d="cosine"; }
if ( $which == 2 ) { $c="tan"; $ctex = "\tan"; $d="tangent"; }

$n = random(2,4,1);

$f1 = Compute("x^$n");
$g1 = Compute("$c(x)");

## pick a domain on which the functions are likely to be well-defined
Context()->flags->set(limits=>[0, 1.57**(1/$n)]);

$mp1 = MultiAnswer( $f1, $g1 )->with(
    singleResult=>0,
    checker=>sub {
	my ( $cor, $stu, $self ) = @_;
	my ( $stuf, $stug ) = @{$stu};
	my ( $corf, $corg ) = @{$cor};

	if ( $corf == $stuf && $corg == $stug ) {
	    return [1,1];
	} else {
	    if ( $stuf == Compute("x") ) {
		$self->setMessage(1, "Write your composition without using " .
				  "the identity function f(x) = x.");
		if ( $stug == Compute("$c^($n)(x)") ) {
		    $self->setMessage(2, "This is just the original " .
				      "function.");
		}
		return [0,0];
	    } else {
		if ( ref($stuf) eq 'Value::Formula' ) {
		    $stufg = Compute($stuf)->substitute( x=>$stug );
		    $corfg = Compute($corf)->substitute( x=>$corg );
		    if ( $corfg == $stufg ) {
			return [1,1];
		    } else {
			return [0,0];
		    }
		} else {
		    return [0,0];
		}
	    }
	}
    }
);

$f2 = Compute("$c(x)");
$g2 = Compute("$c(x)");

$mp2 = MultiAnswer( $f2, $g2 )->with(
    singleResult=>0,
    checker=>sub {
	my ( $cor, $stu, $self ) = @_;
	my ( $stuf, $stug ) = @{$stu};
	my ( $corf, $corg ) = @{$cor};

	if ( $stuf == $corf && $stug == $corg ) {
	    return [1,1];
	} else {
	    if ( $stuf == Compute("x") ) {
		$self->setMessage(1, "Write your composition without using " .
				  "the identity function f(x) = x.");
		if ( $stug == Compute("$c($c(x))") ) {
		    $self->setMessage(2, "This is just the original " .
				      "function.");
		}
		return [0,0];
	    } else {
                if ( ref($stuf) eq 'Value::Formula' ) {
		    $stufg = Compute($stuf)->substitute( x=>$stug );
		    $corfg = Compute($corf)->substitute( x=>$corg );
		    if ( $stufg == $corfg ) {
			return [1,1];
		    } else {
			return [0,0];
		    }
                } else {
		    return [0,0];
		}
	    }
	}
    }
);

$f3 = Compute("$c(x)");
$g3 = Compute("x^($n)");

$mp3 = MultiAnswer( $f3, $g3 )->with(
    singleResult=>0,
    checker=>sub {
	my ( $cor, $stu, $self ) = @_;
	my ( $stuf, $stug ) = @{$stu};
	my ( $corf, $corg ) = @{$cor};

	if ( $stuf == $corf && $stug == $corg ) {
	    return [1,1];
	} else {
	    if ( $stuf == Compute("x") ) {
		$self->setMessage(1, "Write your composition without using " .
				  "the identity function f(x) = x.");
		if ( $stug == Compute("$c(x^($n))") ) {
		    $self->setMessage(2, "This is just the original " .
				      "function.");
		}
		return [0,0];
	    } else {
                if ( ref($stuf) eq 'Value::Formula' ) {
		    $stufg = Compute($stuf)->substitute( x=>$stug );
		    $corfg = Compute($corf)->substitute( x=>$corg );
		    if ( $stufg == $corfg ) {
			return [1,1];
		    } else {
			return [0,0];
		    }
                } else {
		    return [0,0];
		}
	    }
	}
    }
);

Context()->texStrings;
BEGIN_TEXT
The following three functions look very similar, but define very
different functions. Think about how they are defined and write
each as a composition of functions given the information below.
(None of the functions in your compositions should be the same as
the given function.)

$BR
${BBOLD}1.$EBOLD \($ctex^{$n}(x)=f(g(x))\) where

$BR
\(f(x)=\) \{ $mp1->ans_rule(10) \}, and \(g(x)=\) \{ $mp1->ans_rule(10) \}.

$BR
${BBOLD}2.$EBOLD \($ctex($ctex(x))=f(g(x))\) where

$BR
\(f(x)=\) \{ $mp2->ans_rule(10) \}, and \(g(x)=\) \{ $mp2->ans_rule(10) \}.

$BR
${BBOLD}3.$EBOLD \($ctex x^{$n}=f(g(x))\) where

$BR
\(f(x)=\) \{ $mp3->ans_rule(10) \}, and \(g(x)=\) \{ $mp3->ans_rule(10) \}.

END_TEXT
Context()->normalStrings;

# ANS( Compute("x^$n")->cmp() );
# ANS( Compute("$c(x)")->cmp() );
# ANS( Compute("x^$n")->cmp() );

ANS( $mp1->cmp() );
ANS( $mp2->cmp() );
ANS( $mp3->cmp() );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
\($ctex^{$n}(x) =($ctex(x))^{$n}\) means you find \($ctex(x)\) and
then take that to the $n.
$BR
\($ctex($ctex(x))\) means find \($ctex(x)\) and then take the $d of that.
$BR
\($ctex x^{$n}\) by convention is \($ctex(x^{$n})\), which means you
take \(x\) to the $n first and then take the $d.
$BR
Expressing each as a composition: If \(h(x)=$ctex(x)\) and \(k(x)=x^$n\),
then:
$BR
${BBOLD}1.$EBOLD
\($ctex^{$n}(x) =k(h(x))\)
$BR
${BBOLD}2.$EBOLD
\($ctex($ctex(x))=h(h(x))\)
$BR
${BBOLD}3.$EBOLD
\($ctex x^{$n}=h(k(x))\)

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
