##DESCRIPTION
##ENDDESCRIPTION
##Tagged by Nelson Chen
##July 31, 2016.


## DBsubject(Statistics)
## DBchapter(Hypothesis Testing)
## DBsection(Elements of a Statistical Test)
## Date(7/17/2012)
## Institution(UBC)
## Author(Djun Kim)
## Level(5)
## MO(1)
## KEYWORDS('inference', 'hypothesis test', 'difference in two means')


########################################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",	   # Standard macros for PG language
    "MathObjects.pl",
    "PGchoicemacros.pl",
    "parserRadioButtons.pl",
    "unionTables.pl",
    "parserPopUp.pl",
    "RserveClient.pl",
    "answerHints.pl",
 "PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "PGcourse.pl"
    );



# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

@n1 = rserve_eval('
n1 <- round(runif(1, min=7, max=12))
');


@n2 = rserve_eval('
n2 <- round(runif(1, min=6, max=11))
');


@mu1 = rserve_eval('
mu1 <- round(runif(1, min=24, max=24.6), 1)
');


@mu2 = rserve_eval('
mu2 <- round(runif(1, min=20.7, max=21.3), 1)
');


@s1 = rserve_eval('
s1 <- round(runif(1, min=2.9, max=3.3), 1)
');

@s2 = rserve_eval('
s2 <- round(runif(1, min=2.7, max=3.1), 1)
');

@ansb = rserve_eval('
round(mu1-mu2, 1)
');

@var1 =  rserve_eval('
t <- ((n1-1)*s1^2 + (n2-1)*s2^2)/(n1+n2-2) 
round(t, 4)
'); 

@sd1 =  rserve_eval('
round(sqrt(t), 4)
'); 

@sd2 =  rserve_eval('
tem <- sqrt(t)*sqrt(1/n1 + 1/n2)
round(tem, 4)
'); 

@sd3 =  rserve_eval('
tem2 <- (mu1-mu2)/tem
round(tem2, 4)
'); 

@sd4 =  rserve_eval('
tem3 <- qt(0.975, n1+n2-2)
round(tem3, 4)
'); 


$mc_a = RadioButtons( 
    [
    "The difference in the mean BMI of fourth graders at Laura Secord and Charles Dickens schools.",
    "The difference in the mean BMI of the children taken in the samples from Laura Secord and Charles Dickens schools.",
    "The mean BMI of children at Laura Secord and Charles Dickens schools.",
    "The difference in the variances of BMI of the children taken in the samples from Laura Secord and Charles Dickens schools.",
    "The mean BMI of children at schools in Canada.",
    ],
    # correct answer
    "The difference in the mean BMI of fourth graders at Laura Secord and Charles Dickens schools.",
    last => [], # can be a list
    );


$ans_b = @ansb[0];

$multians = MultiAnswer($ans_b)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $a_stu ) = @{$student};
      my ( $f1 ) = @{$correct};
      if (
         $a_stu == $ans_b     ||
          $a_stu == -$ans_b 
          )
      {
         return [1];
      }
      else {
         return [0];  
      }

  }
);


$mc_c =  RadioButtons(
    [
     "There is a difference in the mean BMI for fourth graders at the two schools.",
     "There is no difference in the mean BMI for fourth graders at the two schools.",
     "The difference between the mean BMI of fourth graders at the two schools is @ansb[0].",
     "The mean difference between the BMI of a fourth grader at Laura Secord and the BMI of a fourth grader at Charles Dickens is @ansb[0].",     
     "The mean BMI of fourth graders at Laura Secord is greater than the mean BMI of fourth graders at Charles Dickens.",
    ],
    # correct answer
    "There is no difference in the mean BMI for fourth graders at the two schools.",
    last => [], # can be a list
    );

$mc_d = RadioButtons( 
   [ 
"two-sided",

"one-sided, left-tailed",

"one-sided, right-tailed",

"it does not matter whether we take a one-sided or two-sided alternative",
    ],
    # correct answer
    "two-sided",

    last => [], # can be a list
);


#------------------------------------------------------------
# setup for part e
@eqn = (
    "\( z_0 = \frac{ \bar{y} - \mu_0 } {\sigma / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{y} - \mu_0}{s / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{d}} {s_d / \sqrt{n}} \)",
    "\( t_0 = \frac{\bar{y}_1 - \bar{y}_2} {\mathrm{SE} (\bar{y}_1 - \bar{y}_2)} \)",
    );

@letter = ("A", "B", "C", "D");

$k = random(0,3,1);

@perm = shuffle(4);
@eqn = @eqn[@perm];
@inv = invert(@perm);

# The last (3rd) equation in the list above is the correct choice.
$popup = PopUp(["?","A","B","C","D"], $letter[$inv[3]]);

#------------------------------------------------------------
# setup for part f


$mc_f =  RadioButtons( 
  [

"There is sufficient evidence to suggest that there is a difference in the mean BMI for fourth graders at the two schools.",   

   "There is sufficient evidence to suggest that the difference between the mean BMI of fourth graders at the two schools is @ansb[0].",

   "There is sufficient evidence to suggest that the mean difference between the BMI of a fourth grader at Laura Secord and the BMI of a fourth grader at Charles Dickens is @ansb[0].",

   "There is sufficient evidence to suggest that the mean BMI of fourth graders at Laura Secord is greater than the mean BMI of fourth graders at Charles Dickens.",

   "There is insufficient evidence to suggest that there is a difference in the mean BMI for fourth graders at the two schools.",

   "There is insufficient evidence to suggest that the difference between the mean BMI of fourth graders at the two schools is @ansb[0].",

   "There is insufficient evidence to suggest that the mean difference between the BMI of a fourth grader at Laura Secord and the BMI of a fourth grader at Charles Dickens is @ansb[0].",

   "There is insufficient evidence to suggest that the mean BMI of fourth graders at Laura Secord is greater than the mean BMI of fourth graders at Charles Dickens.",
],

"There is sufficient evidence to suggest that there is a difference in the mean BMI for fourth graders at the two schools.",
    last => [], # can be a list 
    );


##############################################################
#
#  Text
#
#

TEXT(beginproblem());


Context()->texStrings;
BEGIN_TEXT
$BEGIN_ONE_COLUMN
One definition of obesity is in terms of body mass index (BMI). In a
study to compare obesity levels in two schools, random samples of fourth
graders were taken from each school and body mass indices recorded. Here is
a summary of the body mass index data from two schools, Laura Secord and
Charles Dickens.

$PAR
$BCENTER
\{ begintable(4) \}
\{ row( "School", "Number of children measured", "Average body mass", "SD of the body masses" ) \}
\{ row( "Laura Secord", "@n1[0]", "@mu1[0]", "@s1[0]" ) \}
\{ row( "Charles Dickens", "@n2[0]", "@mu2[0]", "@s2[0]" ) \}
\{ endtable() \}
$ECENTER
$END_ONE_COLUMN


$PAR
$BBOLD Part (a)$EBOLD What is the parameter of interest in this study?$BR

$BR
\{ $mc_a -> buttons() \}

$BR
$BR
$BBOLD Part (b)$EBOLD Based on the data provided, what is your estimate of this parameter (to one decimal place)?  \{ ans_rule(10) \}

$BR
$BR 
$BBOLD Part (c)$EBOLD In testing a hypothesis about the parameter of interest, what would your null hypothesis be?

$BR 
$BR 

\{ $mc_c -> buttons() \}

$BR 
$BR 
$BBOLD Part (d)$EBOLD You would take the alternative hypothesis to be:

$BR 
$BR 
\{ $mc_d -> buttons() \}

$BR 
$BR 
$BBOLD Part (e)$EBOLD  Which of the following test statistics would you choose to test your hypotheses?  \{ $popup->menu() \}

$BR
$BR
$BCENTER
\{
BeginTable().
  AlignedRow(["$BBOLD A $EBOLD", $eqn[0], "$BBOLD B $EBOLD" , $eqn[1]]).
  TableSpace(5,5).
#  AlignedRow(["A","B"]).
  TableSpace(25,6).
  AlignedRow(["$BBOLD C $EBOLD" , $eqn[2], "$BBOLD D $EBOLD" , $eqn[3]]).
  TableSpace(5,5).
#  AlignedRow(["C","D"]).
EndTable();
\}
$BR
$ECENTER


$BR 
$BBOLD Part (f)$EBOLD If you use a 5$PERCENT level of significance, which of the following would you conclude?

$BR
$BR
\{ $mc_f -> buttons() \}

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

#install_problem_grader(~~&std_problem_grader);

ANS( $mc_a->cmp() );
ANS( $multians->cmp() );
ANS( $mc_c->cmp() );
ANS( $mc_d->cmp() );
ANS( $popup->cmp() );  # e)
ANS( $mc_f->cmp() );

Context()->normalStrings;

COMMENT('MathObject version.');

## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
(a) The parameter of interest here is the difference in the
mean BMI of fourth graders at Laura Secord and Charles Dickens schools. 
$BR
$BR

(b) The estimate of this parameter from the data is @mu1[0]-@mu2[0]= @ansb[0]
(or -@ansb[0]). 
$BR
$BR

(c) The null hypothesis is that there is no difference in the mean BMI
for fourth graders at the two schools. 
$BR
$BR

(d) The alternative hypothesis would be two-sided, as there is no
information to suggest one school would have higher mean BMI than the other
before the study.
$BR
$BR

(e) The correct test statistic standardises the difference in the
sample means, \(\bar{y}_{1}-\bar{y}_{2}\), by an estimate of the
standard deviation of that difference. 
$BR
$BR


(f) We use a two sample t test with pooled standard deviation, since 
@n1[0] and @n2[0] are small and the two sample standard deviations
are similar. The test statistic is \( t=(\bar{y}_{1}-\bar{y}_{2})/SE(\bar{y}
_{1}-\bar{y}_{2}) \). To calculate the SE, we first find the pooled
estimate of the common unknown population variances:

\begin{eqnarray*}
s_{p}^{2} &=& \frac{(@n1[0]-1)\times @s1[0]^{2}+(@n2[0]
-1)\times @s2[0]^{2}}{@n1[0]+@n2[0]-2} \\ 
&=& @var1[0].
\end{eqnarray*}

The standard error is
\[
SE=s_{p}\sqrt{1/@n1[0]+1/@n2[0]}= @sd2[0].
\]

So the t-statistic is
\[
t=\frac{\bar{y}_{1}-\bar{y}_{2}}{SE}=\frac{@mu1[0]-@mu2[0]}{SE}= @sd3[0].
\]

Using software or a t table, we find the 97.5% point of the \(t_{
@n1[0]+@n2[0]-2}\) distribution to be @sd4[0]. The two-sided test at 5% significance level rejects the null
hypothesis if \( |t|> @sd4[0]\). Otherwise there
is insufficient evidence to reject the null hypothesis that the mean BMI of
fourth graders at the two schools are equal.
END_SOLUTION


ENDDOCUMENT();        