#DESCRIPTION
## DBsubject(Statistics)
## DBchapter(Hypothesis tests)
## DBsection(Independent samples - t)
## Level(3)
## MO(1)
## KEYWORDS('Two sample t test', 'hypothesis tests', 'P-value')
##ENDDESCRIPTION
##Tagged by Nelson Chen, Dec 30, 2016
##Attempts: Five. 

DOCUMENT();

loadMacros(
"PGstandard.pl",
  "PGchoicemacros.pl",
  "parserRadioButtons.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "PGcourse.pl",
 "MathObjects.pl",
 "parserMultiAnswer.pl",
"Parser.pl",
);

#######
@m1 = rserve_eval('
m1 <- round(runif(1, min=0.720, max=0.729), 3)
');

@m2 = rserve_eval('
m2 <- round(runif(1, min=0.734, max=0.739), 3)
');

@s1 = rserve_eval('
s1 <- round(runif(1, min=0.02, max=0.025), 4)
');

@s2 = rserve_eval('
s2 <- round(runif(1, min=0.019, max=0.02), 4)
');

@ans1 = rserve_eval('
round(m1-m2, 3)
');

@ans1al = rserve_eval('
round(m2-m1, 3)
');

$ans1=@ans1[0];
$ans2=@ans1al[0];

@var = rserve_eval('
var <- (23*s1^2 + 34*s2^2)/(24+35-2)
round(var, 6)
');

@sd = rserve_eval('
sd = sqrt(var)
round(sd, 5)
');

@se = rserve_eval('
se = sd*sqrt(1/24+1/35) 
round(se, 4)
');

@t1 = rserve_eval('
tem1 = (m1-m2)/se
round(tem1, 3)
');

@t2 = rserve_eval('
tem2 = (m2-m1)/se
round(tem2, 3)
');

$t1=@t1[0];
$t2=@t2[0];

#######P-value
@p = rserve_eval('
p<-round(2*pt(abs(tem1), df=57, lower.tail=FALSE),3)
');

$mc = new_multiple_choice();
$mc->qa(
"In performing a hypothesis test here, what would you take as the null
hypothesis?",
"There is no difference in the mean humerus lengths between the
sparrows that died and those that survived."
);
$mc->extra(
"There is a difference in the mean humerus lengths between the sparrows
that died and those that survived.",
"There is no difference in the sample mean humerus lengths between the
sparrows that died and those that survived.",
"There is a difference in the sample mean humerus lengths between the
sparrows that died and those that survived.",
"On average, bone length makes no difference to the chance of a bird
dying in a storm. ",
);

$mc1 = new_multiple_choice();
$mc1->qa(
"What type of alternative hypothesis would you take here?",
"Two--sided"
);
$mc1->extra(
"One--sided",
"It does not matter",
);

if ( @p[0] < 0.05 ) {
$corr = "Reject";
$extra = "Do not reject";
} else {
$corr = "Do not reject";
$extra = "Reject";
}

$mc2 = new_multiple_choice();
$mc2->qa(
"Would you reject or not reject your null hypothesis at the 5% significance level?",
$corr
);
$mc2->extra(
$extra
);

#######
BEGIN_TEXT
In the nineteenth century Darwin's theory of natural selection was not
widely accepted. A researcher in 1898 in Massachusetts presented data he had
gathered to explore natural selection. There had been an extreme storm the
winter of that year, and a sample of house sparrows had been brought to
Brown University to study. Some of these sparrows had survived the storm,
and some had perished. Of the adult males in the sample, 24 had perished
during the storm and 35 had survived. It was of interest to determine if
there were differences in the physical characteristics between the sparrows
that survived and those that died. The humerus bone was measured on each
bird (in inches). Suppose the sparrows collected that died had mean humerus
length @m1[0] inches with standard deviation @s1[0] inches.
Suppose the corresponding measurements for the sparrows that survived were 
@m2[0] inches and @s2[0] inches.
$BR
$BR
$BBOLD (a) $EBOLD Interest resides in the difference in the mean humerus lengths of male sparrows between those that died and those that survived the storm. Denote
this difference \(\mu_{1}-\mu_{2}.\) Based on the data provided, what is
your estimate of this parameter? \{ ans_rule(6) \}
$BR
$BR
$BBOLD (b) $EBOLD
\{ $mc->print_q() \}
$BR
\{ $mc->print_a() \}
$BR
$BR
$BBOLD (c) $EBOLD
\{ $mc1->print_q() \}
$BR
\{ $mc1->print_a() \}
$BR
$BR
$BBOLD (d) $EBOLD Neither data set have distributions that appear skewed. As the sample variances are similar it would be reasonable to perform the pooled
two--sample t test here. Find the test statistic for a t test of the null
hypothesis to three decimal places. \{ ans_rule(6) \}
$BR
$BR
$BBOLD (e) $EBOLD Provide the P-value of your test to 3 decimal places. \{ ans_rule(6) \}
$BR
$BR
$BBOLD (f) $EBOLD 
\{ $mc2->print_q() \}
$BR
\{ $mc2->print_a() \}
$BR
$BR
END_TEXT

#######

$showPartialCorrectAnswers = 1;

ANS(Real($ans1)->cmp(checker => sub {
      my ($correct,$student) = @_;
      return $correct == $student || $correct == -$student;
    }));

ANS( radio_cmp( $mc->correct_ans() ) );
ANS( radio_cmp( $mc1->correct_ans() ) );

ANS(Real($t1)->cmp(checker => sub {
      my ($correct,$student) = @_;
      return $correct == $student || $correct == -$student;
    }));

ANS( num_cmp(@p[0], tol=> 0.0105) );
ANS( radio_cmp( $mc2->correct_ans() ) );

## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
(a) The estimate would be the difference in the sample
means, @m1[0]-@m2[0] = @ans1[0] inches ( or alternatively 
@m2[0] - @m1[0] = @ans1al[0] inches). 
$BR
$BR
(b) The null hypothesis would be that there is no difference in the
mean humerus lengths between the sparrows that died and those that survived.
That is, \(H_{0}:\mu _{1}-\mu _{2}=0.\)
$BR
$BR
(c) We would take the alternative hypothesis as being \(H_{A}:\mu
_{1}-\mu _{2}\neq 0\).There is no reason to believe prior to the
study that the birds that survived would differ in bone lengths compared to
those that perished.
$BR
$BR
(d) For the pooled t-test, the test statistic is \(t=(\bar{x}_{1}-\bar{
x}_{2})/SE(\bar{x}_{1}-\bar{x}_{2})\). To calculate the standard
error, SE, we find the pooled estimate of the unknown variance: 
\begin{eqnarray*}
s_{p}^{2} &=&\frac{(24-1)\times @s1[0]^{2}+(35-1)\times @s2[0]^{2}
}{24+35-2} \\
&=& @var[0] \ \ \Rightarrow \ \ s_{p}= @sd[0].
\end{eqnarray*}
The standard error is 
$BR
$BR
$BCENTER
\(
SE=s_{p}\sqrt{1/24+1/35}=@se[0].
\)
$ECENTER
$BR
$BR
So the t-statistic is
$BR
$BR
$BCENTER
\(
t=\frac{\bar{x}_{1}-\bar{x}_{2}}{SE}=\frac{@m1[0]-@m2[0]}{SE}=$t1.
\)
$ECENTER
$BR
$BR
The test statistic can also be found via \(\left( @m2[0]-@m1[0]\right) /SE \).
$BR
$BR

(e) We should compare the test statistic with the t distribution on 57
degrees of freedom. As the test is two--sided, the P-value is @p[0].  
$BR
$BR
(f) If testing at the 5% significance level, if @p[0] >0.05 there is not sufficient evidence to reject the null hypothesis that there is no difference in the mean humerus lengths between sparrows that survived the storm and those that died. Otherwise we would
reject the null hypothesis.
END_SOLUTION

ENDDOCUMENT();
