##DESCRIPTION
##
## Created by the UBC School of Biomedical Engineering and Department of Mechanical Engineering. Project led by Robyn Newell and Agnes d'Entremont.
##
## Contact: robyn.newell@ubc.ca and agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this 
## project provided by UBC Open Education Resources Fund. We also gratefully acknowledge
## additional funding support from BCcampus and the School of Biomedical
## Engineering.  
## Contributors to this project include Thomas Oxland, Peter Cripton, Dave Wilson, Lyndia Wu,
## Pranav Shrestha, Jackson Robinson, Jacqueline Kho, Laura Vargas, Mona Behrouzian
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people. 
##
##ENDDESCRIPTION

## DBsubject(Engineering Biomechanics Statics)
## DBchapter(Mechanics of the Human Body)
## DBsection(Joint Reaction Forces and Muscle Forces)
## Institution('University of British Columbia')
## Author('UBC-SBME')
## Date(July 2021)
## Level()
## KEYWORDS()
## RESOURCES('UBC-BME-STA-20-002-1.png','UBC-BME-STA-20-002-2.png', 'UBC-BME-STA-20-002-3.png', 'UBC-BME-STA-20-002-4.png')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
## Beginner

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "scaffold.pl",
   "parserMultiAnswer.pl",
   "extraAnswerEvaluators.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

##############################################################
Scaffold::Begin();
BEGIN_TEXT
Two critical dimensions of hip prostheses are the caput-collum-diaphyseal (CCD) angle and neck length (L). To visualize this, please reference the coordinate system found \{htmlLink("http://www.wiki.orthoload.com/index.php/File:CoorSysAtLeftFemu_big.png", "here")\}. 
$BR
Use the \{htmlLink("https://orthoload.com/database/", "OrthoLoad Database")\} to determine the hip abductor muscle force, angle of muscle force action, and angle of leg during single leg standing. 
$PAR
\{image("UBC-BME-STA-20-002-3.png", width =>200, height =>200)\}
$PAR
Patient information is listed in the table:
$BR
\{image("UBC-BME-STA-20-002-2.png", width =>500, height =>55)\}
$PAR
Assume the hip abductor muscle inserts at point P1 (see coordinate system link) and the length of the leg (from P1 to the foot) is \(0.52H\). You can also assume the weight of the leg is \(0.17W\) and that the location of the center of mass of the leg is \(0.2H\) distal to the point P1 (along the length of the leg). 
$PAR
Note: The femur is not necessarily perfectly vertical as represented in the coordinate system diagram, and that the z axis aligns with the femoral shaft! $BR
\{ 
knowlLink("Hint", 
value=>image("UBC-BME-STA-20-002-4.png", width =>500, height =>55)) 
\}
END_TEXT

Section::Begin();
Context("Numeric");

$z_high = 240;
$z_low = 180;
$x_high = 25;
$x_low = 10;

BEGIN_TEXT
Using Orthoload, find \(F_z\) and \(F_x\) on the hip joint during single leg standing. 
$BR
Use these search terms in OrthoLoad: 
$BR
\{image("UBC-BME-STA-20-002-1.png", width =>450, height =>80)\}
$BR
and play the video file in the "Now Selected" box. Examine the joint reaction force during single leg stance for patient PFL. 
$PAR
Note: All forces in Orthoload are reported in a right-handed coordinate system of the left femur. The load components are reported as -Fx, -Fy, -Fz with negative signs. Positive values therefore indicate components acting towards the femoral head. In older videos, such as the PFL video, the minus signs are lacking.
$PAR
\(F_z\) = \{ans_rule(10)\} %BW
$PAR
\(F_x\) = \{ans_rule(10)\} %BW 
END_TEXT

ANS(Real($z_high)->cmp(
   checker => sub{
   my ($correct, $student, $ansHash) = @_;
   my ($sz) = $student->value;
   if (($sz < $z_high) && ($sz > $z_low)){
   $stuZ = $sz;
   return 1;
   }
   else {return 0;}
 }));
   
ANS(Real($x_high)->cmp(
   checker => sub{
   my ($correct, $student, $ansHash) = @_;
   my ($sx) = $student->value;
   if (($sx < $x_high) && ($sx > $x_low)){
   $stuX = $sx;
   return 1;
   }
   else {return 0;}
 }));   
Section::End(); 

Section::Begin();
$W = 980; 
$H = 1.75;
$CCD = 135;
$L = 0.062;

$a = $L*cos(($CCD-90)*pi/180);
$b = $L*sin(($CCD-90)*pi/180);

$Fjz = $stuZ / 100 * $W;
$Fjx = $stuX / 100 * $W;

$L_leg = 0.52*$H;
$W_femur = 0.17*$W;
$COM = 0.2*$H; 

$beta = acos(((-$Fjz*$a)+($Fjx*$b))/(($W_femur*$COM)-($W*$L_leg)));
$beta_deg = $beta*180/pi;

$theta = atan(($Fjz+($W_femur*sin($beta))-($W*sin($beta)))/($Fjx+($W_femur*cos($beta))-($W*cos($beta))));
$theta_deg = $theta*180/pi;

$Fm = ($Fjx+($W_femur*cos($beta))-($W*cos($beta)))/cos($theta);

BEGIN_TEXT
Using your values obtained from Orthoload, solve for the hip abductor muscle force and angles of muscle force action (with respect to the x axis) and leg (with respect to the horizontal) during single leg stance.
$PAR
Hip Abductor Muscle Force: \{ ans_rule(10)\} N
$PAR
Angle of Muscle Force with respect to the x axis: \{ ans_rule(10)\} degrees
$PAR
Angle of Leg with respect to the horizontal:  \{ ans_rule(10)\} degrees
END_TEXT

ANS(num_cmp($Fm, , reltol => 0.5));
ANS(num_cmp($theta_deg, , reltol => 0.5));
ANS(num_cmp($beta_deg, , reltol => 0.5));

Section::End();
Scaffold::End();
ENDDOCUMENT();