## DESCRIPTION
## 
## ENDDESCRIPTION


## DBsubject(Trigonometry)
## DBchapter(Polar coordinates & vectors)
## DBsection(Curves)
## Date(6/19/2018)
## Institution(Red Rocks Community College, Colorado Community College System)
## Author(Brenda Forland)
## MO(1)
## KEYWORDS('trigonometry')


###########################
#  Initialization

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"AnswerFormatHelp.pl",
"parserMultiAnswer.pl",
"PGML.pl",
"PGcourse.pl",
"parserPopUp.pl",
"contextFraction.pl",
);

TEXT(beginproblem());

###########################

# Setup
Context("Numeric");
#Context()->variables->are('theta'=>['Real', TeX=>'\theta',limits => [*pi/6,pi/3]]);
Context()->variables->add(y =>'Real');
Context()->noreduce('(-x)-y','(-x)+y');


$a = random(2,9,1);
$b = random(2,9,1);

$fac1 = Formula("-1/$b*x+$a/$b")->reduce;
$fac2 = Formula("y")->reduce;

$multians = MultiAnswer($fac1, $fac2)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
      if ( ($f1 == $f1stu && $f2 == $f2stu) ||
           ($f1 == $f2stu && $f2 == $f1stu) ) {
          return [1,1];
      } else {
          if ($f1 == $f1stu || $f2 == $f1stu) {
              return [1,0];
          } elsif ($f1 == $f2stu || $f2 == $f2stu) {
              return [0,1];
          } else {
              return [0,0];
          }
      }
  }
);


$popup = PopUp(["?", "ellipse", "hyperbola", "parabola", "not a conic"], "not a conic");

###########################
#  Main text

BEGIN_PGML
Convert the given polar equation to a Cartesian equation. Write in the standard form of a conic, if possible, and identify the conic section represented. 

[`` r = \frac{[$a]}{\cos(\theta)+[$b]*\sin(\theta)} ``]

END_PGML

BEGIN_TEXT
$BR
$BR
\{$multians->ans_rule(20)\}
\( = \) 
\{$multians->ans_rule(5)\}
$BR
$BR
The conic section is a(an):  \{ $popup->menu() \}
$BR
$BR
$BR
END_TEXT

BEGIN_PGML
Note that standard forms of conics are [` \frac{(x-h)^2}{a^2}+\frac{(y-k)^2}{b^2}=1 `], [` \frac{(x-h)^2}{a^2}-\frac{(y-k)^2}{b^2}=1 `], [` (y-k)^2=4p(x-h) `], and [` (x-h)^2=4p(y-k) `], and that a circle is an ellipse. If it is a line, enter as [` y = mx+b `].
END_PGML

ANS( $multians->cmp() ); 
ANS( $popup->cmp() );
############################
#  Solution

#BEGIN_PGML_SOLUTION
#Solution explanation goes here.
#END_PGML_SOLUTION

COMMENT('MathObject version. Uses PGML.');

ENDDOCUMENT();

