##DESCRIPTION
##
## Created by the UBC School of Biomedical Engineering and Department of Mechanical Engineering. Project led by Robyn Newell and Agnes d'Entremont.
##
## Contact: robyn.newell@ubc.ca and agnes.dentremont@mech.ubc.ca
##
## This work is licensed under the Creative Commons
## Attribution-ShareAlike 4.0 International (CC BY-SA 4.0)
## <https://creativecommons.org/licenses/by-sa/4.0/>
##
## We gratefully acknowledge the financial support for this 
## project provided by UBC Open Education Resources Fund. We also gratefully acknowledge
## additional funding support from BCcampus and the School of Biomedical
## Engineering.  
## Contributors to this project include Thomas Oxland, Peter Cripton, Dave Wilson, Lyndia Wu,
## Pranav Shrestha, Jackson Robinson, Jacqueline Kho, Laura Vargas, Mona Behrouzian
##
## This work was completed on the traditional, ancestral, and unceded
## territory of the Musqueam people. 
##
##ENDDESCRIPTION

## DBsubject(Engineering Biomechanics Dynamics)
## DBchapter(Applications: Injury Biomechanics)
## DBsection(Experiments and Instruments)
## Institution('University of British Columbia')
## Author('UBC-SBME')
## Date(July 2021)
## Level()
## KEYWORDS()
## RESOURCES('UBC-BME-INJ-20-001.png')
## TitleText('TBD')
## EditionText('TBD')
## AuthorText('TBD')
## Section('TBD')
## Problem('TBD')
## Beginner

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   #"source.pl",        # allows code to be displayed on certain sites.
   #"PGcourse.pl",      # Customization file for the course
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$max = 400;
$min = 350;

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

Estimate the impact duration \(\Delta t\) for the following crash test resultant linear acceleration data:
$PAR
\{image("UBC-BME-INJ-20-001.png", width=>704, height=>432)\}
$PAR
\(\Delta t\) = \{ans_rule(10)\} ms
END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#
ANS(Real($min)->cmp(
   checker => sub{
   my ($correct, $student, $ansHash) = @_;
   my ($sx) = $student->value;
   if (($sx <= $max) && ($sx >= $min)){
   return 1;
   }
   else {return 0;}
 }));

ENDDOCUMENT();        
