##DESCRIPTION
## Simplify Exponents
## 
##ENDDESCRIPTION
## DBsubject('Intermediate Algebra')
## DBchapter('Ch 01: Algebra, Mathematical Models and Problem Solving')
## DBsection('Properties of Integral Exponents')
## KEYWORDS('real numbers','algebraic expressions')
## TitleText1('Essentials of Intermediate Algebra')
## EditionText1('1')
## AuthorText1('Blitzer')
## Section1('1.6')
## Problem1('')
## Author('RA Cruz')
## Institution('The College of Idaho')
## Date: 2007/09
# Pulled from NPL: Library/CollegeOfIdaho/setAlgebra_01_06_Exponents/16IntAlg_35_Exponents.pg
# Edited for Periodic Re-randomization by Rick Lynch @ Mizzou 
# Significant changes made also by Rick Lynch

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "problemRandomize.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
Context("Numeric");

######################################
#  Setup

@alphabet = ("a","b","x","y");
$choose = random(0,2,2);
@var = ($alphabet[$choose], $alphabet[$choose+1]);
Context()->variables->are($var[0]=>'Real',
                          $var[1]=>'Real');

$n1 = random(4,9,1);
$n2 = random(-3,-1,1);
$n3 = random(2,4,1);
$n4 = random(5,9,1);
$coeff = random(2,9,1);
$power = list_random(-1,2);

$zchoice  = random(0,1,1);
$choice = ($zchoice == 0) ? 0 : random(1,3,1);
if ($choice == 1) {
  $n2 = 0;
} elsif ($choice == 2) {
  $n3 = 0;
} elsif ($choice == 3) {
  $n4 = 0;
} else {
  $power = 0;
}

$frac = "\left( \frac{$var[0]^{$n1}$var[1]^{$n2}}{$coeff $var[0]^{$n3}$var[1]^{$n4}} \right)^{$power}";

if ($power == 2) {
  $num = Formula("$var[0]^{2*($n1-$n3)}")->reduce;
  $den = Formula("(($coeff)**2) * $var[1]^{2*($n4 - $n2)} ")->reduce;
} elsif ($power == -1) {
  $num = Formula("(($coeff)) * $var[1]^{($n4 - $n2)} ")->reduce;
  $den = Formula("$var[0]^{($n1-$n3)}")->reduce;
} elsif ($choice == 0) {
  $num = Formula("1")->reduce;
  $den = Formula("1")->reduce;
}
LimitedPowers::OnlyPositiveIntegers();
$multians = MultiAnswer($num, $den)->with(
  singleResult => 0,
  allowBlankAnswers => 1,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $f1stu, $f2stu ) = @{$student};
      my ( $f1, $f2 ) = @{$correct};
 
      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
          return [1,1];
      } elsif ( $f1==$f1stu ) {
          return [1,0];
      } elsif ( $f2==$f2stu ) {
          return [0,1];
      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
          $self->setMessage(1,"Simplify your answer further");
          $self->setMessage(2,"Simplify your answer further");
          return [0,0];
      } else {
          return [0,0];
      }
  }
);
Context()->operators->undefine("+","-","/");
Context()->{error}{msg}{"Operands of '*' can't be words"} = " ";


if ($displayMode eq 'TeX') {
  $displayfrac =
  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
} else {
  $displayfrac =
  ColumnTable(
  "\( \displaystyle $frac = \)",
  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
  indent => 0, separation => 10, valign => "MIDDLE"
  );
}


###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following fraction using positive exponents only, and simplify.
$PAR
$BCENTER
$displayfrac
$ECENTER


END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;
Context()->operators->undefine("+","-");
ANS($multians->cmp() );

Context()->texStrings;
if ($choice != 0) {
  $texanswer = "\displaystyle \frac{$num}{$den}";
} else {
  $texanswer = 1;
}
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \($texanswer\)
  $PAR
  END_TEXT
} else {
  BEGIN_TEXT
$PAR
  ${BBOLD}Note: $EBOLD  If your answer is not a fraction, write "1" in the denominator.
$BR
  END_TEXT
}
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
