##DESCRIPTION
##   Evaluation of a simple algebraic expression
##ENDDESCRIPTION

##KEYWORDS('algebra','exponent')
## hcao tagged and PAID on 12-12-2003

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Exponents and Radicals')
## Date('6/3/2002')
## Author('')
## Institution('')
## TitleText1('College Algebra')
## EditionText1('3')
## AuthorText1('Stewart, Redlin, Watson')
## Section1('1.3')
## Problem1('23')
# Pulled from NPL: Library/Rochester/setAlgebra02ExponentsRadicals/sw1_3_23.pg
# Edited for Periodic Re-randomization by Rick Lynch @ Mizzou 
# Significant changes made also by Rick Lynch

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGunion.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "contextLimitedPowers.pl",
  "contextLimitedNumeric.pl",
  "problemRandomize.pl",
  "PeriodicRerandomization.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");
Context("Numeric")->variables->add(r=>"Real", s=>"Real");

$a = random(-1,1,2)*random(2,3,1);
$b = random(2,3,1)*random(-1,1,2);
$a0 = list_random(-10,-5,-4,-3,-2,2,3,4,5,10);

$c = random(3,8,1);
$c1 = random(2,10,1);
$d = random(4,6,1);
$e = random(4,6,1);

$frac = "($a0 r^{$c}s^{$c1})($a rs^{$d})($b r^{$e}) ";
$num = Formula("($a0)*($a)*($b)* r**($e+$c+1) s**($d+$c1)")->reduce;
# $den = Formula("1")->reduce;

LimitedPowers::OnlyPositiveIntegers();
#$multians = MultiAnswer($num, $den)->with(
#  singleResult => 0,
#  allowBlankAnswers => 1,
#  checker => sub {
#      my ( $correct, $student, $self ) = @_;
#      my ( $f1stu, $f2stu ) = @{$student};
#      my ( $f1, $f2 ) = @{$correct};
 
#      if ( ( $f1==$f1stu &&  $f2==$f2stu) ) {
#          return [1,1];
#      } elsif ( $f1==$f1stu ) {
#          return [1,0];
#      } elsif ( $f2==$f2stu ) {
#          return [0,1];
#      } elsif ( $f1*$f2stu==$f1stu*$f2 ) {
#          $self->setMessage(1,"Simplify your answer further");
#          $self->setMessage(2,"Simplify your answer further");
#          return [0,0];
#      } else {
#          return [0,0];
#      }
#  }
#);
Context()->operators->undefine("+","-","/");


#if ($displayMode eq 'TeX') {
#  $displayfrac =
#  "\[ $frac = ".$multians->ans_rule(10).$multians->ans_rule(10)." \]";
#} else {
#  $displayfrac =
#  ColumnTable(
#  "\( \displaystyle $frac = \)",
#  $multians->ans_rule(20).$BR.$HR.$multians->ans_rule(20),
#  indent => 0, separation => 10, valign => "MIDDLE"
#  );
#}

###################################
# Main text

Context()->texStrings;
BEGIN_TEXT

Write the following expression using positive exponents only:
$PAR
$BCENTER
\($frac =\) \{ans_rule(20)\}
$ECENTER


END_TEXT
Context()->normalStrings;

###################################
# Answers

$showPartialCorrectAnswers = 1;
Context()->operators->undefine("+","-");
ANS($num->cmp() );


Context()->texStrings;
###########################################################################
# use PeriodicRerandomization to write the answer and generate a new
# version of the problem
###########################################################################
if ($attempts_modp == 0 && $actualAttempts != 0) {
  BEGIN_TEXT
  $PAR
  ${BBOLD}Answer:${EBOLD} \(\displaystyle $num\)
  $PAR
  END_TEXT
} 
Context()->normalStrings;
PeriodicStatus(); 

COMMENT('Features Periodic Rerandomization');
ENDDOCUMENT();
