## DESCRIPTION
## Seting up a double integral in polar coordinates  
## ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Double integrals in polar)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(IdentifyLimitsGivenDomainGraph)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('16.4')
## Problem1('4')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('16.4')
## Problem2('4')
## KEYWORDS('Double Integral','Polar Coordinates')

DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGgraphmacros.pl",
  "unionTables.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

$refreshCachedImages=1;

##############################################
#  Setup

Context("Numeric")->variables->are(x=>"Real",y=>"Real");

$r = random(1,3,1);
$R = sqrt(2*$r*$r);

$gr = init_graph(-5.5,-5.5,5.5,5.5,axes=>[0,0],size=>[400,400]);
$gr->lb('reset');
$gr->lb( new Label(5.2,0.2,'x','black','center','middle'));
$gr->lb( new Label(0.2,5.2,'y','black','center','middle'));
# too lazy to change the color names...
$gr->new_color("lightblue",245,234,229); # really a light purple
$gr->new_color("darkblue", 159, 64, 16); # really a red-brown
#$gr->new_color("lightblue",214,230,244);
#$gr->new_color("darkblue", 100,100,255);
$gr->h_grid("gray",-5,-4,-3,-2,-1,1,2,3,4,5);
$gr->v_grid("gray",-5,-4,-3,-2,-1,1,2,3,4,5);
  foreach my $i (1..4) {
    $gr->lb( new Label($i,-5.2,$i,  'black','center','middle'));
    $gr->lb( new Label(-$i,-5.2,-$i,'black','center','middle'));
    $gr->lb( new Label(-5.2,$i,$i,  'black','center','middle'));
    $gr->lb( new Label(-5.2,-$i,-$i,'black','center','middle'));
  }


$select = random(0,3,1);


if ($select == 0) {

$A = Formula("pi/4");
$B = Formula("3*pi/4");
add_functions($gr,
"sqrt($R^2-x^2) for x in <-$r,$r> using color:darkblue and weight:2",
);
$gr->moveTo(-$r,$r);
$gr->lineTo(0,0,"darkblue",2);
$gr->lineTo($r,$r,"darkblue",2);
$gr->fillRegion([0.1,1,"lightblue"]);

} elsif ($select==1) {

$A = Formula("3*pi/4");
$B = Formula("5*pi/4");
$xfunc = sub { my $t = shift(); return $R*cos($t); };
$yfunc = sub { my $t = shift(); return $R*sin($t); };
$fn = new Fun( $xfunc, $yfunc, $gr );
$fn->domain(2.356,3.927);
$fn->steps(90);
$fn->weight(2);
$fn->color('darkblue');
#add_functions($gr,
#"sqrt($R^2-x^2)  for x in <-$R,-$r> using color:darkblue and weight:2",
#"-sqrt($R^2-x^2) for x in <-$R,-$r> using color:darkblue and weight:2",
#);
$gr->moveTo(-$r,$r);
$gr->lineTo(0,0,"darkblue",2);
$gr->lineTo(-$r,-$r,"darkblue",2);
$gr->fillRegion([-1,0.1,"lightblue"]);

} elsif ($select==2) {

$A = Formula("5*pi/4");
$B = Formula("7*pi/4");
add_functions($gr,
"-sqrt($R^2-x^2) for x in <-$r,$r> using color:darkblue and weight:2",
);
$gr->moveTo(-$r,-$r);
$gr->lineTo(0,0,"darkblue",2);
$gr->lineTo($r,-$r,"darkblue",2);
$gr->fillRegion([0.1,-1,"lightblue"]);

} else {

$A = Formula("-pi/4");
$B = Formula("pi/4");
$xfunc = sub { my $t = shift(); return $R*cos($t); };
$yfunc = sub { my $t = shift(); return $R*sin($t); };
$fn = new Fun( $xfunc, $yfunc, $gr );
$fn->domain(-0.785,0.785);
$fn->steps(90);
$fn->weight(2);
$fn->color('darkblue');
#add_functions($gr,
#"sqrt($R^2-x^2)  for x in <$r,$R> using color:darkblue and weight:2",
#"-sqrt($R^2-x^2) for x in <$r,$R> using color:darkblue and weight:2",
#);
$gr->moveTo($r,$r);
$gr->lineTo(0,0,"darkblue",2);
$gr->lineTo($r,-$r,"darkblue",2);
$gr->fillRegion([1,0.1,"lightblue"]);

}



$C = Formula("0");
$D = Formula("sqrt(2*$r*$r)")->reduce;



$multians1 = MultiAnswer($A, $B)->with(
  singleResult => 1, separator => ' to ', tex_separator => '\mbox{ to }',
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $Astu, $Bstu ) = @{$student};
#     my ( $a, $b ) = @{$correct};
      if ( (($Bstu - $Astu) == ($B - $A)) && 
           (cos($Astu) == cos($A)) &&
           (sin($Astu) == sin($A)) &&
           (cos($Bstu) == cos($B)) &&
           (sin($Bstu) == sin($B))
         ) { return 1; } else { return 0; }
  }
);


$multians2 = MultiAnswer($C, $D)->with(
  singleResult => 1, separator => ' to ', tex_separator => '\mbox{ to }',
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $Cstu, $Dstu ) = @{$student};
#     my ( $a, $b ) = @{$correct};
      if ( ($Cstu  == $C) && ($Dstu == $D)
         ) { return 1; } else { return 0; }
  }
);





##############################################
#  Main text

Context()->texStrings;
BEGIN_TEXT
\{
ColumnTable(
"Suppose \(R\) is the shaded region in the figure.
As an iterated integral in polar coordinates, ".
$PAR.
$BCENTER.
"\( \displaystyle
\iint\limits_R f(x,y) \, dA
=
\int_A^B \!\! \int_C^D f(r\cos(\theta),r\sin(\theta)) \, r \, dr \, d\theta
\)".
$ECENTER.
$PAR.
"with limits of integration".
$BR.
$BR.
"A = ".$multians1->ans_rule(20).$BR.
"B = ".$multians1->ans_rule(20).$BR.
"C = ".$multians2->ans_rule(20).$BR.
"D = ".$multians2->ans_rule(20).$BR,
$BCENTER.
image(insertGraph($gr),width=>300,height=>300,tex_size=>700).
$BR.
"(Click on graph to enlarge)".
$ECENTER,
indent=>0, separation=>30, valign=>"TOP"
)
\}
END_TEXT
Context()->normalStrings;


##################################################
#  Answers

ANS($multians1->cmp());
ANS($multians2->cmp());

$showPartialCorrectAnswers = 1;


;
ENDDOCUMENT();
