## DESCRIPTION
## Calculating double integrals in polar coordinates  
## ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Double integrals in polar)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('16.4')
## Problem1('24')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('16.4')
## Problem2('24')
## KEYWORDS('Double Integrals','Iterated Integrals','Polar Coordinates')

DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "contextFraction.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(r=>"Real",theta=>"Real",dr=>"Real",dtheta=>"Real");

#-----------------------
#  limits of integration
#  A < x < B
#  C < y < D

$a = random(3,5,1);
Context()->variables->set(r=>{limits=>[0,$a]});
$asq = $a**2;
$asq2 = 2 * $a**2;


#-----------
#  integrand
#

$f = Formula("(sqrt($asq2 - r**2) - r)*r") -> reduce;

#-----------
#  volume
#

$V = Formula("( 2*pi*$a**3 * (2**(3/2)-2) ) / 3") -> reduce;


#=================================
#  Multiple solutions
#
#  inner and outer differentials
#
#  arrays are (if r inner & theta outer integral, if theta inner & r outer integral)
# 
@id = (Formula("dr"),Formula("dtheta"));
@od = (Formula("dtheta"),Formula("dr"));

# inner limit high, inner limit low, outer limit high, outer limit low (case 0, case 1)
@ilh = (Formula("$a"),Formula("2*pi"));
@ill = (Formula("0"),Formula("0"));
@olh = (Formula("2*pi"),Formula("$a"));
@oll = (Formula("0"),Formula("0"));


$multipart = MultiAnswer( $f, $id[0], $od[0], $oll[0], $olh[0], $ill[0], $ilh[0] )->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $fstu, $idstu, $odstu, $ollstu, $olhstu, $illstu, $ilhstu ) = @{$student};
        if ( 
             ( 
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           return 1;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration.");
           return 0.94;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration and order of integration.");
           return 0.47;




        } elsif (
             ( 
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           $self->setMessage(1,"Check your integrand and dA.");
           return 0.47;
        } elsif (
             (
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration, your integrand, and dA.");
           return 0.47;
        } elsif (
             (
             $f      != $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      != $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration, order of integration, integrand, and dA.");
           return 0.24;
        } else {
             return 0;
        }
    }
  );



##############################################
#  Main text

Context()->texStrings;
Context()->flags->set(reduceConstants=>0);
BEGIN_TEXT

Consider the solid shaped like an ice cream cone 
that is bounded by the functions
\( z = \sqrt{x^2 + y^2} \) 
and 
\( z = \sqrt{$asq2 - x^2 - y^2} .\)
Set up an integral in polar coordinates to find
the volume of this ice cream cone.
$BR
$BR
${BITALIC}Instructions:${EITALIC}
Please enter the integrand in the first answer box,
typing ${BITALIC}theta${EITALIC} for \(\theta\). 
Depending on the order of integration you choose,  
enter ${BITALIC}dr${EITALIC} and ${BITALIC}dtheta${EITALIC} 
in either order into the second and third answer boxes 
with only one ${BITALIC}dr${EITALIC} or 
${BITALIC}dtheta${EITALIC} in each box.  Then, enter the limits of 
integration and evaluate the integral to find the volume. 
$BR
$BR
\( \displaystyle \int_A^B \int_C^D \) 
\{ $multipart->ans_rule(40) \}
\{ $multipart->ans_rule(5) \}
\{ $multipart->ans_rule(5) \}
$BR
$BR
A = \{ $multipart->ans_rule(20) \} $BR
B = \{ $multipart->ans_rule(20) \} $BR
C = \{ $multipart->ans_rule(20) \} $BR
D = \{ $multipart->ans_rule(20) \}
$BR
$BR
Volume = \{ ans_rule(40) \}


END_TEXT
Context()->normalStrings;





##################################################
#  Answers

ANS($multipart->cmp() );
ANS(($V)->cmp());

$showPartialCorrectAnswers = 1;

COMMENT('Allows integration in either order.');

$a3=$asq*$a;
Context("Fraction");
$m=Compute("2*$a3/3");

Context()->texStrings;
BEGIN_SOLUTION
$PAR SOLUTION $BR $BR
The functions \(z = \sqrt{x^2+y^2} \) and \( z =\sqrt{ $asq2-x^2-y^2}\) intersect when 
\( \sqrt{x^2+y^2} =\sqrt{ $asq2-x^2-y^2} \; \Rightarrow \; x^2+y^2 = $asq2 - x^2 - y^2 \;
\Rightarrow \; x^2+y^2 = $asq \). $BR 
Thus the region of integration is the disk $BR
\(\begin{array}{lcl} D& =& \left \lbrace (x,y) \, | \, x^2+y^2 \le $asq  \right\rbrace
\\ &=& \left \lbrace (r, \theta) \, | \, 0 \le \theta \le 2 \pi, \, 0 \le r \le $a 
\right \rbrace\end{array} \).
$BR
In polar coordinates,  the cone \( z = \sqrt{x^2+y^2} \) has 
equation \( z= r\), while the 
hemisphere \( z = \sqrt{$asq2- x^2- y^2} \) has equation \( z= \sqrt{$asq2-r^2}\), so
the volume is given by $BR
\(\begin{array}{lcl} V& =& \displaystyle \int_0^{2 \pi} \int_0^{$a}\left( \sqrt{$asq2-r^2}-r\right) \, r \, dr \, d\theta
\\ &=& \displaystyle 2 \pi \int_0^{$a} \left(\sqrt{$asq2-r^2}\,r-r^2 \right) \, dr  

\\ &=& \displaystyle2 \pi \int_0^{$a} \sqrt{$asq2-r^2}\,r \, dr -2 \pi \int_0^{$a}r^2 
 \, dr  
\quad [Substitution: u= $asq2 -r^2,\, du = - 2r dr ] \\
&=&-  \displaystyle \pi \int_{$asq2}^{$asq} u^{1/2} \, du  - 2\pi \left[\frac{r^3}{3}\right]_0^{$a}
\\ & = & \displaystyle \frac{2}{3} \pi \left[ u^{3/2} \right]_{$asq}^{$asq2} - 2 \pi \left(\frac{$a3}{3} \right)
\\ &=& \displaystyle \frac{2}{3} \pi \left( $asq2^{3/2} - $asq^{3/2} \right)
- 2\pi \left(\frac{$a3}{3} \right)
\\ & = &  \displaystyle $m \pi \left( 2^{3/2} -2\right) 
\end{array} \)

END_SOLUTION


ENDDOCUMENT();
