## DESCRIPTION
## Compute a basis for the range of a 3x3 matrix.
##   All matrix entries will be small real integers. The range will be 2-dimensional, so 2 basis vectors are expected.
##   There will be a solution (i.e., a basis) with small integer entries.
##   There are infinitely many different solutions (i.e., bases). They are all accepted.
## ENDDESCRIPTION

## DBsubject('Linear algebra')
## DBchapter('Euclidean spaces')
## DBsection('Row, column, and null spaces')
## Date('March 22, 2020')
## Author('Nathaniel Johnston')
## Institution('Mount Allison University')
## TitleText1('Introduction to Linear and Matrix Algebra')
## EditionText1('1')
## AuthorText1('Nathaniel Johnston')
## Section1('2.4: Bases and Rank')
## Problem1('')
## KEYWORDS('linear algebra','matrix','range','basis')
## http://www.njohnston.ca/publications/introduction-to-linear-and-matrix-algebra/

##############################################################
#  Initialization
##############################################################

DOCUMENT();      

loadMacros(
    "PGstandard.pl",
    "MathObjects.pl",
    "parserMultiAnswer.pl",
);

TEXT(beginproblem());

##############################################################
#  Problem Setup
##############################################################

Context("Vector");
Context()->lists->set(Vector => {open => "(", close => ")"});
Context()->parens->set("(" => {type => "Vector", close => ")"});

$flg = random(1,2);
if($flg == 1) {
    do {
        $c1 = random(-2,3);
        $c2 = random(-2,3);
        $c3 = random(-2,3);
        $c4 = random(-2,3);
        $c5 = random(-2,3);
        $c6 = random(-2,3);
        $c7 = random(-2,3);
        $c8 = random(-2,3);
        $c9 = random(-2,3);
        $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
        
        $r1 = random(-2,4);
        $r2 = random(-2,4);
        $R = Matrix([[1,0,$r1],[0,1,$r2],[0,0,0]]);
        
        $A = $C * $R;
        $v1 = Vector($A->column(1));
        $v2 = Vector($A->column(2));
    } until (abs($C->det) > 0.1);
} else {
    do {
        $c1 = random(-2,3);
        $c2 = random(-2,3);
        $c3 = random(-2,3);
        $c4 = random(-2,3);
        $c5 = random(-2,3);
        $c6 = random(-2,3);
        $c7 = random(-2,3);
        $c8 = random(-2,3);
        $c9 = random(-2,3);
        $C = Matrix([[$c1,$c2,$c3],[$c4,$c5,$c6],[$c7,$c8,$c9]]);
        
        $r1 = random(-2,4);
        $R = Matrix([[1,$r1,0],[0,0,1],[0,0,0]]);
        
        $A = $C * $R;
        $v1 = Vector($A->column(1));
        $v2 = Vector($A->column(3));
    } until (abs($C->det) > 0.1);
}
$CI = $C->inverse;
$covec = Vector($CI->row(3));

$multians = MultiAnswer($v1,$v2)->with(
    singleResult => 1,
    allowBlankAnswers => 0,
    separator => ', ', 
    tex_separator => ', ', 
    checker => sub {
        my ( $correct, $student, $answerHash ) = @_;
        my ( $s0, $s1 ) = @{$student};
        if(abs($s0 . $covec) < 0.001 && abs($s1 . $covec) < 0.001 && $s0->norm > 0.001 && $s1->norm > 0.001) {
            return 1 - ($s0->isParallel($s1));
        } else {
            return 0;
        }
    }
);

##############################################################
#  Text
##############################################################

Context()->texStrings;
BEGIN_TEXT
Compute a basis of the range of the matrix \( \displaystyle $A \).
$BR
$BR
\( \Bigg\lbrace \)\{$multians->ans_array(3)\}, \{$multians->ans_array(3)\}\( \Bigg\rbrace \).
END_TEXT
Context()->normalStrings;

##############################################################
#  Answers and Solutions
##############################################################

$showPartialCorrectAnswers = 1;

ANS( $multians->cmp() );

ENDDOCUMENT();
