## DESCRIPTION
## Calculating double integrals  
## ENDDESCRIPTION


## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Double integrals over general regions)
## Date(10/31/09)
## Institution(Fort Lewis College)
## Author(Paul Pearson)
## MLT(EvaluateDblIntOverTriangularRegion)
## MLTleader(1)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('4')
## Section1('16.2')
## Problem1('46')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('5')
## Section2('16.2')
## Problem2('46')
## KEYWORDS('Double Integrals','Iterated Integrals')

DOCUMENT();        

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);


TEXT(beginproblem());

##############################################
#  Setup

Context("Numeric")->variables->are(x=>"Real",y=>"Real",dx=>"Real",dy=>"Real");


#-----------------------
#  limits of integration
#  A < x < B
#  C < y < D

$a = random(4,10,2);



#-----------
#  integrand
#

$f = Formula("x*y");

#-----------
#  volume
#

$V = Formula("$a**4 / 96");


#=================================
#  Multiple solutions
#
#  inner and outer differentials
#
#  arrays are (if x inner & y outer integral, if y inner & x outer integral)
# 
@id = (Formula("dx"),Formula("dy"));
@od = (Formula("dy"),Formula("dx"));

# inner limit high, inner limit low, outer limit high, outer limit low (case 0, case 1)
@ilh = (Formula("$a - 2*y"),Formula("($a - x) / 2"));
@ill = (Formula("0"),Formula("0"));
@olh = (Formula("$a / 2"),Formula("$a"));
@oll = (Formula("0"),Formula("0"));


$multipart = MultiAnswer( $f, $id[0], $od[0], $oll[0], $olh[0], $ill[0], $ilh[0] )->with(
    singleResult => 1,
    checker => sub {
        my ( $correct, $student, $self ) = @_;
        my ( $fstu, $idstu, $odstu, $ollstu, $olhstu, $illstu, $ilhstu ) = @{$student};
        if ( 
             ( 
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu  
             )
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu  
             )
           ) {
           return 1;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             $ill[0] == $illstu &&
             $ilh[0] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             $ill[1] == $illstu &&
             $ilh[1] == $ilhstu 
             )  
             ||
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             $oll[0] == $ollstu &&
             $olh[0] == $olhstu &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu) 
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             $oll[1] == $ollstu &&
             $olh[1] == $olhstu &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu) 
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration.");
           return 0.94;
        } elsif (
             (
             $f      == $fstu &&
             $id[0]  == $idstu && 
             $od[0]  == $odstu &&
             ($oll[0] != $ollstu || $olh[0] != $olhstu) &&
             ($ill[0] != $illstu || $ilh[0] != $ilhstu)   
             )  
             ||
             (
             $f      == $fstu &&
             $id[1]  == $idstu && 
             $od[1]  == $odstu &&
             ($oll[1] != $ollstu || $olh[1] != $olhstu) &&
             ($ill[1] != $illstu || $ilh[1] != $ilhstu)   
             )  
           ) {
           $self->setMessage(1,"Check your limits of integration and order of integration.");
           return 0.47;
        } else {
             return 0;
        }
    }
  );



##############################################
#  Main text

Context()->texStrings;
Context()->flags->set(reduceConstants=>0);
BEGIN_TEXT
Set up a double integral in rectangular coordinates
for calculating the volume 
of the solid under the graph of the function 
\( f(x,y) = $f \) over the triangular region 
\( x + 2 y \leq $a, x \geq 0, \) and \(y \geq 0 \).
$BR
$BR
${BITALIC}Instructions:${EITALIC}
Please enter the integrand in the first answer box. 
Depending on the order of integration you choose,  
enter ${BITALIC}dx${EITALIC} and ${BITALIC}dy${EITALIC} 
in either order into the second and third answer boxes 
with only one ${BITALIC}dx${EITALIC} or 
${BITALIC}dy${EITALIC} in each box.  Then, enter the limits of 
integration and evaluate the integral to find the volume. 
$BR
$BR
\( \displaystyle \int_A^B \int_C^D \) 
\{ $multipart->ans_rule(40) \}
\{ $multipart->ans_rule(5) \}
\{ $multipart->ans_rule(5) \}
$BR
$BR
A = \{ $multipart->ans_rule(20) \} $BR
B = \{ $multipart->ans_rule(20) \} $BR
C = \{ $multipart->ans_rule(20) \} $BR
D = \{ $multipart->ans_rule(20) \}
$BR
$BR
Volume = \{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;





##################################################
#  Answers

ANS($multipart->cmp() );
ANS(($V)->cmp());

$showPartialCorrectAnswers = 1;

COMMENT('Allows integration in either order.');


;
ENDDOCUMENT();
