#DESCRIPTION
##  Algebra problem: combining powers
##ENDDESCRIPTION

## DBsubject(Arithmetic)
## DBchapter(Irrational numbers)
## DBsection(Simplify radical numbers)
## Date(6/3/2002)
## MLT(RatDenom_various)
## Level(2)
## MO(1)
## TitleText1('Precalculus')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.2')
## Problem1('57')
## TitleText2('College Algebra')
## AuthorText2('Stewart, Redlin, Watson')
## EditionText2('3')
## Section2('1.3')
## Problem2('57')
## KEYWORDS('algebra', 'powers')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
"PG.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGauxiliaryFunctions.pl",
"parserMultiAnswer.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = random(1,4);
$b = random(5,9);

$ans1 = $a*$b;
$ans2 = $b;

$mp = MultiAnswer(Real($ans1), Real($ans2))->with(
  checker=> sub {
    my ( $correct, $student, $self ) = @_;
	my ( $u, $v ) = @{$student};
	my ( $c1, $c2 ) = @{$correct};
	$u = Real($u);
	$v = Real($v);
	if( $u != int($u) || $u <= 0 ) { return [0,0];}
	if( $v != int($v) || $v <= 0 ) { return [0,0];}
	return [1,1] if $u*$c2*$c2 == $v*$v*$c1;
	return [0,0];
  }
);

BEGIN_TEXT
Rationalize the denominator of expression 
\[\sqrt{\frac{$a}{$b}},\] 
i.e., write it in the form
\[\frac{\sqrt{a}}{b}.\] 
$BR
Your answer for \(a\) is : \{$mp->ans_rule(5) \}
$BR
Your answer for \(b\) is : \{$mp->ans_rule(5) \}
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();        # This should be the last executable line in the problem.
