#DESCRIPTION
##  Algebra problem: combining powers
##ENDDESCRIPTION

## DBsubject(Algebra)
## DBchapter(Properties of exponents, rational exponents and radicals)
## DBsection(Properties of rational exponents and radicals)
## Date(6/3/2002)
## MLT(rationalize_basic)
## MO(1)
## Level(2)
## TitleText1('Precalculus')
## AuthorText1('Stewart, Redlin, Watson')
## EditionText1('3')
## Section1('1.2')
## Problem1('57')
## KEYWORDS('algebra', 'powers')
DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
"PG.pl",
"PGbasicmacros.pl",
"PGchoicemacros.pl",
"PGanswermacros.pl",
"PGauxiliaryFunctions.pl",
"parserMultiAnswer.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

$a = random(1,3,1);
$b = random(5,7,1);

$ans1 = $a*$b;
$ans2 = $b;

$mp = MultiAnswer(Real($ans1), Real($ans2))->with(
  checker=> sub {
    my ( $correct, $student, $self ) = @_;
	my ( $u, $v ) = @{$student};
	my ( $c1, $c2 ) = @{$correct};
	$u = Real($u);
	$v = Real($v);
	if( $u != int($u)) { return [0,0];}
	if( $v != int($v)) { return [0,0];}
	return [1,1] if $u*$c1*$c1 == $v*$v*$c2;
	return [0,0];
  }
);

BEGIN_TEXT
Rationalize the denominator of expression 
\[\sqrt{\frac{$a}{$b}},\] 
i.e., write it in the form
\[\frac{\sqrt{a}}{b}.\] 
$BR
Your answer for \(a\) is : \{$mp->ans_rule(5) \}
$BR
Your answer for \(b\) is : \{$mp->ans_rule(5) \}
END_TEXT

ANS($mp->cmp);

ENDDOCUMENT();        # This should be the last executable line in the problem.
