## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/30/06

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Complex eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(5)
## Static(1)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('6.5')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');

$A = Matrix([[0,-1],[1,0]]);

Context()->texStrings;
BEGIN_TEXT
Give an example of a \(2 \times 2 \) matrix with no real eigenvalues.
$BR
$BR
\{ $A->ans_array \}
END_TEXT

ANS($A->cmp(
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      $student = Matrix($student);
      my $a = $student->element(1,1);
      my $b = $student->element(1,2);
      my $c = $student->element(2,1);
      my $d = $student->element(2,2);
      
      return Real(($a + $d)**2 - 4 * ($a * $d - $b * $c)) < Real(0);
  }
));

;
ENDDOCUMENT();
