## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 4/30/06

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## MLT(comp_char_poly)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('6.1')
## Problem1('')

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix')->variables->are(x=>'Real');

$a11 = non_zero_random(-5,5,1);
$a12 = non_zero_random(-5,5,1);
$a13 = 0;
$a21 = 0;
$a22 = non_zero_random(-5,5,1);
$a23 = non_zero_random(-5,5,1);
$a31 = non_zero_random(-5,5,1);
$a32 = non_zero_random(-5,5,1);
$a33 = 0;

$A = Matrix([[$a11, $a12, $a13], [$a21, $a22, $a23], [$a31, $a32, $a33]]);

$k = $a11 + $a22;
$m = $a23 * $a32 - $a11 * $a22;
$n = $a12 * $a23 * $a31 - $a11 * $a23 * $a32;

$ans = Formula("x^3 - $k * x^2 - $m * x - $n")->reduce;

Context()->texStrings;
BEGIN_TEXT
Find the characteristic polynomial of the matrix 
\[ A = $A. \]
\( p(x) = \) \{ans_rule(40)\}.
END_TEXT
Context()->normalStrings;

ANS($ans->cmp( checker=>sub { 
		  my ( $c, $s, $a ) = @_;
		  return $c == $s || -1*$c == $s; } ) );

;
ENDDOCUMENT() ;        
