## DESCRIPTION
## ENDDESCRIPTION

## DBsubject('Precalculus')
## DBchapter('Linear Functions and Change')
## DBsection('Linear Functions')
## KEYWORDS()

## TitleText1('AUCI')
## EditionText1('1')
## AuthorText1('Petrillo')
## Section1('1.2')
## Problem1('1')
## Author('Darwyn Cook')
## Institution('Alfred University')



########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#
Context("Numeric");

$x0 = Compute(random(-4,4));
$x1 = Compute(random(1,5))+$x0;

$y0 = Compute(random(-9,9));
$y1 = Compute(non_zero_random(-9,9))+$y0;

$deltax = $x1-$x0;
$deltay = $y1-$y0;
$m = $deltay/$deltax;


$multians = MultiAnswer($y0,$m,$x0)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $ystu, $mstu,$xstu ) = @{$student};
      my ( $f1, $f2,$f3 ) = @{$correct};
      my $xcorrect = 0;
      my $ycorrect = 0;
      my $mcorrect = 0;
      if ($mstu == $m){
         $mcorrect = 1;
      }
      if ($ystu == $y0){
         $ycorrect = 1;
         if ($xstu == $x0){
            $xcorrect = 1;
         }
      }
      elsif ($ystu == $y1){
         $ycorrect = 1;
         if ($xstu == $x1){
            $xcorrect = 1;
         }
      }
      elsif (($xstu == $x0) || ($xstu == $x1)){
         $xcorrect = 1; 
      }
      return [$ycorrect,$mcorrect,$xcorrect];   
  }
);

Context()->texStrings;
BEGIN_TEXT
Suppose that the straight line y = f(x) passes through the two points P = ($x0,$y0) and Q = ($x1,$y1). Using one of these points, find a slope intercept form equation for y:
$BR
$BCENTER
y - \{$multians->ans_rule(3)\} = \{$multians->ans_rule(3)\}\(\big(\)x - \{$multians->ans_rule(3)\}\(\big)\)
$ECENTER
END_TEXT
Context()->normalStrings;

ANS($multians->cmp);

ENDDOCUMENT();

