# ************************************
#   Comments + Keyword Clarification
#
#    Anton, Calculus, 8.2.2
# ************************************
# DESCRIPTION
#   Solve the separable differential equation dy/dx = A(1 + y^2)x^n
#
# ENDDESCRIPTION
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Institution(John Wiley & Sons)
## MLT(jimf_the_number_slayer)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Anton')
## EditionText1('9')
## Section1('8.2')
## Problem1('2')
## KEYWORDS('Differantial Equations', 'Separation of Variables', 'General Solution, 'Specific Solution')
#***************************************************************
#   Initialization - Load standard PG and MathObjects macros   
#***************************************************************
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "littleneck.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);
#
rand_button();
#
# Set Up - Set up variable context and formulation of problem
#
$CT = Context("Numeric");
Parser::Numbers::NoDecimals;
#*********************************************************************************
#    dy/dx = A(1 + y^2)x^n
#    Separable ==>  dy/(1 + y^2) = Ax^n
#    Solution: arctan y = Ax^(n+1)/(n+1) + C  ==>  y = tan(Ax^(n+1)/(n+1) + C)
#*********************************************************************************
$Exp = random(2,9,1);
$A = non_zero_random(-6,6,2);
$CT->variables->add(C=>'Real',c=>'Real');
$Solution1 = Formula("tan($A x^($Exp+1)/($Exp+1) + C)")->reduce;
$Solution2 = Formula("tan($A x^($Exp+1)/($Exp+1) + c)")->reduce;
#*******************
#   Display stuff
#*******************
$Diffeq = "\frac{dy}{dx} = $A (1 + y^2) x^{$Exp}";
$Dform = "\frac{dy}{1 + y^2} = $A x^{$Exp} dx";
$N = $Exp + 1;
$Afterint = "\tan^{-1} y = \frac{$A x^{$N}}{$N} + C";
$CTF = Context("Fraction");
$CTF->variables->add(C=>'Real');
$frac = Compute("$A/($Exp+1)")->reduce;
$Solutiondisp = Formula("tan($frac x^($Exp+1) + C)")->reduce('(-x)+y'=>0);
#***********************
#   Problem Statement
#***********************
$CT = Context("Numeric");
TEXT(beginproblem());
#
# Strings in Tex form (formula)
#
$CT->texStrings;
BEGIN_TEXT
Using separation of variables, solve the differential equation,
\[ $Diffeq. \]
Use $BITALIC $BBOLD C $EBOLD $EITALIC to represent the arbitrary constant.
$PAR
\(\small{y}\) = \{ ans_rule(30) \}
$PAR
END_TEXT
#
# Expect strings in standard form
#
$CT->normalStrings;
# *******************************
#   Answer and Solution section
# *******************************
#
# Check student solution
#   : Accept upper or lower case C in solution
#
ANS( $Solution1->cmp( checker=>sub
                {
                   my ($correct, $student, $ansHash) = @_;
                   return ( ($student == $Solution1) || ($student == $Solution2) );
                }
              )
   );
#***************************
#  Explanation of solution  
#***************************
$CT->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
After separating and writing in differential form, the equation becomes,
\[ \small{$Dform}. \]
Integrating on both sides yields,
\[ \small{$Afterint}. \]
Therefore,
\[ \small{y = $Solutiondisp}. \]
$PAR
END_SOLUTION
$CT->normalStrings;
ENDDOCUMENT();
