# ************************************
#   Comments + Keyword Clarification
#
#    Anton, Calculus, 8.2.4
# ************************************
# DESCRIPTION
#   Solve the separable differential equation (A + x^n)dy/dx = x^(n-1)/y
#
# ENDDESCRIPTION
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Institution(John Wiley & Sons)
## MLT(jimf_the_number_slayer)
## MLTleader(1)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Anton')
## EditionText1('9')
## Section1('8.2')
## Problem1('4')
## KEYWORDS('Differantial Equations', 'Separation of Variables', 'General Solution, 'Specific Solution')
#***************************************************************
#   Initialization - Load standard PG and MathObjects macros   
#***************************************************************
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "littleneck.pl",
  "contextFraction.pl",
  "PGcourse.pl"
);
#
rand_button();
#
# Set Up - Set up variable context and formulation of problem
#
$CT = Context("Numeric");
Parser::Numbers::NoDecimals;
#*********************************************************************************
#    (A + x^n)dy/dx = x^(n-1)/y
#    Separable ==>  y dy/(1 + y^2) = x^(n-1)/(A + x^n) dx
#    Solution: y^2/2 = (1/n)ln(A + x^n) + C  ==>  y^2 = (2/n) ln(A + x^n) + C
#
#    Note: The exponent is kept even so that A + x^n is strictly positive
#*********************************************************************************
$Exp = random(2,12,2);
$N = $Exp - 1;
$A = random(1,10,1);
$CT->variables->add(C=>'Real',c=>'Real');
$Denom = $Exp/2;
$Solution1 = Formula("(1/$Denom) ln($A + x^$Exp) + C")->reduce;
$Solution2 = Formula("(2/$Denom) ln($A + x^$Exp) + C")->reduce;
#*******************
#   Display stuff
#*******************
$Diffeq = "($A + x^{$Exp})\frac{dy}{dx} = \frac{x^{$N}}{y}";
$Dform = "y dy = \frac{x^{$N}}{$A + x^{$Exp}} dx";
$Afterint = "\frac{y^2}{2} = \frac{1}{$Exp}\ln($A + x^{$Exp}) + C";
$CTF = Context("Fraction");
$CTF->variables->add(C=>'Real',c=>'Real');
$Solutiondisp = Formula("(1/$Denom) ln($A + x^$Exp) + C")->reduce;
#***********************
#   Problem Statement
#***********************
$CT = Context("Numeric");
TEXT(beginproblem());
#
# Strings in Tex form (formula)
#
$CT->texStrings;
BEGIN_TEXT
Using separation of variables, solve the differential equation,
\[ $Diffeq. \]
Use $BITALIC $BBOLD C $EBOLD $EITALIC to represent the arbitrary constant.
$PAR
\(\small{y^2}\) = \{ ans_rule(30) \}
$PAR
END_TEXT
#
# Expect strings in standard form
#
$CT->normalStrings;
# *******************************
#   Answer and Solution section
# *******************************
#
# Check student solution
#   : Accept upper or lower case C in solution
#
ANS( $Solution1->cmp( checker=>sub
                {
                   my ($correct, $student, $ansHash) = @_;
                   return ( ($student == $Solution1) || ($student == $Solution2) );
                }
              )
   );
#***************************
#  Explanation of solution  
#***************************
$CT->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR
After separating and writing in differential form, the equation becomes,
\[ \small{$Dform}. \]
Integrating on both sides yields,
\[ \small{$Afterint}. \]
Therefore,
\[ \small{y^2 = $Solutiondisp}. \]
$PAR
END_SOLUTION
$CT->normalStrings;
ENDDOCUMENT();
