# ************************************
#   Comments + Keyword Clarification
#
#    Anton, Calculus, 8.2.7
# ************************************
# DESCRIPTION
#   Solve the separable differential equation e^(-y) f(x) + y' g(x) = 0,
#   where f(x) and g(x) are specific trig functions.
#
# ENDDESCRIPTION
## DBsubject(Differential equations)
## DBchapter(First order differential equations)
## DBsection(Separable)
## Institution(John Wiley & Sons)
## Level(2)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Anton')
## EditionText1('9')
## Section1('8.2')
## Problem1('7')
## KEYWORDS('Differantial Equations', 'Separation of Variables', 'General Solution, 'Specific Solution')
#***************************************************************
#   Initialization - Load standard PG and MathObjects macros   
#***************************************************************
DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "littleneck.pl",
  "PGcourse.pl"
);
#--
rand_button();
#
# Set Up - Set up variable context and formulation of problem
#
Context("Numeric");
Context()->variables->add(y=>'Real', dy=>'Real', dx=>'Real', C=>'Real', c=>'Real');
#*********************************************************************************
#    e^(-y) f(x) + y' g(x) = 0,
#    Separable ==>  f(x)/g(x) dx = -e^y dy
#    f(x)       g(x)             Solution
#    ----       ----           -------------		
#    sinx       -cos^2(x)      ln|sec x + C|
#    cosx       sin^2(x)       ln|csc x + C|
#    secx       -cos x         ln|tan x + C|
#    sinx       1              ln|cos x + C|
#    cosx       -1             ln|sin x + C|
#*********************************************************************************
@Fs = ('sin x', 'cos x', 'sec x', 'sin x', 'cos x');
@Gs = ('-cos^2 x', 'sin^2 x', '-cos x', '1', '-1');
@Sol1 = ('ln|sec x + C|', 'ln|csc x + C|', 'ln|tan x + C|', 'ln|cos x + C|', 'ln|sin x + C|');
@Sol2 = ('ln|sec x + c|', 'ln|csc x + c|', 'ln|tan x + c|', 'ln|cos x + c|', 'ln|sin x + c|');
@SolTex = ('\ln|\sec x + C|', '\ln|\csc x + C|', '\ln|\tan x + C|', '\ln|\cos x + C|', '\ln|\sin x + C|');
@Dfsimp = ('\sec x \tan x', '-\csc x \cot x', '\sec^2 x', '-\sin x', '\cos x');
@Aint = ('\sec x', '\csc x', '\tan x', '\cos x', '\sin x');
$idx = random(0,$#Fs,1);
$Fx = $Fs[$idx];
$Gx = $Gs[$idx];
$Solution1 = Formula("$Sol1[$idx]");
$Solution2 = Formula("$Sol2[$idx]");
$Diffeq = Formula("e^(-y)*($Fx) + dy/dx*($Gx)")->reduce;
#*******************
#   Display stuff
#*******************
$Dform = "e^y dy = $Dfsimp[$idx]\;dx";
$Afterint = "e^y = $Aint[$idx] + C";
$Solutiondisp = $SolTex[$idx];
#***********************
#   Problem Statement
#***********************
TEXT(beginproblem());
#
# Strings in Tex form (formula)
#
Context()->texStrings;
BEGIN_TEXT
Using separation of variables, solve the differential equation,
\[ $Diffeq = 0. \]
Use $BBOLD C $EBOLD to represent the arbitrary constant.
$PAR
\(y\) = \{ ans_rule(30) \}
$PAR
END_TEXT
#
# Expect strings in standard form
#
Context()->normalStrings;
# *******************************
#   Answer and Solution section
# *******************************
#
# Check student solution
#   : Accept upper or lower case C in solution
#
ANS( $Solution1->cmp( checker=>sub
                {
                   my ($correct, $student, $ansHash) = @_;
                   return ( ($student == $Solution1) || ($student == $Solution2) );
                }
              )
   );
#***************************
#  Explanation of solution  
#***************************
Context()->texStrings;
BEGIN_SOLUTION
$PAR SOLUTION $PAR
After separating and writing in differential form, the equation becomes,
\[ $Dform. \]
Integrating on both sides yields,
\[ $Afterint. \]
Applying the natural logarithm to each side we get,
\[ y = $Solutiondisp. \]
$PAR
END_SOLUTION
Context()->normalStrings;
ENDDOCUMENT();
