## DESCRIPTION
## The geometry of vectors in R^n
## ENDDESCRIPTION


## DBsubject(Geometry)
## DBchapter(Vector geometry)
## DBsection(Vectors and vector arithmetic)
## Date(December 2015)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(2)
## MO(1)
## KEYWORDS('vectors','vector operations')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGgraphmacros.pl",
  "contextLimitedVector.pl",
  "AnswerFormatHelp.pl",
  "parserMultiAnswer.pl",
  "PGML.pl",
  "scaffold.pl",
  "PGcourse.pl"
);
TEXT(beginproblem());
$refreshCachedImages = 1;
$showPartialCorrectAnswers = 1;

Context("LimitedVector-coordinate")->variables->are(x=>'Real',y=>'Real');
Context()->flags->set( ijk=>0 );

#  construct graph canvases
$pixels = 300;
($xmin,$xmax) = (-7,7);
($ymin,$ymax) = (-7,7);

foreach my $i (0) {
    $gr[$i] = init_graph($xmin,$ymin,$xmax,$ymax, axes=>[0,0], size=>[$pixels,$pixels]);
    $gr[$i] -> lb('reset');

    #  manually draw the grid
    foreach my $j ($xmin..$xmax) {
      $gr[$i] -> moveTo($j,$ymin);
      $gr[$i] -> lineTo($j,$ymax,'gray',1);
    }
    foreach my $j ($ymin..$ymax) {
      $gr[$i] -> moveTo($xmin,$j);
      $gr[$i] -> lineTo($xmax,$j,'gray',1);
    }

    #  axes labels
    $gr[$i] -> lb(new Label ( $xmax - 0.5,0,'x','black','left','bottom'));
    $gr[$i] -> lb(new Label ( 0.25,$ymax - 0.5,'y','black','left','bottom'));

    # grid labels
    foreach my $x (($xmin+1)..-1,1..($xmax-1)) {
        $gr[$i] -> lb(new Label ( $x,$ymin+0.75,$x,'black','center','top'));
    }
    foreach my $y (($ymin+1)..-1,1..($ymax-1)) {
        $gr[$i] -> lb(new Label ( $xmin+0.75,$y,$y,'black','right','middle'));
    }

} # end constructing graph canvas

#  construct points A, B
($ax,$ay) = (non_zero_random(-5,5,1), random(2,5,1));
do { 
    ($bx,$by) = (non_zero_random(-5,5,1), random(-5,-2,1)); 
} until (
    $bx != $ax && 
    $ax * $by != $bx * $ay &&            # A, B, and origin not collinear
    abs(($ax - $bx) % 2) == 0 && abs(($ay - $by) % 2) == 1
);

$AB = Vector("<$bx - $ax, $by - $ay>");
$BA = Vector("<$ax - $bx, $ay - $by>");

foreach my $i (0) {

#  plot points A, B, C, D
$gr[$i]->stamps(closed_circle($ax, $ay, 'blue'));
$gr[$i]->lb(new Label($ax+0.25, $ay+0.25, 'A', 'blue', 'left', 'bottom'));

$gr[$i]->stamps(closed_circle($bx, $by, 'blue'));
$gr[$i]->lb(new Label($bx-0.25, $by+0.25, 'B', 'blue', 'right', 'bottom'));

}

$multians2 = MultiAnswer(Vector("<- $ax, - $ay>"), Vector("< $bx,  $by>"))->with(
  singleResult => 1,
  separator => ' + ',
  tex_separator => ' + ',
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $a, $b ) = @{$student};
      return 0 if (Vector(0,0) == $a || Vector(0,0) == $b);
      return 0 if ($a->isParallel($b) || $a->isParallel($b));
      return ($AB == ($a + $b));
  }
);

$multians3 = MultiAnswer( Vector("< -$ax,  -$ay>"), Vector("< $bx, 0>"), Vector("<0,  $by>"))->with(
  singleResult => 1,
  separator => ' + ',
  tex_separator => ' + ',
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $a, $b, $c ) = @{$student};
      return 0 if (Vector(0,0) == $a || Vector(0,0) == $b || Vector(0,0) == $c);
      return 0 if ($a->isParallel($b) || $a->isParallel($c) || $b->isParallel($c));
      return ($AB == ($a + $b + $c));
  }
);

############################

$multians4 = MultiAnswer(Vector("<$ax - $bx, 0>"), Vector("<0, $ay - $by>"))->with(
  singleResult => 1,
  separator => ' + ',
  tex_separator => ' + ',
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $a, $b ) = @{$student};
      return (
        (Vector("<$ax - $bx, 0>") == $a && Vector("<0, $ay - $by>") == $b) || 
        (Vector("<$ax - $bx, 0>") == $b && Vector("<0, $ay - $by>") == $a)
      );
  }
);

# [ 1 -1; 1 1] * [x1; x2] = [ ax-bx; ay-by ]
$x1 = 0.5*($ax - $bx) + 0.5*($ay - $by);
$x2 = -0.5*($ax - $bx) + 0.5*($ay - $by);

$multians5 = MultiAnswer(Vector("<$x1,$x1>"), Vector("<- $x2, $x2>"))->with(
  singleResult => 1,
  separator => ' + ',
  tex_separator => ' + ',
  checker => sub {
      my ( $correct, $student, $ansHash ) = @_;
      my ( $a, $b ) = @{$student};
      return (
        (Vector("<$x1,$x1>") == $a && Vector("<- $x2, $x2>") == $b) ||
        (Vector("<$x1,$x1>") == $b && Vector("<- $x2, $x2>") == $a)
      );
  }
);


Scaffold::Begin(
    can_open => "always",
    is_open  => "incorrect",
    hardcopy_is_open => "always",
);

BEGIN_PGML
The figure below shows two named points [`A`] and [`B`] in [`\mathbb{R}^2`].

>> [@ image(insertGraph($gr[0]), width=>$pixels, height=>$pixels, tex_size=>900).$PAR @]* <<
END_PGML

Section::Begin("Part 1");
BEGIN_PGML
a. Find the coordinate vector for [`\overrightarrow{AB}`].

    [`\overrightarrow{AB} = `] [_________]{$AB}

b. Write the vector [`\overrightarrow{AB}`] as a sum of two nonzero, nonparallel vectors.

    [`\overrightarrow{AB} = `] [_________]{$multians2} + [_________]{$multians2}

c. Write the vector [`\overrightarrow{AB}`] as a sum of three nonzero vectors none of which is parallel to the others.

    [`\overrightarrow{AB} = `] [_________]{$multians3} + [_________]{$multians3} + [_________]{$multians3}
END_PGML
Section::End();

Section::Begin("Part 2");
BEGIN_PGML
a. Write the vector [`\overrightarrow{BA}`] as a sum of two vectors, one parallel to [`\langle 1, 0 \rangle`] and the other parallel to [`\langle 0, -2 \rangle`].

    [`\overrightarrow{BA} = `] [_________]{$multians4} + [_________]{$multians4}

b. Write the vector [`\overrightarrow{BA}`] as a sum of two vectors, one parallel to [`\langle 1, 1 \rangle`] and the other parallel to [`\langle -1, 1 \rangle`].

    [`\overrightarrow{BA} = `] [_________]{$multians5} + [_________]{$multians5}
END_PGML
Section::End();

Scaffold::End();
ENDDOCUMENT();
