## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('linear algebra','matrix','multiplication')
## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter('Matrices')
## DBsection(Matrix algebra)
## Date('July 2013')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('3.2')
## Problem1('')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "AnswerFormatHelp.pl",
  "PGcourse.pl"
);
$showPartialCorrectAnswers = 0;
TEXT(beginproblem()); 

Context('Matrix');

$A = Matrix([[1,2,3],[0,4,5],[0,0,6]]);

Context()->texStrings;
BEGIN_TEXT
Give an example of a \( 3 \times 3 \) upper triangular matrix \( A \) that is not diagonal.
$BR
$BR
\( A = \)
\{ $A->ans_array(3,3,5) \}
END_TEXT
Context()->normalStrings;

install_problem_grader(~~&std_problem_grader);

ANS( $A->cmp(  checker => sub {
    my ( $correct, $student, $ansHash ) = @_;
    $S = Matrix($student);
    if ( 
      $S->element(1,2) == 0 and 
      $S->element(1,3) == 0 and 
      $S->element(2,3) == 0 and
      $S->element(2,1) == 0 and 
      $S->element(3,1) == 0 and 
      $S->element(3,2) == 0 ) 
    {
      return 0;
    } elsif (
      $S->element(2,1) == 0 and 
      $S->element(3,1) == 0 and 
      $S->element(3,2) == 0 ) 
    {
      return 1;
    } else {
      return 0;
    }
  }
) );

COMMENT('MathObject version.');

ENDDOCUMENT();
