# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Given a line's slope and a point on the line, find the line's equation in 
# both slope-intercept and point-slope form. Slope is negative. Y-intercept is # positive or negative.
#
# Last edited: Jordan 8/30/13; Yao 7/8/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Linear Equations')
## KEYWORDS('slope','linear','equation','point-slope','slope-intercept')
## DBCCSS('8.F.3','A-CED')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserAssignment.pl",
  "bizarroArithmetic.pl",
  "PGcourse.pl",
);

##############################################

Context("Numeric");
Context()->variables->are(x=>'Real', y=>'Real');
parser::Assignment->Allow;
Context()->flags->set(showExtraParens=>0);
Context()->noreduce('(-x)-y','(-x)+y','(-x)*y','x*(-y)','-n');

Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
);

$m=random(-5,-2,1);
$b=non_zero_random(-5,5,1);

$x1=non_zero_random(-5,5,1);
$y1=$m*$x1+$b;

$ansPS = ($y1!=0) ? Formula("y=$m(x-$x1)+$y1") : Formula("y=$m(x-$x1)");

$ansSI=Formula("y=$m*x+$b");


##############################################

TEXT(beginproblem());

BEGIN_PGML

A line's slope is [`[$m]`]. The line passes through the point [`([$x1],[$y1])`]. Find an equation for this line in both point-slope and slope-intercept form.

    An equation for this line in point-slope form is: [_______________].

    An equation for this line in slope-intercept form is: [_______________].

END_PGML

##############################################

Context()->flags->set(reduceConstants=>0);
Context()->flags->set(reduceConstantFormulas=>0);

ANS($ansPS -> cmp(
      checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is a line equation for the right line, but it is not in the expected point-slope form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
      return 1;
}));

ANS($ansSI -> cmp(
      checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is a line equation for the right line, but it is not in the expected slope-intercept form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
      return 1;
}));

$s1 = $m*$x1;

BEGIN_PGML_SOLUTION

It's easier to find the point-slope form first, since we have been given a _point_ on the line and the _slope_ of the line. The generic formula for the point-slope form is:

    [`y=m(x-x_{0})+y_0`]

where [`m`] is the slope and [` (x_{0},y_{0}) `] is a point on the line. We have been given that [`m=[$m]`] and that the line passes through the point [` ([$x1],[$y1]) `]. We substitute these numbers into the formula, and we have:

    [`
\begin{aligned}
   y &= m(x-x_{0})+y_0 \\
   y &= [$m](x-[$x1])+[$y1]
\end{aligned}
`]

So the line has an equation in point-slope form: [`y = [$m](x-[$x1])+[$y1] `].

----

The slope-intercept form of a line equation looks like [`y=mx+b`], where [`m`] is the slope and [`b`] is the [`y`]-intercept. We can take the point-slope form and solve for [`y`].

    [`
\begin{aligned}
   y &= [$m](x-[$x1])+[$y1]\\
   y &= [$m]x-[$m]([$x1])+[$y1]\\
   y &= [$m]x-[$s1]+[$y1]\\
   y &= [$m]x-[$s1-$y1]\\
\end{aligned}
`]

So the line's equation in slope-intercept form is [`y=[$m]x+[$b]`].



END_PGML_SOLUTION

ENDDOCUMENT();

