# WeBWorK problem written by Carl Yao
# Portland Community College
#
# Change a line's equation from point-slope form to slope-intercept form.
# Slope is a negative fraction.
#
# Last edited: Jordan 8/30/13; Yao 7/8/13
#
# ENDDESCRIPTION

## DBsubject('Algebra')
## DBchapter('Basic Algebra')
## DBsection('Linear Equations')
## KEYWORDS('slope','linear','equation','point-slope','slope-intercept','fraction')
## DBCCSS('8.F.3','A-CED')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')
## Author('Alex Jordan, Carl Yao, Chris Hughes')
## Institution('PCC')

##############################################

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGML.pl",
  "parserAssignment.pl",
  "bizarroArithmetic.pl",
  "contextFraction.pl",
  "PGcourse.pl",
);

##############################################

Context("Fraction");
Context()->variables->are(x=>'Real', y=>'Real');
parser::Assignment->Allow;
Context()->flags->set(showExtraParens=>0);
Context()->noreduce('(-x)-y','(-x)+y','(-x)*y','x*(-y)','-n');

Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
);

$num=random(1,9,1);
$den=random(2,9,1);
while (gcd($num,$den)!=1) {$den=random(2,9,1);}

$m= -Fraction($num,$den);
$b=non_zero_random(-5,5,1);

$x1=$den*non_zero_random(-3,3,1);
$y1=$m*$x1+$b;

$left1 = Formula("y")->reduce;
$right1 = Formula("$m*(x-$x1)+$y1")->reduce;

#for solution display
$s1 = $m*$x1;
$s2 = $s1+$y1;
$right2 = Formula("$m*x-$s1")->reduce;

Context("Numeric");
Context()->variables->are(x=>'Real', y=>'Real');
parser::Assignment->Allow;
Context()->flags->set(showExtraParens=>0);
Context()->noreduce('(-x)-y','(-x)+y','(-x)*y','x*(-y)','-n');

Context()->operators->set(
'+' => {class => 'bizarro::BOP::add', isCommand => 1},
'-' => {class => 'bizarro::BOP::subtract', isCommand => 1},
);

$ansSI=Formula("y=x+$b")->reduce;
$ansSI = $ansSI->substitute(x=>Formula("$m x"));


##############################################

TEXT(beginproblem());

BEGIN_PGML

Change this equation from point-slope form to slope-intercept form.

    [`` [$left1]=[$right1] ``]

    In slope-intercept form, this line's equation would be [_______________].

END_PGML

##############################################

ANS($ansSI -> cmp(
      checker=>sub{
      my ( $correct, $student, $ansHash ) = @_;
      return 0 if $ansHash->{isPreview} || $correct != $student;
      $student = $ansHash->{student_formula};
      $correct = $correct->{original_formula} if defined $correct->{original_formula};
      $student = Formula("$student"); $correct = Formula("$correct");
      return 0 unless ($correct == $student);
      Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
      delete $correct->{test_values};
      delete $student->{test_values};
      Value->Error("Your answer is a line equation for the right line, but it is not in the expected slope-intercept form") unless (($correct == $student) or ($student == $correct));
      Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
      return 1;
}));

$absY1 = abs($y1);
$y1Output = $y1>0 ? "+$y1" : "-$absY1";

BEGIN_PGML_SOLUTION

    [`
\begin{aligned}
   [$left1]     &= [$right1] \\
   [$left1]     &= [$m]x+([$m])([$x1*-1])+[$y1] \\
   [$left1]     &= [$right2]+[$y1] \\
   y           &= [$m]x+[$b]
\end{aligned}
`]

The line's equation in slope-intercept form is [``y=[$m]x+[$b]``].

END_PGML_SOLUTION

ENDDOCUMENT();





