# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 12/09/08

## DBsubject(Calculus - single variable)
## DBchapter(Applications of differentiation)
## DBsection(Linear approximation and differentials)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(LinearApprox_model)
## MLTleader(1)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## Section1('3.9')
## Section2(3.Review)
## Section3(3.Review)
## Problem1('21')
## Problem2(116)
## Problem3(115)
## KEYWORDS('calculus','derivative','linear approximation','tangent line','approximation')

## Textbook tags
## HHChapter('Short-Cuts to Differentiation')

DOCUMENT();
loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserPopUp.pl",
  "PGcourse.pl"
);
 
Context("Numeric");
Context()->variables->add( r=>'Real' );
Context()->constants->add( G=>6.67, M=>5.98 ); ## times 10^{-11} and 10^{24},
                                               ## respectively, in N m kg
# this means that student answers are
#    not reduced to the numerical value
#    specified in the Context
Context()->flags->set( formatStudentAnswer=>'parsed' );

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

## approximate heights of mountains in Colorado, in km
@mtHeights = ( [ "Mount Elbert", 4.40 ],
	       [ "Blanca Peak",  4.37 ],
	       [ "Uncompahgre Peak", 4.36 ],
	       [ "Grays Peak", 4.35 ],
	       [ "Mount Shavano", 4.34 ],
	       [ "Mount Belford", 4.33 ],
	       [ "Mount Bross", 4.32 ],
	       [ "Capitol Peak", 4.31 ],
	       [ "Pikes Peak", 4.30 ],
	       [ "Missouri Mountain", 4.29 ],
	       [ "Culebra Peak", 4.28 ],
	       [ "Pyramid Peak", 4.27 ] );

$whichMt = random(0, scalar(@mtHeights)-1, 1);
$posneg = PopUp( ["?", "positive", "negative"], "negative" );

Context()->texStrings;
BEGIN_TEXT

The acceleration due to gravity, \(g\), is given by
\[ g = \frac{G M}{r^2}, \]
where \(M\) is the mass of the Earth, \(r\) is the distance from
the center of the Earth, and \(G\) is the uniform gravitational
constant.
$PAR

${BBOLD}(a)$EBOLD
Suppose that we increase from our distance from the center of the
Earth by a distance \(\Delta r = x\).  Use a linear approximation
to find an approximation to the resulting change in \(g\), as a
fraction of the original acceleration:
$BR
\( \Delta g \approx \) \(g \times\) \{ ans_rule(35) \}
$BR
${BITALIC}(Your answer will be a function of \(x\) and \(r\).)$EITALIC

$PAR
${BBOLD}(b)$EBOLD
Is this change positive or negative?
$BR
\( \Delta g \) is \{ $posneg->menu() \}
$BR
${BITALIC}(Think about what this tells you about the acceleration
due to gravity.)$EITALIC

$PAR
${BBOLD}(c)$EBOLD
What is the percentage change in \(g\) when moving from sea level
to the top of \{ $mtHeights[$whichhMt]->[0] \} (a mountain over
14,000 feet tall in Colorado; in km, its height is
\{ $mtHeights[$whichMt]->[1] \} km; assume the radius of the Earth
is 6400 km)?
$BR
percent change = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS( Compute("-2*x/r")->cmp() );
ANS( $posneg->cmp() );

## accept either + or - percent
ANS( Compute("2*$mtHeights[$whichMt]->[1]/64")->cmp( checker=>sub {
	my ( $correct, $student, $ansHash ) = @_;
	return ($correct == $student || $correct == -1*$student);
    } ) );


$mt = $mtHeights[$whichMt]->[0];
$ht = $mtHeights[$whichMt]->[1];
$pchg = round(10000*2*$ht/64)/10000;

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

${BBOLD}(a)$EBOLD
The linear approximation for \(g\) around a given distance
\(r\) is, with \(g = f(r) = \frac{GM}{r^2}\),
\(g \approx f(r) + f'(r)\Delta r\).  Thus the change in
\(g\) is just
\(\Delta g = f'(r)\,\Delta r = -\frac{2GM}{r^3}\,\Delta r\).
Writing \(\Delta r = x\), this is
\(\Delta g = -\frac{2GM}{r^3}\,x = -2\frac{GM}{r^2}\,(\frac{x}{r}) =
  -2g\,(\frac{x}{r})\).

$PAR
${BBOLD}(b)$EBOLD
This is clearly negative, which implies that the gravitational
acceleration is smaller the higher we go!

$PAR
${BBOLD}(c)$EBOLD
If we go from sea level to the top of $mt,
the change in \(r\) is \(x = $ht\) km,
so the percent change in \(g\) is
\( p = -2\frac{$ht}{6400}\cdot 100 \approx -$pchg\).

END_SOLUTION
Context()->normalStrings;

;

ENDDOCUMENT();
