# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus','taylor series','error','polynomials')
## Tagged by glr 01/27/09

## DBsubject('Calculus')
## DBchapter('Infinite Sequences and Series')
## DBsection('Taylor and MacLaurin Series')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## AuthorText1('Hughes-Hallett')
## Section1('10.4')
## Problem1('1')

## Textbook tags
## HHChapter('Approximating Functions Using Series')
## HHSection('The Error in Taylor Polynomial Approximations')


DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
$showPartialCorrectAnswers = 1;

$a = random(0.1,0.5,0.1);
$n = random(2,4,1);
$nth = ( 'second', 'third', 'fourth' )[$n-2];
$np1 = $n+1;
$np1fac = Compute("$np1!");

$val = Compute("e^($a)*($a)^($np1)/$np1fac");

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Use the Error Bound for Taylor Polynomials to find a reasonable
(smallest) upper-bound for the error in approximating the quantity
\( e^{$a} \) with a ${nth}-degree Taylor polynomial about \(x=0\).

$PAR
error bound = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($val->cmp( checker=>sub{
		    my ( $correct, $student, $ansHash ) = @_;
		    if ( $correct == $student ) {
			return 1;
		    } elsif ( $student > $correct ) {
			Value->Error("Your answer isn't the smallest " .
				     "upper-bound you can obtain with the " .
				     "theorem.");
			return 0;
		    } else {
			return 0;
		    }
		} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Let \(f(x)=e^x\), so \(f($a)=e^{$a}\). The
error bound in the Taylor approximation of degree $n for
\(f($a)=e^{$a}\) about \(x=0\) is:
\[
|E_$n|= |f($a)-P_$n($a)| \leq \frac{M\cdot|$a-0|^$np1}{$np1!} =
\frac{M($a)^$np1}{$np1fac},\]
where \(|f^{($np1)}(x)| \le M\) for \(0\leq x \leq $a\).
Now, \(f^{($np1)}(x) = e^x\).  By looking at the
graph of \(e^x\), we see that \(|f^{($np1)}(x)|\) is
maximized
for \(x\) between 0 and $a when \(x=$a\). Thus,
\[
|f^{($np1)}| \leq e^{$a},
\]
so
\[
|E_$n|\leq \frac{e^{$a}\cdot ($a)^$np1}{$np1fac} \approx $val.
\]

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
