# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 01/27/09

## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Taylor polynomials)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(ErrorBound)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('10.4')
## Section2(10.4)
## Section3('10.4')
## Section4('10.4')
## Problem1('1')
## Problem2(1*)
## Problem3('1')
## Problem4('1')
## KEYWORDS('calculus','taylor series','error','polynomials')

## Textbook tags
## HHChapter('Approximating Functions Using Series')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
#  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
$showPartialCorrectAnswers = 1;

$a = random(0.1,0.5,0.1);
$n = random(2,4,1);
$nth = ( 'second', 'third', 'fourth' )[$n-2];
$np1 = $n+1;
$np1fac = Compute("$np1!");



$val = Compute("3*($a)^($np1)/$np1fac");

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Imagine that you need to compute \(e^{$a}\) but you have no calculator or other aid to enable you to compute it exactly, only paper and pencil.  You decide to use a ${nth}-degree Taylor polynomial expanded around \(x=0\).  Use the fact that \(e^{$a}<e<3\) and the Error Bound for Taylor Polynomials to find an upper bound for the error in your approximation.

$PAR
\(\lvert\)error\(\rvert \leq \)\{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($val->cmp( checker=>sub{
		    my ( $correct, $student, $ansHash ) = @_;
		    if ( $correct == $student ) {
			return 1;
		    } elsif ( $student > $correct ) {
			Value->Error("Your answer isn't the smallest " .
				     "upper-bound you can obtain with the " .
				     "theorem.");
			return 0;
		    } else {
			return 0;
		    }
		} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

Let \(f(x)=e^x\), so \(f($a)=e^{$a}\). The
error bound in the Taylor approximation \(P_{$n}($a)\) of degree $n for
\(f($a)\) about \(x=0\) is:
\[
|E_$n|= |f($a)-P_$n($a)| \leq \frac{M\cdot|$a-0|^$np1}{$np1!} =
\frac{M($a)^$np1}{$np1fac},\]
where \(|f^{($np1)}(x)| \le M\) for \(0\leq x \leq $a\).
Now, \(f^{($np1)}(x) = e^x\) is a positive, increasing function (since its derivative, \(e^x\), is never negative), so  \(|f^{($np1)}(x)|\) is
maximized
for \(x\) between 0 and $a when \(x=$a\). Thus,
\[
|f^{($np1)}| \leq e^{$a}<e < 3
\]
so
\[
|E_$n|\leq \frac{3\cdot ($a)^$np1}{$np1fac} \approx $val.
\]

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
