# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION

## KEYWORDS('calculus','taylor series','error','polynomials')
## Tagged by glr 01/27/09

## DBsubject('Calculus')
## DBchapter('Infinite Sequences and Series')
## DBsection('Taylor and MacLaurin Series')
## Date('')
## Author('Gavin LaRose')
## Institution('University of Michigan')
## TitleText1('Calculus')
## TitleText2('Calculus')
## EditionText1('4')
## EditionText2('5')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## Section1('10.4')
## Section2('10.4')
## Problem1('15')
## Problem2('15')

## Textbook tags
## HHChapter1('Approximating Functions Using Series')
## HHChapter2('Approximating Functions Using Series')
## HHSection1('The Error in Taylor Polynomial Approximations')
## HHSection2('The Error in Taylor Polynomial Approximations')



DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"PGchoicemacros.pl",
# "parserNumberWithUnits.pl",
# "parserFormulaWithUnits.pl", 
# "parserFormulaUpToConstant.pl",
# "PGcourse.pl",
);

Context("Numeric");
Context()->variables->add( n=>'Real' );
$showPartialCorrectAnswers = 1;

$whichF = random(0,1,1);
$f = ( $whichF ) ? Compute( "cos(x)" ) : Compute( "sin(x)" );
$fp = ( $whichF ) ? Compute( "sin(x)" ) : Compute( "cos(x)" );
$x0 = random(0.8, 1.5, 0.1);
$err = Compute( "($x0)^(n+1)/(n+1)!" );
$err->{test_points} = [[1,1],[2,2],[3,3],[4,4],[5,5]];

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

Use the Error Bound for Taylor Polynomials to give a good
bound for the error for the \(n^{\rm th}\) degree Taylor
polynomial about \(x = 0\) approximating \($f\) on the interval
\([0,$x0]\).
$PAR
error bound = \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($err->cmp( checker=>sub{
		    my ( $correct, $student, $ansHash ) = @_;
		    if ( $correct == $student ) {
			return 1;
		    } elsif ( $correct > $student ) {
			Value->Error("Your answer isn't the smallest bound " .
				     "you can obtain with the theorem.");
			return 0;
		    } else {
			return 0;
		    }
		} ) );

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

The maximum possible error for the \(n^{\rm th}\) degree Taylor polynomial
about \(x=0\) approximating \($f\) is
\(|E_n|\leq \frac{M\cdot|x-0|^{n+1}}{(n+1)!}\),
where the \(n^{\rm th}\) derivative of \($f\) is in magnitude
less than or equal to \(M\) on \(0 \leq x \leq 1\).
Now the derivatives of \($f\) are
simply \($f\), \($fp\), \(-$f\), and \(-$fp\).  The largest magnitude
these ever take is 1, so \(M = 1\), and thus
\( |E_n| \leq \frac{|x|^{n+1}}{(n+1)!}\leq \frac{$x0^{n+1}}{(n+1)!}\).

END_SOLUTION
Context()->normalStrings;


COMMENT('MathObject version');
ENDDOCUMENT();
