# DESCRIPTION
# Problem from Calculus, single variable, Hughes-Hallett et al., 4th ed.
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>

## Tagged by glr 01/27/09

## DBsubject(Calculus - single variable)
## DBchapter(Infinite sequences and series)
## DBsection(Taylor polynomials)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## MLT(FindTermForError)
## MLTleader(1)
## Level(3)
## MO(1)
## TitleText1('Calculus')
## TitleText2('Calculus')
## TitleText3('Calculus')
## TitleText4('Calculus')
## AuthorText1('Hughes-Hallett')
## AuthorText2('Hughes-Hallett')
## AuthorText3('Hughes-Hallett')
## AuthorText4('Hughes-Hallett')
## EditionText1('4')
## EditionText2('5')
## EditionText3('6')
## EditionText4('7')
## Section1('10.4')
## Section2('10.4')
## Section3('10.4')
## Section4('10.4')
## Problem1('15')
## Problem2('16')
## Problem3('19')
## Problem4('25')
## KEYWORDS('calculus','taylor series','error','polynomials')

## Textbook tags
## HHChapter1('Approximating Functions Using Series')
## HHChapter2('Approximating Functions Using Series')


DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "PGcourse.pl"
);

Context("Numeric");
Context()->flags->set(reduceConstants=>0);
$showPartialCorrectAnswers = 1;

$n1 = random(3,5,1);
$n2 = $n1 + 3;
$whichF = random(0,1,1);
$f = ( $whichF ) ? Compute( "cos(x)" ) : Compute( "sin(x)" );
$fp = ( $whichF ) ? Compute( "sin(x)" ) : Compute( "cos(x)" );
$f1 = ( $whichF ) ? "\cos(1)" : "\sin(1)";

# the required degrees are
%degrees = ( 3 => 6,   4 => 7,   5 => 8,
	     6 => 9,   7 => 10,  8 => 11 );
$deg1 = $degrees{$n1};
$deg2 = $degrees{$n2};

# correct for odd/even polynomials
if ( ( $whichF && ( $n1 == 4 || $n1 == 6 || $n1 == 11 ) ) ||
     ( ! $whichF && ( $n1 == 3 || $n1 == 5 || $n1 == 7 ) ) ) {
    $deg1a = $deg1 - 1;
    $deg2a = $deg2;
} else {
    $deg1a = $deg1;
    $deg2a = $deg2 - 1;
}

TEXT(beginproblem());
Context()->texStrings;
BEGIN_TEXT

What is the minimal degree Taylor polynomial about \( x= 0 \)
that you need to calculate \( $f1 \) to $n1 decimal places?
$BR
degree = \{ ans_rule(10) \}
$PAR
To $n2 decimal places?
$BR
degree = \{ ans_rule(10) \}

END_TEXT
Context()->normalStrings;

ANS(Compute( $deg1 )->cmp( checker => sub {
			       my ( $c, $s, $ansh ) = @_;
			       return $c == $s || Compute($deg1a) == $s; } ) );
ANS(Compute( $deg2 )->cmp( checker => sub {
			       my ( $c, $s, $ansh ) = @_;
			       return $c == $s || Compute($deg2a) == $s; } ) );

$fn1 = Compute("0.5*10^(-$n1)");
$fn2 = Compute("0.5*10^(-$n2)");

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
$PAR SOLUTION $PAR

By using the Error Bound for Taylor Polynomials, if we approximate
\( $f1 \) using the \(n^{\rm th}\) degree
polynomial, the error is at most \(\frac{1}{(n+1)!}\). For the answer to
be correct to four decimal places, the error must be less than
\( $fn1 \).  Thus, the first \(n\) such that
\(\frac{1}{(n+1)!} < $fn1\) will work. In particular, this is first
true when \(n=$deg1\).

$PAR
For $n2 decimal places, we need
\(\frac{1}{(n+1)!} < $fn2\), for which \(n=$deg2\) works.

END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
