## DESCRIPTION
## Jacobian matrix
## ENDDESCRIPTION

## KEYWORDS('Jacobian','Jacobian matrix')

## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('08/24/2012')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('Multivariable Mathematics')
## EditionText1('Version 3.0')
## AuthorText1('Darin Stephenson')
## Section1('10.2')
## Problem1('3a')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


#############################
#  Setup

Context("Numeric")->variables->are(
x=>'Real',
y=>'Real',
z=>'Real',
);

# construct the function
$a = random(2,3,1);
$b = random(4,6,1);
do { $c = random(2,3,1); } until ($c != $a);
$f = Formula("x^$a + $b x y^$c z + y^$c + z");

$fx = $f->D('x');
$fy = $f->D('y');
$fz = $f->D('z');

# coordinates of the point
$Px = non_zero_random(-2,2,1);
do { $Py = non_zero_random(-2,2,1); } until ($Py != $Px);
do { $Pz = non_zero_random(-2,2,1); } until ($Pz != $Px && $Pz != $Py);

# entries of Jacobian at the point
$fPx = $fx->eval(x=>$Px,y=>$Py,z=>$Pz);
$fPy = $fy->eval(x=>$Px,y=>$Py,z=>$Pz);
$fPz = $fz->eval(x=>$Px,y=>$Py,z=>$Pz);


$multians1 = MultiAnswer($fx, $fy, $fz)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rrr} %s & %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..2) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians2 = MultiAnswer($fPx, $fPy, $fPz)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rrr} %s & %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..2) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Consider the function
$PAR
\[
f(x,y,z) = $f.
\]
$BR
(a) Find a general expression for the Jacobian matrix \( \displaystyle Jf = \frac{\partial(f)}{\partial(x,y,z)}. \)
$BR
$BR
\{ 
mbox(
  "\( Jf = \)",
  display_matrix([ [$multians1->ans_rule(20),$multians1->ans_rule(20),$multians1->ans_rule(20)] ], align=>'ccc')
);
\}
$BR
(b) Evaluate \( Jf \) at the point \( \boldsymbol{\vec{a}} = \left( \begin{array}{r} $Px \\ $Py \\ $Pz \end{array} \right). \)
$BR
$BR
\{ 
mbox(
  "\( Jf(\boldsymbol{\vec{a}}) = \)",
  display_matrix([ [$multians2->ans_rule(20),$multians2->ans_rule(20),$multians2->ans_rule(20)] ], align=>'ccc')
);
\}
END_TEXT
Context()->normalStrings;

########################
#  Answer evaluation

#install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );


COMMENT("MathObject version.  (Uses the same randomization as 10-02-Jacobian-06a.pg)");

ENDDOCUMENT();


