## DESCRIPTION
## Jacobian matrix
## ENDDESCRIPTION

## KEYWORDS('Jacobian','Jacobian matrix')

## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('08/24/2012')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('Multivariable Mathematics')
## EditionText1('Version 3.0')
## AuthorText1('Darin Stephenson')
## Section1('10.2')
## Problem1('3c')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


#############################
#  Setup

Context("Numeric")->variables->are(
x=>'Real',
y=>'Real',
z=>'Real',
);

# construct the function
$a = random(2,7,1);
do { $b = random(2,7,1); } until ($b != $a);
do { $c = random(2,7,1); } until ($c != $a && $c != $b);

$f1 = Formula("$a e^(x y z^2) + $b y");
$f2 = Formula("$c x^2 + x z");

# construct the Jacobian
$f1x = $f1->D('x');   $f1y = $f1->D('y');   $f1z = $f1->D('z');
$f2x = $f2->D('x');   $f2y = $f2->D('y');   $f2z = $f2->D('z');

# coordinates of the point
$Px = non_zero_random(-1,1,1);
$Py = 0;
$Pz = -1 * $Px;

# entries of Jacobian at the point
$f1Px = $f1x->eval(x=>$Px,y=>$Py,z=>$Pz);   $f1Py = $f1y->eval(x=>$Px,y=>$Py,z=>$Pz);   $f1Pz = $f1z->eval(x=>$Px,y=>$Py,z=>$Pz);
$f2Px = $f2x->eval(x=>$Px,y=>$Py,z=>$Pz);   $f2Py = $f2y->eval(x=>$Px,y=>$Py,z=>$Pz);   $f2Pz = $f2z->eval(x=>$Px,y=>$Py,z=>$Pz);


$multians1 = MultiAnswer($f1x, $f1y, $f1z, $f2x, $f2y, $f2z)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rrr} %s & %s & %s \\  %s & %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..5) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians2 = MultiAnswer($f1Px, $f1Py, $f1Pz, $f2Px, $f2Py, $f2Pz)->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td><td> %s </td></tr>
  </table>",
  tex_format =>  "\left\lbrack\begin{array}{rrr} %s & %s & %s \\ %s & %s & %s \end{array}\right\rbrack",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..5) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Consider the function
$PAR
\[
f(x,y,z) = \left( \begin{array}{c} $f1 \\ $f2 \end{array} \right).
\]
$BR
(a) Find a general expression for the Jacobian matrix \(\displaystyle Jf = \frac{\partial(f_1,f_2)}{\partial(x,y,z)}.\)
$BR
$BR
\{ 
mbox(
  "\( Jf = \)",
  display_matrix([ [$multians1->ans_rule(20),$multians1->ans_rule(20),$multians1->ans_rule(20)], [$multians1->ans_rule(20),$multians1->ans_rule(20),$multians1->ans_rule(20)] ], align=>'ccc')
);
\}
$BR
(b) Evaluate \( Jf \) at the point \( \boldsymbol{\vec{a}} = \langle $Px, $Py, $Pz \rangle. \)
$BR
$BR
\{ 
mbox(
  "\( Jf(\boldsymbol{\vec{a}}) = \)",
  display_matrix([ [$multians2->ans_rule(20),$multians2->ans_rule(20),$multians2->ans_rule(20)], [$multians2->ans_rule(20),$multians2->ans_rule(20),$multians2->ans_rule(20)] ], align=>'ccc')
);
\}
END_TEXT
Context()->normalStrings;

########################
#  Answer evaluation

#install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 1;

ANS( $multians1->cmp() );
ANS( $multians2->cmp() );


COMMENT("MathObject version.  (Uses the same randomization as 10-02-Jacobian-06c.pg)");

ENDDOCUMENT();


