## DESCRIPTION
## Jacobian matrix
## ENDDESCRIPTION

## KEYWORDS('Jacobian','Jacobian matrix')

## DBsubject('Calculus')
## DBchapter('Partial Derivatives')
## DBsection('Functions of Several Variables')
## Date('08/24/2012')
## Author('Paul Pearson')
## Institution('Hope College')
## TitleText1('Multivariable Mathematics')
## EditionText1('Version 3.0')
## AuthorText1('Darin Stephenson')
## Section1('10.2')
## Problem1('')


##############################
#  Initialization

DOCUMENT(); 

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserImplicitPlane.pl",
  "parserVectorUtils.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());


#############################
#  Setup

Context("ImplicitPlane")->variables->are(
x=>'Real',
y=>'Real',
z=>'Real',
);

# construct the function
$a = random(2,3,1);
$b = random(4,6,1);
do { $c = random(2,3,1); } until ($c != $a);
$f = Formula("x^$a + $b x y^$c");
$g = Formula("z - ($f)");

$gx = $g->D('x');
$gy = $g->D('y');
$gz = $g->D('z');

$gradg = Vector("<$gx, $gy, $gz>");

$x = non_zero_random(-5,5,1);
do { $y = non_zero_random(-5,5,1); } until ($y != $x);
$z = $f->eval(x=>$x,y=>$y);

$P = Point($x,$y,$z);

$n = $gradg->eval(x=>$x,y=>$y);

$plane = ImplicitPlane($P,$n);


######################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Consider the function
$PAR
\[
z = f(x,y) = $f.
\]
$BR
(a) Find \( f($x,$y) \).
$BR
\( z = f($x,$y) = \)
\{ ans_rule(30) \}
$BR
$BR
(b) Find a function \( g(x,y,z) \) whose level zero set is
equal to the graph of \( z = f(x,y) \) and such that the 
coefficient of \( z \) in \( g(x,y,z) \) is \( 1 \).
$BR
The level set 
\( g(x,y,z) = \) 
\{ ans_rule(40) \} 
\( = 0 \) is the same as the graph of \( z = f(x,y) \).
$BR
$BR
(c) Find the gradient of \( g \).  Write your answer as a 
row vector of the general form \( \langle a, b, c \rangle \).
$BR
\( \nabla g(x,y,z) = \) 
\{ ans_rule(40) \}
$BR
$BR
(d) Use \( \nabla g \) to find a vector \( \vec{n} \) 
perpendicular (or normal) to the graph of \( z = f(x,y) \) 
at the point \( $P \).  Write your answer as a 
row vector of the general form \( \langle a, b, c \rangle \).
$BR
\( \vec{n} = \)
\{ ans_rule(40) \}
$BR
$BR
(e) Find an equation for the tangent plane to \( z = f(x,y) \) 
at the point \( $P \).  Enter your answer as an equation.
$BR
\{ ans_rule(40) \}
END_TEXT
Context()->normalStrings;

########################
#  Answer evaluation

#install_problem_grader(~~&std_problem_grader);

$showPartialCorrectAnswers = 1;

ANS( $z->cmp() );
ANS( $g->cmp() );
ANS( $gradg->cmp() );
ANS( $n->cmp(checker=>sub {
  my ($correct, $student, $ansHash) = @_;
  return $correct->isParallel($student);
}, showCoordinateHints => 1) );
ANS( $plane->cmp() );


COMMENT("MathObject version");

ENDDOCUMENT();


