##DESCRIPTION
##ENDDESCRIPTION
## DBsubject(Statistics)
## DBchapter(Confidence intervals)
## DBsection(One sample proportion)
## Level(3)
## MO(1)
## KEYWORDS('sample proportion', 'confidence interval')
##Tagged by Nelson Chen. Jun 26th, 2016. 
##Attempts: Three. 


########################################################################

DOCUMENT();      

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGchoicemacros.pl",
  "parserMultiAnswer.pl",
  "PGnumericalmacros.pl",
  "PGstatisticsmacros.pl",
  "RserveClient.pl",
  "answerHints.pl",
  "PGcourse.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());


##############################################################
#
#  Setup
#
#
Context("Numeric");

@nR = rserve_eval('
nR <- round(runif(1, min=330, max=340))
');

@mR = rserve_eval('
mR <- round(runif(1, min=35, max=45))
');

@aR = rserve_eval('
aR <- mR/nR
round(aR, 3)
');

$num = random(30,40,1);
$n = random(330, 340, 1);

$mc1 = new_multiple_choice();
$mc1->qa(
"The proportion of the @nR[0] people who have recently had the flu, @mR[0]/@nR[0], is a: ", 
"statistic."
);
$mc1->extra(
"parameter.",
"variable of interest.",
);


# Part ii)
# Acceptable answers use z = 2 or z = 1.96

$phat = @mR[0] / @nR[0];

$left_endpoint = $phat - 2*sqrt($phat * (1-$phat) / @nR[0]);
$right_endpoint = $phat + 2*sqrt($phat * (1-$phat) / @nR[0]);

$left_endpoint = sprintf "%.3f", $left_endpoint;
$right_endpoint = sprintf "%.3f", $right_endpoint;

$left_ans1 = $phat - 2*sqrt($phat * (1-$phat) / @nR[0]);
$right_ans1 = $phat + 2*sqrt($phat * (1-$phat) / @nR[0]);

$left_ans1 = sprintf "%.3f", $left_ans1 ;
$right_ans1 = sprintf "%.3f", $right_ans1 ;

$left_ans2 = $phat - 1.96*sqrt($phat * (1-$phat) / @nR[0]);
$right_ans2 = $phat + 1.96*sqrt($phat * (1-$phat) / @nR[0]);

$left_ans2 = sprintf "%.3f", $left_ans2 ;
$right_ans2  = sprintf "%.3f", $right_ans2 ;

$left_ans1_lower=$left_ans1 - 0.001;
$left_ans1_upper=$left_ans1 + 0.001;
$right_ans1_lower=$right_ans1 - 0.001;
$right_ans1_upper=$right_ans1 + 0.001;

$left_ans2_lower=$left_ans2 - 0.001;
$left_ans2_upper=$left_ans2 + 0.001;
$right_ans2_lower=$right_ans2 - 0.001;
$right_ans2_upper=$right_ans2 + 0.001;

$multians = MultiAnswer($left_endpoint, $right_endpoint)->with(
  singleResult => 1,
  checker => sub {
      my ( $correct, $student, $anshash, $self ) = @_;
      my ( $stu_ans_left, $stu_ans_right ) = @{$student};
      my ( $ans1, $ans2 ) = @{$correct};
      $ansHash->{correct_ans} = "($left_ans1, $right_ans1) OR ($left_ans2, $right_ans2)";
     if((($stu_ans_left >= $left_ans1_lower && $stu_ans_left <= $left_ans1_upper) &&
($stu_ans_right >= $right_ans1_lower && $stu_ans_right <= $right_ans1_upper)) ||
(($stu_ans_left >= $left_ans2_lower && $stu_ans_left <= $left_ans2_upper) &&
($stu_ans_right >= $right_ans2_lower && $stu_ans_right <= $right_ans2_upper)))
      {
          return [1,1];
      } 
      else {
          return [0,0];
      }
  }
);

##############################################################
#
#  Text
#
#

Context()->texStrings;
BEGIN_TEXT

An epidemiologist is worried about the prevalence of the flu in East Vancouver and the potential shortage of vaccines for the area. She will need to provide a recommendation for how to allocate the vaccines appropriately across the city. She takes a simple random sample of @nR[0] people living in East Vancouver and finds that @mR[0] have recently had the flu.

$PAR
$BBOLD Part i) $EBOLD
\{ $mc1->print_q() \}
$BR
\{ $mc1->print_a() \}


$PAR
$BBOLD Part ii) $EBOLD
Use the sample data to compute a 95$PERCENT confidence interval for the true proportion of East Vancouver residents that have recently had the flu.

$PAR
(Please carry answers to at least six decimal places in intermediate steps. Give your final answer to the nearest three decimal places).

$PAR
95$PERCENT confidence interval = (\{ $multians->ans_rule(6) \}, \{ $multians->ans_rule(6) \})

END_TEXT
Context()->normalStrings;

##############################################################
#
#  Answers
#
#

$showPartialCorrectAnswers = 1;

ANS( radio_cmp( $mc1->correct_ans() ) );
ANS( $multians->cmp( tol=> 0.002)); 

## Solution (Optional but recommended)
Context()->texStrings;
BEGIN_SOLUTION
(i) The sample proportion is a function of the data, and so
is a statistic. (ii) The sample proportion, @aR[0], is an
estimate of the population proportion, \(p\) say, who have recently had
the flu. Since the sample proportion is approximately Normal with mean 
\(p\) and standard deviation \( \left( 1-p\right) p/@nR,\) we
can approximate a 95% confidence interval for \(p\) via
\[
@mR[0]/@nR[0] \pm 1.96 \sqrt{\frac{\left( @mR[0]/@nR[0] \right) \left( 1-@mR[0]/@nR[0]\right) }{@nR[0]}}.
\]

END_SOLUTION



ENDDOCUMENT();        
