## DESCRIPTION
## KEYWORDS('vectors', 'parametric', 'equation', 'line')
##
## ENDDESCRIPTION

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros( "PG.pl", 
"PGbasicmacros.pl", 
"PGanswermacros.pl", 
"Parser.pl", 
"parserVectorUtils.pl", );

TEXT(beginproblem());

Context("Vector"); 
## $t = Formula("t");

$p = non_zero_point2D(-6,6,1); 
$q = non_zero_point2D(-5,5,1);
while($q==$p){
$q = non_zero_point2D(-5,5,1);
}
$p1=Vector($p)->with(ColumnVector=>1);
$v= Vector($q - $p)->with(ColumnVector=>1);

Context()->texStrings; 

BEGIN_TEXT

Enter a the vector equation for the line \(l\) going through the two points
\($p, $q\).
$PAR

\( l: \) \{$p1->ans_array\} \( + t \) \{$v->ans_array\} 

$BR
$BR
END_TEXT 

Context()->normalStrings;

ANS($p1->cmp(showCoordinateHints=>0, checker=> sub { my $correct = shift; my $student = shift; return $student == $p1 || $v->isParallel($student-$p1); }));

ANS($v->cmp(showCoordinateHints=>0, checker=> sub { my $correct = shift; my $student = shift; return $v->isParallel($student); }));

ENDDOCUMENT(); 
