##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('trigonometry', 'trig functions', 'trigonometric functions','fraction', 'analytic trigonometry', 'angle', 'angles', 'radians', 'magnitude', 'direction', 'vectors')

## DBsubject('Trigonometry')
## DBchapter('Polar coordinates & vectors')
## DBsection('Polar and rectangular coordinates')
## Date('7/31/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "AppletObjects.pl",
   "PGML.pl",
   "PGgraphmacros.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Vector");
Context()->flags->set(tolType=>'absolute', tolerance=>0.0005);

##############################################
#  Begin trig functions in degrees

Context()->functions->remove("cos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cos {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub cos {Parser::Function->call('cos',@_)} 
#  Add the new functions to the Context
Context()->functions->add( cos => {class => 'NewFunc', TeX => '\cos'}, );


Context()->functions->remove("sin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sin {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub sin {Parser::Function->call('sin',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sin => {class => 'NewFunc', TeX => '\sin'}, );


Context()->functions->remove("tan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub tan {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180)/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub tan {Parser::Function->call('tan',@_)}
#  Add the new functions to the Context
Context()->functions->add( tan => {class => 'NewFunc', TeX => '\tan'}, );


Context()->functions->remove("sec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sec {
  shift; my $x = shift;
  return 1.0/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub sec {Parser::Function->call('sec',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sec => {class => 'NewFunc', TeX => '\sec'}, );


Context()->functions->remove("csc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub csc {
  shift; my $x = shift;
  return 1.0/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub csc {Parser::Function->call('csc',@_)} 
#  Add the new functions to the Context
Context()->functions->add( csc => {class => 'NewFunc', TeX => '\csc'}, );


Context()->functions->remove("cot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cot {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180)/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub cot {Parser::Function->call('cot',@_)}
#  Add the new functions to the Context
Context()->functions->add( cot => {class => 'NewFunc', TeX => '\cot'}, );



#sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))}
#sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])}
#sub atan {CORE::atan2($_[1],1)}
#sub acot {CORE::atan2(1,$_[1])}
#sub asec {acos($_[0],1.0/$_[1])}
#sub acsc {asin($_[0],1.0/$_[1])}


Context()->functions->remove("acos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acos {Parser::Function->call('acos',@_)}
#  Add the new functions to the Context
Context()->functions->add( acos => {class => 'NewFunc', TeX => '\arccos'}, );


Context()->functions->remove("asin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub asin {Parser::Function->call('asin',@_)}
#  Add the new functions to the Context
Context()->functions->add( asin => {class => 'NewFunc', TeX => '\arcsin'}, );


Context()->functions->remove("atan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub atan {CORE::atan2($_[1],1)*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub atan {Parser::Function->call('atan',@_)}
#  Add the new functions to the Context
Context()->functions->add( atan => {class => 'NewFunc', TeX => '\arctan'}, );


Context()->functions->remove("asec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asec {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
#sub asec {acos($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub asec {Parser::Function->call('asec',@_)}
#  Add the new functions to the Context
Context()->functions->add( asec => {class => 'NewFunc', TeX => '\arcsec'}, );


Context()->functions->remove("acsc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acsc {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
#sub acsc {asin($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub acsc {Parser::Function->call('acsc',@_)}
#  Add the new functions to the Context
Context()->functions->add( acsc => {class => 'NewFunc', TeX => '\arccsc'}, );


Context()->functions->remove("acot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acot {CORE::atan2(1,$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acot {Parser::Function->call('acot',@_)}
#  Add the new functions to the Context
Context()->functions->add( acot => {class => 'NewFunc', TeX => '\arccot'}, );


#  End trig functions in degrees
###################################################

###################################################
#  Make a subclass of MultiAnswer
#

package myMultiAnswer;
our @ISA = ('parser::MultiAnswer');

sub new {
  my $self = shift;
  my $ma = $self->SUPER::new(@_);
  $ma->{part} = 1;
  $ma->{answerName} = 'answerBox';
  $ma->{answerNames}{0} = $ma->{answerName};
  $ma->{id} = $MultiAnswer::answerPrefix.$ma->{answerName};
  $ma->{singleResult} = 1;
  $ma->{namedRules}  = 1;
  return $ma;
}

sub ANS_NAME {
  my $self = shift;
  my $i = shift;
  $self->{answerNames}{$i} = ($i == 0 ? $self->{answerName} : $self->{id}.'_'.$i);
  return $self->{answerNames}{$i};
}

package main;

#  End subclass of MultiAnswer
###################################################

$answer = Vector("<2,2>");

$multiAns = myMultiAnswer->new(List(1,2), $answer)->with(
  format => '%2$s',
  tex_format => '%2$s',
  checkTypes=>0,
  allowBlankAnswers=>1,
  checker => sub {
      my ( $correct, $student, $self, $ans_hash ) = @_;
      my ( $ggbStu, $vecStu ) = @{$student};
      my ( $r, $theta ) = split(/,/,$ggbStu);
#      $self->{ans}[0]->{preview_text_string} = "hidden";
#      $self->{ans}[0]->{preview_latex_string} = "hidden";
      return [0,0] if $self->{ans}[0]->{isPreview};
      Value::Error("Very clever, but you must place \(\vec{v}\) in a quadrant to receive credit.") if ( abs($vecStu.Vector(1,0)) < 0.005 || abs($vecStu.Vector(0,1)) < 0.005 );
      Value::Error("The vector you entered is not a unit vector.") if ($vecStu->abs != 1);
      my $vecGGB = Vector("<cos($theta), sin($theta)>");
      Value->Error("Your vector doesn't match the image.") if $vecGGB != $vecStu;
      return [1,1];
  }
);

######################################
# Create  link to applet:
###################################
# You can name your applet (anything reasonable :-) )
# Adjust the height and width as desired
# Paste the geogebra parameters in at the bottom of the page just above the
# command end command
# so that they don't get in the way
###################################


$appletName = "ggbfoo";

$applet =  GeogebraWebApplet(
   code       => "geogebra.GeoGebraApplet",
   archive    => "",
   codebase                 => "",
   appletName              => $appletName,
   appletId                    => $appletName,
   submitActionAlias     => 'getAppletValues',  # default actionAlias
   initializeActionAlias   => 'configure',  # default actionAlias
   setStateAlias             => 'setXML',
   getStateAlias             => 'getXML',
   setConfigAlias           => '',
   getConfigAlias           => '',
   returnFieldName       => '',
   width                 => 613,              # may want to modify width
   height                => 525,              # may want to modify height
  # mayscript             => "true",
   debugMode             => 0,   # set debugMode to 0 for no debug
                                              # to 1 to make xml representation visible
                                              # to 2 to add alerts detailing progression
                                              # through the applet
   onInit                => 'ggbOnInit',
   type                  => 'geogebraweb',
   submitActionScript    => qq{ getQE('answerBox').value = getAppletValues() },
   #submitActionScript => '',
   selfLoading           => 1,
   params     => GEOGEBRA_PARAMS(),# paste parameters in

);


###############################################
#
#  PNG picture of GeoGebra applet for hardcopy

 $im = init_graph(-4,-4,4,4,axes=>[0,0],grid=>[8,8],size=>[400,400]);
 $im->moveTo(0,0);
 $im->arrowTo(2,2,'blue',2);


##################################
# Setup GeogebraWebApplet --
###################################

###############################################
#
#  For persistence of GeoGebra applet state 
#  when answers are submitted or page is refreshed.
#  Uses JavaScript interface to GeoGebra.

HEADER_TEXT(
qq! 
<script language="javascript">
   function ggbOnInit(param) {
       if (param == "$appletName") {
           applet_loaded(param,1);  // report that applet is ready. 
           ww_applet_list[param].safe_applet_initialize(2);
       }
   } 
</script>
<noscript>For this question to work properly, 
you must enable JavaScript, or use another browser 
(or computer) on which JavaScript works.</noscript>
!);


#######################################################
#
#  For setting initial values in GeoGebra from WeBWorK
#  Uses JavaScript interface to GeoGebra.

TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );
<script language="javascript">
function setAppletCoefficients() {
   var applet=getApplet("$appletName");   // alert("Updating coefficients");
   if (typeof(applet)!="undefined") {
     if (typeof(applet.setValue)!="undefined") {
     
     ////////////////////////////////////////////////////////////////
     //  List the values of the parameters to be set in the applet here
     
          // Initialize anything?
          applet.setFixed('text1',true,false);
          applet.setRounding(3);
     } else {
          setTimeout("setAppletCoefficients()", 1000);
     }
   } else {
      setTimeout("setAppletCoefficients()",1000);
   }
}
ww_applet_list["$appletName"].setConfig = function() {setAppletCoefficients()};
</script>
END_SCRIPT

HEADER_TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );

<script language="javascript">
function getAppletValues() {
    console.log("getAppletValues() entered");
	// get a handle for the applet object
	var applet = $appletName;
	var mag = applet.getValue("mag");
	var theta = applet.getValue("theta");

    return( mag + ',' + theta );	
}
 </script>
END_SCRIPT

@completed = split(/:/,$main::inputs_ref->{_data});

###################################
#insert applet into body  -- this does not need to be changed
# except to adjust the insertion of the reinitialize_button or
# a hidden AnswerBox
###################################
 TEXT($PAR,  MODES(TeX=>$im,
                      HTML=>$applet->insertAll(
                                debug =>0,
                                reinitialize_button => 0,
                                includeAnswerBox=>1,
 )));

##############################################################
#
#  Text
#
#

BEGIN_PGML

[$PAR]*

Drag [`\vec{v}`] into one of the four quadrants, then give the *unit* vector, [`\vec{u}`] in [`\left<x,y\right>`]-form.

[`\vec{u}`] = [@$multiAns->ans_rule(30)@]*

You must complete this problem once for each quadrant:

    [`[@ ($completed[1]==1)?'\boxtimes':'\square'@]*`] Quadrant I

    [`[@ ($completed[2]==1)?'\boxtimes':'\square'@]*`] Quadrant II

    [`[@ ($completed[3]==1)?'\boxtimes':'\square'@]*`] Quadrant III

    [`[@ ($completed[4]==1)?'\boxtimes':'\square'@]*`] Quadrant IV

END_PGML

##############################################################
#
#  Answers
#
#

NAMED_ANS($multiAns->cmp);

$needed = 4;

install_problem_grader(sub {
  my ($answers, $state) = (@_);
  my $result = {
    score => 0,
    errors => '',
    type => 'persistent grader',
    msg => '',
  };

  my $name = $main::PG->{flags}{ANSWER_ENTRY_ORDER}[0];
  my $ans = $answers->{$name};
  my $score = $answers->{$name}{score};
  $result->{score} = $score;
  my $data = $main::inputs_ref->{_data} // "0:0:0:0:0";
  my ($correct,@quadrants) = split(/:/,$data);
  my ($MAG,$THETA) = split(/,/,$ans->{original_student_ans});
  if ($main::inputs_ref->{submitAnswers}) {
    if ($score == 1) {
      # Check whether credit has been given for this quadrant already
      foreach my $Q (0..3) {
        if ($THETA > $Q*90 && $THETA < ($Q+1)*90 && $quadrants[$Q]==1) {
           $ans->{score} = $result->{score} = $score = 0;
           $ans->{ans_message} = "You have already received credit for a vector in this quadrant. <BR> Reposition the vector to another quadrant and try again.";
           last;
        }
      }
    }
    if ($score == 1) {
      # if we're still correct, record credit for the quadrant
      foreach my $Q (0..3) {
        if ($THETA > $Q*90 && $THETA < ($Q+1)*90) {
           $quadrants[$Q] = 1;
           last;
        }
      }
      $correct++ if $correct < $needed;
      $state->{num_of_correct_ans}++;
      my $need = $needed - $correct;
      $ans->{ans_message} = "Good job! <BR>".
        "You need to get the problem correct $need more time".($need == 1 ? "" : "s")." for full credit."
            if $need && $ans->{ans_message} eq '';
    } else {
      $state->{num_of_incorrect_ans}++;
    }
    $state->{recorded_score} = $state->{sub_recorded_score} = $correct/$needed;
    $data = join(':',$correct,@quadrants);
  }

  INSERT_RESPONSE($name,'_data',$data);
  $main::inputs_ref->{_data} = $data;
  push @{ $main::PG->{flags}{KEPT_EXTRA_ANSWERS} }, '_data';
  $result->{msg} = '<input type="hidden" name="_data" value="'.$data.'" />' .
    ($correct ?
      ($correct == $needed ?
        'You have completed the problem.' :
        'You have '.$correct.' out of '.$needed.' correct so far.') :
      'You must get this problem correct '.$needed.' times for full credit.');

  return ($result, $state);
});

sub GEOGEBRA_PARAMS {
    my $result = {
showResetIcon=>"true" ,
enableRightClick=>"false" ,
enableLabelDrags=>"false",
showMenuBar=>"false" ,
showToolBar=>"false",
showAlgebraInput=>"false",
enableShiftDragZoom=>"false",
useBrowserForJS=>"true",
ggbbase64=>"UEsDBBQACAgIAKy2/kwAAAAAAAAAAAAAAAAXAAAAZ2VvZ2VicmFfZGVmYXVsdHMyZC54bWztmltz4jYUgJ+7v0Kjp/YhYBkMbCbOTnZnOs1MNptpMjt9FbYwaoTkWnIw/PqVJfBlgRQMKSRNHiIfWTd/5+joSOLiUzZh4IkkkgruQ9RyICA8ECHlkQ9TNTobwE+XHy4iIiIyTDAYiWSClQ+9vGRRT0ut7sDUxnHsw4BhKWkAQcywyqv4cAoByCQ95+IWT4iMcUDugzGZ4BsRYGVaGSsVn7fb0+m0teyvJZKoHUWqlckQAj1WLn24eDjXzdUqTTumuOs4qP3X1xvb/BnlUmEeEAj0d4RkhFOmpH4kjEwIV0DNYuLDWFCuIGB4SJgP73IJ/DpKCPkNgkUljceBlx9+uZBjMQVi+DcJdJ5KUlLUM0I7L6NffxFMJCDxYb8PQWSToQ9dz9OUWDzGPnRsYYZnJAFPmBU5OFUiMPVN7ggzSZZldU9fRUjsm+6iPKcTQxFIRbQCEAQyJiQ0T/YLkdHGzCi20l4gRBJKkPnwFt9CMFukc5uaIobNPZ0vuvSquWrGSGXkF+0F1u0AhyQmPNSFapRRI8q9gaGcJ0ObvGbI3ZeG3HuHvAky2p3yN15l6zZii1zPwDXpu6Oo0L3mf5JIj7nKuPPO+KCM6xbcbUTXMWydV0rWFLEMZf5fhzFiEjOSHRA8o7yEeGOEArrbLLqoQneOgtxpjDzHYeGpMQ0eOZEyJ1u2mz/8QUO9epn+hA4bqdItof7AtkD+4TWVUa0xqss8r4ZRygNlHMoC7Zc0earqotN1jqGNss1DK2Nf0ptZShLlUsHlfimXht0soPt/G7ZIFct7vuZK77KIMVe58mmPhMQPuqlv/CHBXOZbrbolbdZbgmfP6cx719np6Wzpt26/46TQRKqD+5Eee1hVX7PwaOMC3nK9Y+twB2++lsj+Ic1JGfSO1noYs+o18wqu012PsdU/YbN60p8nSh7fF2IZI7zHa7u7yDVBNk4UkRTzf9uysFlUmeN3S7nQR9/qY/8x7ryp9DpGpx5asW/k2D/U/egg1EPusdX8PODa9uSuyCgRoyMhPtFJs5lmIHh+/r3cXlip4Nh9Y67jAPs4GhFuPa4EIHNMsZljKs+dxc1Ehow8Q+btHNlsU18PPKEZuLI1rmzBK9cmHZt0beIVgJptHo1qY+21KuHzT0tDt9mO5zU5kjep9P8ghOfphCQV13C7lAvj8axz0O2lpKbaLVzBJjvZbBWS0VCb0IRqJZ1p7U1wZrSIh1KwVJH7ICGEl5d01oynNFTjPLDTfY9olpuLfTEWCZ0LrgoYIJ8EV8zc5tUOONZZj/tc/Fqz1f28M+YRKyfjlZVKBdjze1Po56O9dXqpInQWBHstd9BBA6/j9FH/ozfobUkUDRoTrXkbC2Gb5QQ5W1vR/t5mJ5/hrvMZOAnKQ9iOc2DDWNlK/l5klJugUzwZNCazUvTFDv2YCFJZHmRbqSA0eGPhDU4zyihOZqs9vRhhRbIywHgwQuV3CCcIePOnaOxRObRrK1Wu++3HjKimyPFEV7CdUP4ZB49RIlIeri5bB/l0dGzb2gxtKAQjuHREn5dy5Zp5JVDYBGj71eDFZl8wJsHjUGS1xe15H0NlOQNujFC5/l0zA/ZZ886ObgpNjvS2vZU8WzkAbFd+CdVe/trq8gdQSwcIBq3M5MsEAAAPJgAAUEsDBBQACAgIAKy2/kwAAAAAAAAAAAAAAAAXAAAAZ2VvZ2VicmFfZGVmYXVsdHMzZC54bWztl91u2yAUx6/Xp0Dc14bETpsqbhV1F5vUTpt6s1uCTxI2B1wgH+6r7R32TMPguE7bVGvUbZq2XPjwcc4Bfn8byOhisyjQCrQRSmaYRgQjkFzlQs4yvLTT41N8cX40moGawUQzNFV6wWyG09qzjXO1KDn10awsM8wLZozgGJUFs3VIhtcYoY0RZ1J9YAswJeNww+ewYFeKM+uzzK0tz+J4vV5H2/EipWfxbGajjckxcnOVJsNN4cyl2wla9717jxAaf76+CumPhTSWSQ4YuXXkMGXLwhpXhAIWIC2yVQluxkoK3ndjFGwCRYbfS+sWB7yeGeJLvXLxTXCG+zQl+PzozcjM1RqpyRfnl2Grl9DG+0pc+7juS1UojXSGe2mKkQNLeycYTTJc8yrKOXOliJLwo8mQUDqgvRBesAo0WrGi9vYtbGkV9xl965QVBra+buxrlUPoSRp/KRaeMDIWnDgUI1MC5L4UVk+8UpUXvZtPSLixVQHIzgX/KsE4+mknqC68E3kO9bsTYkDMQK4cEKWNU5z4USri3e9I84ptqK9X1Pfe0dDs491UtdigcYgYB8dxL5h+MEkwaYsEbmWYp6mfGS6Zdi+ZS8Tr/lHcaP1IdbYRpiP6uK6+3RGa9A8SmniZyUOR/1JJ99NFTRncmr9/ex62/4o40xaMYLKD/bLueMh98K9z3w/S5ZfQ4ffR13f4uT3wIH7DoQfYo0OP0Nt2h0pfCyNXSucGbcImELYG/1y3KaesPoKaUSL6DERyIERVVHPItZL3HDtN9yj7DcpDvpyX4qdp3/NP6aPjIWmOh3Q4IMkgeTUtDn2l95OVYNtlfqjLXZbpf5YvYXm7ZLnfZ5ulftrWu0zpYXcRkjy9WUYnrwbkV1wgnrw+1I3hjlAFc9drE770RoHGg2BOgjkNZrj3tiEWZSG4sM9LaZZ66i69Tx2ATdeuqsmfUvU+8W85BOnP7t9x5/Ieb/8gnP8AUEsHCMQHrm/gAgAAwgwAAFBLAwQUAAgICACstv5MAAAAAAAAAAAAAAAAFgAAAGdlb2dlYnJhX2phdmFzY3JpcHQuanNLK81LLsnMz1NIT0/yz/PMyyzR0FSoruUCAFBLBwhFzN5dGgAAABgAAABQSwMEFAAICAgArLb+TAAAAAAAAAAAAAAAAAwAAABnZW9nZWJyYS54bWzVGu1y2zbyd/oUGF6mE9/ZMgAS/Eitdpz8aWeSNlPnMp1rOjcQCUmMKVIlIVnK9Wbuke4Z7gXulW4XIClKsmTJdnuJExkEuNjFfu9CvvhmMcnIXJVVWuR9h/WoQ1QeF0maj/rOTA/PQuebr7+4GKlipAalJMOinEjddwRCtvtg1vNCs1tOp30nzmRVpbFDppnUuKXv3DgkTfqO9EQgg2RwFotQnHlePDgLowE9Y2HoBQNBAzdUDiGLKn2eF9/LiaqmMlZX8VhN5KsiltrQG2s9fX5+fnNz02tO1ivK0floNOgtqsQhwFVe9Z364TmgW9t04xpwTik7/+n1K4v+LM0rLfMY6CPHs/TrL55c3KR5UtyQmzTR477jM9chY5WOxigCLhxyjkBTkMNUxTqdqwq2dqaGZz2ZOgZM5vj+iX0iWcuOQ5J0niaq7Du053Kf+x4PPOqzKKLcIUWZqlzXsKymed5gu5in6saixSdD0aNRAOpJq3SQqb4zlFkFXKX5sASJtvNKLzM1kEBVlzOYrw7ETs0/AEk/KkQHirWS6DuhR08FpacBfISg9jgd2oLBkXVRZAYzJb8RRgSFD2EROSV+ACucMEE8WAlhJSAurgnmEZcgCHOJ58Ho4TLz8Z2A/YISxmCZcEo4J5wR7sJUCCJ8IgLcyAHWjwwyCh+EhuPAx8U114WPWXM9+HB8AkTCooFDCNc3TwKhAb/geHyz6IbEi4AQLoiAERfOAPOAEsDoInpmmPAowf+MeIieB4SHBPAB34gZtXm4VuqFDbU0ShFdpTBQBn58+BhtbSjFW1cJaIACb6c4MDvgcX3fvqJ2jbp24Hbw7CAsjGe3exbUcks9C+O5D2WzYdI9hsmwwyRDJkApeHozuATPzcz5cfDqqW+nxtQoo/VqiL8inIBM/NA8PJAn9148sQ5V66VHeHFD0WfscIoPM9GWS34Ulw8UbkOUiQ5RgXHKhCpxC0n3KD53ivYIiv6aE/7R1BkPH0fGXngwzYDeGnbsyOrxD9fDQ6NTK4g7SF6cN5n5ohYCqcYI29DUalKhXNyIBC7xeZsqfUxmdb4MOAkECfxO1jzFvOmLVerExBmupU4RdvInJE8fFwOTjIEeZj+bS7nXpNPTOqH+tpVQIf95qxQIB0RUjBDI1yY0NrkQTsHbbMgFJkQOwRMSMSc+BuMdiRGqxKJKW+GOVTZtZWTkmObTmV6XXTxJGmPQBYDLzBSB9YakiK9fbIpbyUp38UIFtarTbEW1VsY9ucjkQGVQCF+hMRAylxn6raEwLHJNGtPzHYPOVIwXahZnaZLK/B0ovynPvp9NBqok5rFALg0S3E5WpSXrlpauBYmLokyulhWYCln8TZWFSR89zqHa9l0RCephWFnaV1yInid8KFtEEAlPcCgCq1iinQdhLwj9MGQc7BN+gNhy65UX4U99OjW/UloD+xWRC1U14h6VadJ9/q56UWRJK9hpkeb6pZzqWWlaCPC1Enm6zEeZMoI0SoaCO74eFIsrmyt8i+vtcor50dIfjF4WWVGS0vAFAPU4sKOBwYO1UNTAUANRqwmRtu9ZxA2EGQd2NFCgY3u0mlHWcEkbKmlF7Lw2qToMoX1g9T7LU/2qmeg0vl4xivBW+40E11GyR0J5cb5heBfXqsxVZo0oBz3OilllbdiqypxjVqk3Uo8v8+RHNQIPfCMxDGpAbUFXJ05UnE5go12vJSdRq3+Fo9rVRI1K1XBoPdLKtXYdUk1LJZNqrJRupWtNvAtm2GmOf1HFZTpFOyQDCMLXamVrSVpJCOJJhyPktYJDxxhOQG4aZWYO+A4cuwC2zojMExIXk2mRQ1dVQcc60+OiNF2Y1ABuerWborw253yrFprIQTGHN+9kmSLFCoIOmc6yjAzLYoItb6b0l39i9Cvza16T+vLXWaG/mtnBkLWPcztYYPs8kaMarI9pcQ6/Om/1WGnZvn9vpmDCKkPDqdmeTJBAbpKY5dVZxU5J+84z3qMs9NzI84PICzzmQhjuUe76IoJ2kYeRCL0Tq5Nippttc6uOmgAGuExNQHBEG2e1vLaGOzdE0UJJMfgA7zYM207UHDsSS2nwYcOBGW1cmMhsOpatoWRyiTG0E4UN0h+Gw0ppsug7ZxDuljDwztvXRVLbVTeiVgi+JQ+ze0skDvm4IthGCj0Gn4QGvLK9lZEFtQ/fpkmibC4rpjJO9RKrr2DL2ivr0bEstarAb2sAiJ1ou+bMT9+/kuVIkfcg5n/M//m09nSrgS21v0zLOFObaqenhJ6ALFmbdzrajfdrF7wJb3Fq5cXbym2qo452d6j1OKWu1LZP8uIQyQdhnc/Skcob1yQLasCW1Gz+SOu7rwUz86URFvnI7LLZP5G6TBfk0m68tBCXHOwNR9fsu/TsINpzq1/zNW1jgZEO03hTk2oBobHC27RGmJfgJwsobZ4tns1PTskSfp80/Josa6uYdYXVL1oU91KYHxqN4TCww/10dj9Xq43UcILlgUXpdVe3ksUuj/hRLve4w+WWOwz3u0MJ6BrRDf+PzrAS7NktMlzeKu87gthRrgQSy5aQkb/LsVJQJtlu1xbXSk2xovshf1vKvMJr4PWiYrfeWsQb2out4oZbinthX+xX37pzvDhIgbty1eP7CO0FeOPr0sj1GWWBa7MZLocUatUw8n3PNxfCj+wll2V8u5yfsdpZOIprU+bJAbljCsmtFWjykOKAUd5xHU69o8sDauTJxL7qoAIejDlfZlBcqWTLqPc7D/88E9GjefSu8sVUrHdUL69UPtLjDUOcb1kdVMr77S6fTVTZqVpwg+nDZTarlc17oQhZFEDTG3DGmF9XjNVyMiiyNO4wRQ6N4D1G7Q/zIgooGb87xdeFvE3zlv+fk1/InwkLKTkn//3XyW05fpNBi2WDRU/0aOBGoR8JwUXk/REMHtKEdNLvdkSZHdVrzB7eawhmbwKEe3yzgXfWyzYQ3zOc3Df870vl4SGpXKxdMO3pRyoN8zeYUMgqiNJuDtrdsczu6liu1AjX9xgJVL8vTsx001juCANVjbqxhtEd1tLcEu7IPmghrVeg2dwv3R9jEK3+N83i0FKOfRK13A4dt3rdEQvGx6l3/Lmp9+w2z25di/bMXWzoulxQCOPthfDnoPE3RbYcFflBXr2jttTMAg4sQAoDFDYf7qrvLd02ATBnp8joUSJ7qF3VeaabTXmdTXkUQgfHufeolrZWd+x1nMFBjnMEZ8fxIWeLNEtlufw9IuFxSv4d/GK/6NNPX/SPGaU+cW18+PS1sUMZZ5+la2x3Q2qh2d9Z3Q/Zrzuevs/eqVI3F+DkfYnTp8/rb1LIX8izV/Kt+ulnOaiezU9+OYEV+8p8fbK2fXbH9tkt223jug3f6dawT7MtW3f/f/5tx9saOGTU2eT60F4m3DbBR7lTTCvD2WY/b77orqDfHDaHwG+tX+N1hNf8IUTgNL1D3RGANItsptVVXCqVN3+2aS6CI3s/GW2aw3n3izvzbXr9d51f/w9QSwcI9VjjHEoKAACjKgAAUEsDBBQACAgIAKy2/kwAAAAAAAAAAAAAAAAWAAAAZ2VvZ2VicmFfdGh1bWJuYWlsLnBuZ2S8BzTb//c/HpuqUWrPllJV1J4x2tKittqjpVaJvYldqqgatYk9atfeWqtmbCq22jNiJ5Jf9N3P9/zP+Z+Tc0jyfD1fz+fzde/jPh73XiI11V+Q3WK8BQAAyJRfPtcGAHCxv+LUExNif5DgyVLcvHPTfvEUUDHCvA0AyAOUnyvoeuPsZRJl3H193GfvZeOGmxT9h5nzY25u0P036l6EurqTXV3SPzKe5JyLL2wzHIxk1DzcYFZQ0HUdYUhqJB3cQB3uFNjYWNt4tnxaWCTp6LNJPcjSr/zWPrWidr6rr0VEzlz8Sk3tgTLP48f53d3PIz8aisvJQeITEhRdm5qaOOuc0KhLEe+jZcIZHTMzdRoamrtMTLDs5oGBgaEzU/7kl1C2QJTCs/Dbd/rv3r2rqa0d/zKcV9WydCovLl3ai5/FY1LrPkSkemqvdG72vXskgLmsbCwv8k7pxC+7vcH+/nufnWqshx04xrZchOoZPdMQoScSDQD7vYknMybpIpsR4+Panz5/LkxJoR9UNIZoQxiNNETMO/MVObqkRFwZGUrkjGCPeHkjMfmKawlKFGy3x3PHuqTyZZV3jQNs/1z84eBjJGeEifjO0pVUDA29ej+WSxDDIq3HyTuVVz38KUan+WIVGhifWz1cMvs+1X6EgpHcARL+QSlZScktwHmAm0y8H738iIcnZ3XVIqna1EDpyj6RaWYnfq5POlhXyTL9lFeLkbOEfzc5zHpz2tn5fr024fRIydcE1zcbzPokbjEgPYcCOs7ICtMWz7FirbyhFqcIFk/6GDa5rCy/A0kEXXSP6GjR89ipcKVwK1HsdtCve1ofIMP0NHR1ExtrcsYK4t5uUBGaQH4rWxl83k/I0heRtTupAmtDUkKpq60CL/hDT9RDBaSTbPzKEWTIoZ6Pkar+pMYcbpqpHM2zWtsgnxiNsggBSrLyqNOjnpffIlJVrUSYRao9VChgvYoDCl0mu9S7FobK1aX2yLP9ySINtjRRe1q1LJlUVp8/d5p3K1q2vg/y550v+lWqf/U89RbLj3vr8zlORJ3R9HHX0INm7LkWhUfzNTjRkEknfX2yfGVEzigZE63DnPKBunod0V6Itzz4GVSHnvUsbbyQJMQof1IKh838sns3q8Xt/93dAOLJoAJevWrleHlqMLQ/ZzU1sxP3y8VUGTwY8SyuZq4K0e+bPyRaX/3q2UHl+VKgHLzG7AC/ilU2VeZk9Jk7Ylh8YrxA1a4zCKdudfSq23V8/AuzDOgtUNvUiFwb6FCEfYyxBYMJaqaDEbkvmjXwhYnq1bqe7iS8hsjZ8/pXrLLlofoDZZRi99fkFtttoKzvF1t9Ytwq9mxegGiXpSkBwdkcRz9JBS7qtsdesJ22sDApAPCfRgDM93zeWOvp0YIWW1O/19QMk8urqNzy2JnsZz6vsewjN2r1Fsz0O2i6y8Z2uM4sBTCTqMLFZQfcvH5oBPVVO8zVRlpDs9SiaPmEAheM6bxZ8zoDZTYipAnU0kQEpQnbzXLwllXxd/+7hJ2yG/9itcN55ekVCgEdVihENYtjVjf1Ny7r5ZZoll8EqiY9FmxbVgj2b+hsyP7/XYj5bQPRUlIizPQ/7p+stnyh8vLlE74NUuDSmfjGBfw46JD4YmQplDyI7n5ZEID45iV/RLWMA+5AX5X6DGIy3EZyzvZh6R1XfFGXr+UOAO7b4x+kvA4cHDde4LVXAo7ptgAGk0ck//an8CoocKHSuGGE+chrI9VGO10M1MqPdpDqwJ1tcDJHQF0/sJFe+K1TswCiKGWxhymPfbF/JLlLQ5NKFxgwvVRjo9FB59E+i+hW2t7e9vDVRg5+8PiUh4td1RPKh0PS8oa/kSE/7s5ibzYqj+99tdHmd649mMS3AEXtF69KE2XghXITrqmHvR7FqXq5TI3/NLFTzSZLXrlke1Imj6Gg4UDJ+Hj9Evh00iOFDPeOYpDTGsLt+mChZca7qEOMEoUKxBrARuePhRavltJcvGWlYAJKYiDu1wepJ0Obwjx9WytuApjV1tOTUfuOD1kcVfT48tnquBlqXt0f77jPu/RFHgCOHIL8N99sjmRamxGq3pEC0JL4h8gTU338pF6k0XERpx2Dh+pcBJ24yWptdu0VNZEpJL8VIn7+GCeLKsXmci1xv06NYggPB0i6dLbU6JIneY/ShqBOZTKaQfj9CPkbcJTGjVFyfUlKUsIjIs+/hMipvcj0KH9rf/f6O1f6XN96Sk2YDyO0PyjLd4cHqmZAiENJLIgr01iJxZ/CqaxMACvKRg2/uxLHUciaJDoqKo8V6Ds68JV3u4tCLv9tV/jYaM4zu6lvr60/BxVfI8/15xj4tvo/2jCsGt8Zj09Pt2d8GLcWTKqN3IBCNEKJKJKKBYlevjw9cqNpr/v55YHiDvXuCqwFbbZWrHY4lCo8DFOsm4Jfj5TMbqyvpwINza89yu3/YPFnqMDFyry8NrcdeNYsYdCsQupqO1XoINWqMYVsF0aD71xfpxsYoYd3RdOWjSOGnmo/BlIxqEU3gmMyemWzGL/aaCNQD3STAxmBEYywZodjlFmXupuJ7ciKQ9FGa+7birG8lWVdB7wCPP6HD59ANt3IAfz+8MPp8vRL26BDcgOX1V1qL/xF8LlpzchXfTPlcqM6IasLuUMk/sGwAQB7YgS4SL2AF9exMkvAYRtFuywXv6orAAUB2xsO4qCgaIDGXlfDyWWsagQ5S7H9TMVBKtpbT759H8og1xTaxZIkPCLRtbsKdd0CPCG4S0mshctHBHmxABP30iCkokmXF57eI9FyeiOGnYoFN2M6PfDUbUWjnaOTJLgeBLOuQLDkkDGJjU6XGw//CMkg+cJNdMy3Bbig+iFFHEYfLGp8NpqktltSKzN3QaN5+pC92ySoCet3T3FlSNkx3zfN0fbTZYLmCIRDCBvv2eESJJyCzXGE/lcm4VIIeZDAnbJQrC1ujZbYExuRJWD9rxY/TdxJvY/ZLVpyU24JZ/a7Fek+rAHqz/zL8kTKY+fdZwii1Xi5lJQFYH7fgAiHEns6gGB27k6Oq8unUCfcI5j89dWpm/FQKNH+bfU+GYDju5sD/DtqlL2zEuXTTrhlJ6rB4FCUjfcmgmrnfxCTckeeiH0I95aX/H8Y0EmO4w7Y03zKQvnf5cTCuFuAMEP5FXLi/y6Zvbe8OhaSdRm8T0xDR5fQ4XVrN95KErrt4LU/p4jFKl44J4He5U5p5dDlJ8kVHLqkuLm+Cbbvr6ytucqxcORPMf4eU96wmFsDlW3pFallBE67V5dpaOc3PwbKgEtqX/6uBp3egTXvvQOObzs1qZVnJIaGN4lU7MV6qmjF8O82vmnzjFKKFVPCRjew/8iapfJ3DdsTCQZjMulFIs7bjbK7cJkIhJj0h9OXUF7GjcccU11nqpISNDB5HZyZnnl+aTqihCKaxLuy3vlC87LSazzn46oCY2uxcuqWfZ89SiFCVrfUUp4EOzauELoxdQws86xVbNU8sRlhbj4ZVx0mlY2Nfpfqd9iWPF1hqorKfN1oig5AXcCnuRATFawa99qD4MqvfeVfehhK126E2EVkFmn4ryQMVMmwUVJJ2oEHlcozusWme4YKLunyXepOjopE6hlNp11+XO3RxrCCokWQWb2ySXE741tvF2VU06mrv15Z9pj7muD7OdTNcPEQFVkMIikgUWX2QTQNzbL7B5X8ep2B47lf4dybTTEazSgb0JqG/eVuZYdYkcyJDS9RVpEss6QrJxAboh0ZKuz7jV2z9o1hMqcGsL2Aq3rz5AxduplzGzdT3wtkeP3wBGEkPw+eKjTnlyFo9rLWv1p2f+DljjHovf/4ryrZEjAjkJFBg3b6e6ZGOX9JRq0szLH/qCjfwJWKoYJ3aq02erL9zcab39OjsqK1WV/SiD8mMBx9aK8Xsuy9PdfqUyP5TBt0esTyz4qeCwV/YcT/1dBUJo2fBUHZBg/C+9icFufWaj8+/Bc0AbjBJGG6ROxfmFC30W3ni3nmAZeKqDJOzfRZ8n8DVl6x/5ECBBdKMAtaPJfoI2kP/mfVbVTEqeSdRGz+b6yBWF4ZNfyD36kVHk4S0ofvvjcbZ9TgqLaZ6WXW9jNORC8D8myQTiB3nXdaWzl1YQaX126zXs2adbMnKnm83YQJIc6Ex0mxB1Mt1a8sVIkMLib9d3NW/Ke5eJ2Hx8aVF5kBV9uFl14LHsmXqVhsLlT48/V3jU06LAN9WF6yqf1UhXpH3JnydBcWmlZ7IC9XMVwxFf39utROM/lssvD6oZGdyRC3WYCYqKpqLin3wi7mp5nWcy+IF3kdKFVhm2MkT1J1+qJSv1TP/3/R/TX+01U+QIAXJlYO/WRpCXMYiLE3knXf0kTUKCdw1bqsKSEO5LFknjl297DSXGXtADn7qs3g/W2lP+brHwxh2/msDLraWsrV6EXAt6Gi2cK2V8b6h0MFrOXaW7mkfV94vzGoIQ+v3ed+002s5/k2gw2Fd0E7fqvrqCyKNhE3UVPJshak9ON/GBIULv9yHudiCYO22ZDaSMWjEzBzuPI5X/Bqud6fLKGG4z/FShIXJ7d3Oa1B/C37vFeecGu38LJM3kC4NvI63L/Sw94KJuBkHFVrvRkcOqjGM/ea5vOH3ZmmDkgARvQ1XaDloI3q2erqIVcowX+I9uweJRVLEDmLVAr0ukgD4ng1MIF/vj3RN3hIZ+YdXNG886gSFNH8PWDYaA6iz1yViucntH3QEJg/di2d4mF0JpXVDr4S8rquMz5rTgbSPXr0zcz4D8iGVTroWzOJoq6Ps8pKNTeLvd5S2wl185udmo3j0UHWbOErycCILt3PmP+Q8udn+ZfugMc7sbJXxU2zleZy9gmX/RM61lZWqQiAN9iFmLePfEdpwfcsu7CEpPZeXyA3e7mHOKOXRaAAX6YfhVclRCw/z++3W07XEaJUn/RcW0UXN1NomVgFtBCXde1/j2UiRuCERmTRe8NJRvHJ4BQds4YcBiXMZ9AaPVAeM1X+hgT/v4U8x3/6Da8TetVk4jYsaqZ+vB8IfYFrZv5I//i7Z/ODmUPEEo6z/pHBxdlpxHBdvWVFjmM/q7k3v1JD/mmv2/WO9qulx4bKLL4sCslZUqPzJkkdse1GOqZryNilewFTLkedXf1TAeygXzXIQgUPJgPc/4IKTlRQaXZ2p1H9df/hKqLMfzweAnHiglb/kbC8A91KTkiwBAl1bA69mq5vbZU598qU9hIjQsEFAnaRaulid5bQm9uB5SDUe3YBxO1gh6K6TCYLsEPHptjiQUnEDKRxzYLxtIEuusBM5OECTAmdRboLVydMVA0M/ABZaih8cCir+yyC5x2qr1noKlNJNcA77gpyQmcvzIeGc6iMbGm0iKGf9rQkNTY0bHncBsKPMGMdB8wvQFJsGolNYUrMwLCk66vQFvVsAnM5vKmaenvLQWQcJ+/9mUp7I5tBLqgGNubgwXqRnzjjog5l/9h5ZVvtSsRvrdTXNuczi/jWCq6CNZo1bPdqRp5R0TICmY2zBk3TfinmqXZ9PucgwasQAI4Yhkdh+c9K5wz5bsq3cYfDwL/YPSOlLjOpifzCy8MTOT4xsa8lZRClodffHxywSOj4iUVZTc1Df/n2K/8g751NKRISEjNY2aFPFkYhlHCsO5I6kgj/WP10JYItrRKnsKCgCWZl8LiLIW5WTzWLt6q9u/t5JSgp26YtocIn61FMx9ytQ3z/MUvgcyt3s/PFxPjpmGlHcwUGnuq6j0njjffUttzW6NCqwvqjB5hXhdKyF+3f0OiCk84WxNbIyhQesXAaFzVy870r5wOGusbBetAa/yd1taJJqYMV7yvdKIWDyh3e+gNCvMfO0WWJUJxvFwtlWKW9POtKJJ+2rfUtgsJ5kbuABRDOMbXu9/iCYM8mTxcbso0kS8Dz/cty2ODRnlDt8SnzUHIeerDqlXDWet/qIJG1yKJUjWhIoaqr4vNgxJKB2JwS6+1Gv3NHoxprskzPObuN7I71weR0uj4l4Z3sVZvDppDCSntFmGdlzjSRAHqp44JqPYnV/+Ph22U26IKtGYrnmrj5nYFT8/W103YMz2azvczp8cwflRHnSqZzvqH1YttkF7nVNzyPs+PnEPcNReH/Q1dRYmzcCSrfzOKXu0Uv4czuIWSGsS/W1vMOkkUvA6/AJEvjlW0OaLhh88r6VGKmeLrFIGkM7N1lr8DJA/fjfUxeWCg1zt5cbXwhw+UtoRqEbU8Gy2uyentFlw0ZHC2Rg2B9nl6gjfzdAusTZ97hRadhyq2x1JfKeflzjkKNBdX/C3+iJGHdz+VFhx0v+J4RnRc28QqrKSI8NyaW9fwRbda/L2CDk7v53xQih7Vb2aJu0/uIwUWby5zl6g78j+p6EkdO5r+2lTLOtadfECq85mafbqRxcBCbroQ4C/UUihQNvjjLngkweYasOf4YuOdndtjiwaQU+bp8OfAf8aOkJm6SBlhQfMVV9JmCLtfcGkTVx7dY0pwyXYp3LMF7+fv0PW4PqKqU7tLjAV4snapM+Y74X0G5SYfWnTAILyEdYndgOGvT+/zq4ZTTw8HURCoN3e9+DHyL/kEDn6KKv4mNhTRV6tUFY/6JT3mFbK1DwAXTm55QCW/KN+C7rO9+herJkFI0fmOs3Wyq2zjuc7pO5m9FK6v0S6Sp7xgAO0drNlYNIWOpNpk6zyJeb9aXRqRGwbeHt+E8+mMtS3yFln2vcqxd4Bv5chqyzrzgo1+iIvs/XF0f5I+vHQZozJh7xz9QF4SX8zQUxztK7ro3b7Rq6TLQqOmjrtV9z/oDNU1OEbhe4P8Lc6a4MXM4jfBepm0UAsjVI1iG3sqL3dn/jabWkHmrTvmHZC3ZPCDHrG3+2bZMx7HRgZAD83Cm3Ih9+uEAfHPp6TOMQxvBMrJ80/euKaYzqtd8SgzxlTqPw3XBFCOEltm9NA/Y6tLQ5Ri/bggv/8eB5CeLWQABX0fPLeax+to3irTDegP5zRCDvt4tB4vg/zOUHJVOJNq72Ben5d9Z3gSFURzKladW/2PxFDcs/gXzCsE//vWROBuP/eeHgf8lKFjxGUUZvrCEpYVRykAbRJpC9WK9hNilWfupmwz6+S5olO3iXEc599+k5ys/rSWxrHXVJm+luMfwxd+87oxgg5svLUwsbY22P+gbBweQP7+vqUDZzWys++aDjzCrmeL+oLLMvAoNaWUyGdjydkiB0GOBr1hI/2OnsDBIZTrXLcIIqe0XuZm7X7XWlWpxSG3Q5FtfSf79Peo+vfdsg8pWtYLjA8r7ayMkGOX1/NPM6HKdvfw+yUM5k9pNGh63ohMJGmNOvG9zgUdyxP+TNkAq20lBgL+xLWJj2OzAAVav9hX6u/SOHO9gEh9ibkejQZ/g/zIt9NS2k+SdlrFAJayMQvz0G+ApGnWDAjfrhg8GKkxbbEezFQJK2QIuVrrpwKU3Ka/JBidzexrI/85PC5+xDhuExf1J5fzWg5nTxr9bqd7AlqnzD4KAniZQaZUj658P/0s+xb3M1pIGaC/LHIUUna2HyaLR3kG3YrUXqt7iLQQg98c3tNj+Pcf74UGotu9Kz7t+47m5uUVK46JvLY1vRmGwHF892cQ1suPRLOG/SdnfUBKXSAOYBXkzBYpOTq9bbnhvH1E7STil3T8zAOA44d6bwSE+2YUsNB+2Iz8eUmX/OwEA8aQ6+zspgHG9Q05trcjOQkuDJPmbs73fZlmgpfZMibajfOWE1npLgv8yRULERywWd5E3FLjKaUl2Z7qcCc6NGdgskRFoO9EWEPA727MjB2CNtmHGs0//Ea8tvOxw64RjNaaPJkJL/TSK6HM67of/Hd7HGyZV0YH2TRMDMcLdk6EyiOhtVYgsuNV/VR0n2b6ybdq+xZL3q9XJtKW+zjl5ue1yleDJ+s9v+sxjbB6PcTX/cwR29uVVr5WCwy77d7OF8q7tV8Bd+PGf/gS75c7gDRbA7jSGr2F+tT596brh1hLJch9BO+hqp6H9HFWbVF6/c/XnatpccSxAarWVyG69m+jhvYIHrwueVrefO/bHwxcZvWOoyV3i5ori4uasGT1HxSLB6tvwWqhRPcLu98zoTS74OimDxbxM6sfjCXpZyTvGgwK8hJ6ymZ8SdOoNfTlb64T8KxJ3qt3hRQs/fS/q3Er47RzKr8gK+noS1u7S0cFOYWV9CVvO4wR2byU1TPV+0kBBew+5d6rq4dlyAW2xcljy56ZjZWFxnwV8MnYgJdB+kdvU9rpUb0MCx7Er/rx2TKfja/nQt75948EsoQ1XH8H9Ef7sJL40X3APqNGrsfjuzvzAy/s6hs8Xoo1KKaYf8/AUWkC6bet2jsc16yzU3rzzwEfpkjTyZGvgcNZihecfhihA4X9xvcqC9y2/VSf5e8vSMZqtxiPSq2Wlv3K1WURoaB7BrJOOmfYyAj52yrQvwx7k5BZWL7KMrz0tfYEFAewbW0lnouwcy0F0FWs7ko/HU7LRe9L/6dOQfNCEXclt3jd5MuuN1D9CxWQa4FJODfVw8vd8COS1iZTn3tZi7+4N6w0crXrbNbR/eyhesEPiwtzfv9WvZSq4esdnB57IA+tI0xa+0owbW7MXXhEc8Ecrnc/4318hxe/4vMWvtJl3VR94t6ePWzNEw5znC78XkyxpWYTGYpo482OMf9rtxgHpWt1wJkQMSJ6AqOMNBEvwXLVNGTfGG43my8Dt/NZZS/buBFL8MrBT0+mtsd6XHEXVsdykCd9Ht1Ysf5bMvp91KSvo+3Ai0XDGPeyj2JUB7dOGiDBfhOw9dyiP6qjFe5Ns+jIap8ozly0QNWZywrwNlZ4OISAd2reBgOsvfYYE+z76WSeRL0HJOsBWZigf5Ge6wsuAGI//Qbcy7r0pnOUZQNPJ4dISWboAeh657trLYHOTLRzWJygAcnmFypKafpaGUK19c+zI8FpY8m/vefTPq8O4gr/sAFrPF/02QHTto4cLLckHTZusUogh4f60S6fpMkMRrGz0Pftts8vHWDHpqatp4kGKvubmL23wHKP8N4nI30miGYTDXV1dt7AUtPBGZ6nLYRVik1vRZSmmCBKQi7XF1EBUGRvBP+iSZ3+CA1mxC+YlHBQvKSkRcVntZtzIcP5J0h1ocpNgC5iGBRIY1lrsyu5MJPSd5/0Pireob2AsJrjFedDhk4qjOtjPo6g4Z2YK9LQA/rnepnxHH1emwYLh310YvwKkAB7LS/UWLh4xkh36AkQGc81bdOfMgOcwu56NPYBo9f8Njb0ZWnUYcGz+kK9kB4jxM/e//2b/eE/0rs1Qyejbo34zj8/7aQolfOSIuXAaymJ93B7jIS+Iz56qb83K6ZwIIx72cZjjxkgRLqfZVwia/0Z0I5dDKcb7uf3MnwM7ZviX/A7buteIyn5OkFLQP+k0IkMx0LyQvPSAJtJPlBF9pacOzA4KqsEfKLEqZRVxPK45OQFvjm79ANoNNSK3YeKHLtO3bW/uIYkb40uUjVVK7mOKdD4KyAkI+l3h6emcAQWvTpH66fsCEKxe1N3vdect4jDc4PG2xGSdVwIoN5meB2I552YpT0sbDKm1IkLwqqoFSRz6F8dTdnJa4kieD6cKF8UXkT3HwbnJpWdw5/kgNgpXDI1pdRCf8bbnWmTz+g6BH8q6DZP0/Wd2S4qyvuvBGX7g31Il3Z6og4w3AaEFjfSxJUMcKcrPm9sDMQGWaT6G/YODCSTMZeDlTpzqjN527JSSuX0pdOnSGmQVUaf0/SLdvtZDKfeSXB7VNzU1WcAM5DmKVrLfVvzWl6woUzVq9aDR+4xBkPQOcqs9YvNe7wtSI2OTlVn2Cm4/mqYuTjr3VItuJIjBzlxWVjaw0OLlx2Tal6akcd1dT95FHgBKNvVR9Fm3iZIrzbjanhwgWvTvNFtjrRPZdKadydAXGYqpPWKdnjOezGDL+MRnbh6d6rgvuao1yN+ixIJHGQfklyDDPX2VvfmWjyfnmvMR/3XwqknMlW5lFvCZ26nym4UDj5ODea/+LXzNV7q3WSYOUq4cWEEf2HHcAU+f2D5hwH+Mu244aNt1ZtJ8Fm/RvJ+UPjRcwkI4ei0e3fBtWfgnyZ7ArEblsE/MxcSjonO5THIHmTflq8JPSbLxKHOCI2ArCtlCOaJqyKNnAjbaxvXnRn0qf6gubFde1L6DUQkA4BXejJNvKYfnJSbspfXu39hvHJ2bN7MSvknv46+bWX6O9D74+juhnYVlaidXlZVMNfG+w6rvxcC7PA3TEbdq7xBiaSBxKPvPv3mTKmIZDAJRy2ETUbdz6vM8guOh1K21563Zz49unYURk8uaEGHtQg53LeLzuV+cV0/j1mRAUTik6+ilGoAFICpFHPK3JHXno1vPmF0ERKlmeuRY6GCyvYjKC1dTvvM+pQhL0HC6+ESjy9pWfoJ6Lgk1l6aWVpwVx7/ykgOVQaPVPXD7HzfG3/b5SHV5V38hljE4cTpftlE+JfGng2WsyN1u2swStuFnknJ/eLo1FqqtrW271A6OoyK6+CvOIl69gdf3nGyvavDcEwt2ZHWrkSQ+yhFgy7VXMnH7BcoLPuIY3fIC8LcePUWsYq52K6nNUU2ooy4KXTMz9RlTL9OLOxEvbrjCKAA/CWSuxCaASqGypN7uHRPXfuV+uynXVGS9lqfbC4DdULJ8nicgy/9Y+YaV1L2f+sTf9IeqAXRTGUw299sXPRf7pwwuuKu8kRZp7W+kHvQ8GCPLc8qeTTHwUjwgxjq7HG7Fp0+3EDP1HQJ4oIlfzsZZrr2oXvfNTK+0R47N21xChYKeMEE36yl8GpMlHa4vwHYYEHZ8/Xkm+RpFu4CeFWlehTiNQ/z0CoitgbCk4aLL3YYpSpeC01ydcK/S6oXvlW8wjr57P0pOJlpfixzPk7h3hSGijZ2nwKVQOcy8an12dlB8fDy5o2qB8Aj9W34nA4mEA98jhm8v7dTO5US++tZW7BDDFyj5p4uaLBByJdHh4un3vyZMN8xPDShLOoMXY3PzStdsf1Nu6VgyDTpHYjczWba0z2jqwUKHvmt38kSax2m6fKzga9yA8m6N3FDQ2txotWm1JQG4+oRURUO9SJlbPdd2DWKK9tytljMRYTb4fCqRo5rX9z1Wbp92umszPPxMea7dqmNBIpUF0E6veovpwdeVM+D6YiAGLEQlnkLpMwqx8RPoQH2HMWswnCsxkWbt++hV/jJUw3th5ugodb1I1WA27bZbHljg1Z52aLhgEnAU9VpkgIIj6XfdMfeO6x+Gyd1GVbtFOicdlQSxk6Ui1pxT1QBhxw4G7tNxPvcRyc1J24k55z91AuM8RcwvKoEuTNenZ9wlppi1EtvdhsnNmM3tzMvdrfl2vgUxI6F4qDwzbt2P0wKh+SQR9fGrpDYjyZKbyn5KfEJDr1iJacTP4Ou4nYmL/LixvWdGDLKSx5CwD73cSXF7cmmz/L/bw2SNjxl6+ron6rhNVLW+at3v4Lxd21rnarShDZmVLbkN60VSsunE3z+N3HWzewjkz9PEsp1uIuok3XAaeobE+WbHVriSDjeGqt7x0ljJHwXv27DAssKe3t5mq7LD3as/oSNYHjFZuKvl6KJvktFiphVAe33mMRwlp99+lu+o1JbWMDaggthKzIMKwZBIMCb7z8n0yfGfH+7rOggey9kuXRGnR6B0RJ0xpOMDZhFlk7f/0Bc9xh3JH5hV3yb146MnAZmgIwwIr+yz25PG3LSDuBG9oR1fhTxb1ruQ/zLgbJ9VTTbveyquRCjzlbU8Y5J7o5MoM0yyo+n+VFfHfe/A9TVN7drm/DO7E1vL/LiJ2b4yPAgH6oVCXlGOZYv3rsWV62VyYsaXV/psfqIPuAcgCNVzom0Xx7HJGzy/DOE5jQQlDoi4O/A5tkKOLyEYzh5+aeS5+IM50ORCjU2sYRTZ/mhO6AxkCta0lRvlUPuuQYVATWLde2c0L888GJyurxWleiCocVW2hBCYgvZAmAyFkXhGOmYepA0a57TWLddKGdYBtNCO2LNEKvD1KlJkJwReK+rGvT9veNrg2EVYQy4KjCesE3dEMP1ucj2nMpXcuPi23/vKtI9NPdH8jUjQX5kni59kw18USwd+Ba+Gu3c9PLRBy76uvoWLQ/IrwBPfQemsZ7rceAH3ht2K2k8PtcnaTz++GjRHW1Qg8oLnBljoslSF6w8KA9uucFiW8tRAr+3GIihGKwsy7a5inhWfN9aJ7qDhC6q1iIZ1EeshdUbMjNwYBMpoxWxSd94se304TASLzJ22e+2l9c1xVhcUcAUTFbHcV7/aPoe0nXwd8lsywSB89LUzfhOo5t+zuZ2bLLD9bMFVLr9Y8nB0TTWyAuJfmMysd8OcgpYQm/rXZWNomSaNSDgFAX4+PmMB3hPwlXtDZh+rX5lfKh1rwMVbt0yn6ttY1tNat/d6ZXngzHwpP6LoNfAR/TTi1RodUiZwaghUQJBue0SAyWUhIHroN/eu4xTtBoPV9lxYDzb3nCx4QSce9LpG1dtMTA+zafsHMe+aq0NdG77Ga9bUHO3yLtHTYUQPpKZ/dnldwwyic42LzOxAueKyxvI10ixK5+Ze914r9YzHJcs+jcVYRSP9/oPnjh0bU1QlGxhofdbF7/LyZulVDr+FJmwD2lcML+qRdOb+bx1vakSRfYAjYe7AB/w8Bb5uvqcsYTsxRMmPouycY++dDByWFlsxa3nG3Gq9l2x9+budG0G94zWkWOu76NfbOJdWMBCVBivn+0ojaiNc9njhxSH0IsaCOtr8GtTWZ+YwWztM1cN9Eyq87yxHyXXs9CWgTOR2fD9FCarbPmR3JsZHmc6Q50j65U5/HOQ3bfYmKTOsCS+rxsFq6fgZr3ZePsKsZ6SBaern8+fs+/PHqYczB6f6rK1pRll+dKZijeV+xsMf3s3FpSr//u1nd4n67e7QsdO8u6GRB2rqEVWM5mtDOZaD19IgsjlAJbeiRaTlPg2c72Y9zJydeOihqA7eH/4bUMR1XlBnDpVtInnnn/7XDSvGtcWMEjaSgp57Bm7wD3Rm2RKWYQGmVsskgbnIg6ZQIxZ/9EEl+qtXm7qifUf+hLr0MEra7yfpwfuEXkwq+tCufY1SLbIXRqyWbq0J4Qgfr1RtmKiwWmreLR85XsKxTduoWPW6l+MFZYSxte3fcoPtNrDprEk7/JU95ri8Qvop2LWZo3tBgSYoMMlbU9VsLSnAC2aXMMkZDOoSseAph75Yq7Udu9+8Xw+bDwUsmLOu/txzas+pbL8aQBmvxb0rdhiLrXXt4e8ob9iTTksKgBdfmrpOvWxjPB1Jl9hZrleZizdFeGI+BB6zLnXGMpqZvPtYy6S2LWrURXhDmI64Ot2QHXIYMKJeDnpEd3gc4LBlymTPzs4x+tASMNuNJeha5wcLw+IAb36AOkrYZoRZaoM+vSOad7AEnE5bC5qR9d3osxKBd7cHADF9QEz2bP7O+WYyEK3h66Y/9TLIUqA5c0Smi942pS6Ja+mYc8F7JlMkW5XJS7rcj83ucYutPCK05ux1otbkDx6xDuOOCKXYlsgtyFL6KUU72UtFQ0nHy+rQtkPnIbB1/OUCqIGjL5iJihbrsygqgyRD0QY585PL13IyiPMoEiaKYPk8D4CQkFC0zcHvg5K7BdQsDqnX4T5wm6VKQeuzdPplvqCWg5ZBUug3T9j2sCrG0rFlqBH2xFOeWqpa/0VatS9pG3oAs/WxBqPC1nYh8D7dRGkNKhN8AIZbfj8APPYL/NRSPmYqkDs+LrClEQqbiEA0jIvXzIx+JEVOxELbLKclr+MO+5FtQSpSSMcA5L1Ifhm3jy0sf700SAAnABp4hdre2A9w62TLukTe/+kLpLRNfAHI8rS3lrS45ALJEFJ67u/8fscySCDaHvhnBH4AIpX1+SxV/zLvItHavWu31HbntU1Lh/0CKNsv9YriSVLedjmENg2W1Xp1f64K5XfVhwKK7MTbb9cYq3+rsH+87pil1vL9uPg6SGGySXlhjak1rBoL3ZSC+BfJGCKE3BOn8yWRKy7EWWC3BMULlgejtl4A8qnB2oouHnsElVGVLx0ZmxUwqrJlymzNLF9lDVpzu+lEds6r16PWZtPx0UOews+mzVoJk2kT78XY/AWtKnbUVRtox4u1ukp+pvCLSclOZYg6L6w21b7b9zRqAKY314N2frc2ZczBvR/JSTi5q6ufWoYw0lTjYKW0BUlKJDiJKRKsPb7dnFDRFdTxcYE6w2W+J3U7v6CqSjNDJtjD+ZhNPZudx8fHB6e0L98rOy/cPfsnnqTullbUZmiS2TsE6a0SK1goFBNjcx38J8HoPDvPmKljKyZKYvdtoHf7BW2m72hkL1uy/PekyLrLd1NwVK1ZeuIIHRQ0aRVVxGo23Zt0zqHs0NT+devnTPjcdJ2rNrkDE1T13lv+GF5kSUy5XnJi6E1rRuuh8jlJDHja5+XPp3kRlvbD62YwWRC4wXVT5cp9u0jX0ExZn9mDgIbuLjWmRkdkr9H0/1s1ezuiM+imTmIX5/qld/r7ix1q5M+EVWW2ONjU2iqO9/b2dhK4XvYyPyXDBEHnavpz7+X79+9Np977C9TFvlW2MijRDffCbd/67WM1NsMTyT/dRcehbMGoCnlOMShBY/xZS+VT3zv1pB8e/u/X85lrhwp8LwCbUWH5eW/vxe9YsGgzTD/4Otb7kpe7LlaaAcWNpcLJ3UmQy8CaDSYVQS6JlquU0hKfmg29FBBYqkfdj3wmSz9/6JLZMPP9Aq1UTbxuLGd9u95Ld/hzb2ZjdAKttC8j+HbQJM/kTctOsPP9Ktc+FURqYBk8/A42DEAesi/jB39hwVvGZ976W9kO56zCKr8rrALcqAlgPcIGCvZqjWdUL2BZWCV3lB/8N6Hv/BGAOd7EBpLNWP1r5hByUbObL38G09x8aeuNKzOHAXPwSSwddDhOenhPVa7EPQtlw90CXASHqP0VP5r4SU7m6myPLzpbL5ZENrheHcI0s8ccM/GfEFepPy/+m5Y9fdbJfOR1tAptuBZsCLx2VuSVwm3Ey14j5PlbHbpjUA9OTR1PdkKuJXcQszwSp8RqYkbix2kEf6vD2AVC2rPfOwaq74Wao2kZ/MW4cGRwbV9djZJhPewMX2/fRHFiQ41DP3owY1H8nti9PXB4aI4CwAvwlEO9igpr8fgkR7cD2QTQKSZj/M0+J8N32ZjYQsw6/zg/CJIGrDAaXhBgN1KMP71n9IyNHzXKvA8VRq2sLhWvBLnK4Zv09PRYhxDJLRHMd5PQeaNd7d5Kbmf0PubIxINsbm7m08TnjnUhImv4s8TNtPOby8580HZYg1vOZ1HvTmI4al0dz/2vnokmIHGzM/ntxVs/bgJgNN1LwFo4xeEz9R49SzYIIv9xvq/eoRmofRPagcL4/jSPkWeKdB/Yy0zIb/aWs2AgcxtGswMjasBpiE/heLtjd3mIZAWeOPGQlmiQlffRsLEJqGd1ciup/SSQ3cp3vS3/qCs4bsD25KjfOOH2y29Rp7YJOMqSWKvd4KqLoXAoz9BN/tsKRTFoG49j5w7vVxlQyHvBxHHTjCdLogpVpAKihiA1Vw+6uUdj8pzyE/ENx3kj+RcjGm1nqzPcA7+4WrKYiGzSkBFNPzhV5uiql3SZf5+A89sqQhDADL5781yLiXBLvSoKzzfabdBSGSIWVA8anyxXslblRhLeZ+8ggsmHZxv+tbJXuMDDHTu09tRc5wSyXfrc9FPk5xzj3vmn/fDjxG2jhS94OGQ48zijmqG3iYPkf935xTOUM3SdmpiQYCmLO6tT3zcYkSfFptiFlowOaNqn+ySEVr6FOfaC7gdqT82j/ZloCfBDUgFSgB/MLDcVDk0nfOuyLAUIHzram/KNohXBAy7OZtiGzeTEpkZR4sXdRzy8d9Qma3xn3iOSA911a1OnBHj4+xf4aeXZU26SLUNjwaAbo1Za3gu2ocZY2dikewGMam1XHJH12w1js/cst7/CKxsbfhsum6P9YRCisAXf43LvKSi5yCQJI9Z38NUeviHG7sTyTuNqI2+x/q7Z/dPpHkjg5Row8I6ULjnVzRhl9jeE2DFrd76YobCBl3UfaxZOkpe0P0/aY+ikbaepO40Bvp5f8HEoie/gGn9vGO1EEmrP8JhuT5mbn36/qjKmyJM9P4HP6cFuQY/66AKzYNinubJqzsTnTzd0k8rZ2OoUu+l/4cLSNq4O4VTahZaPrNIYwWbkYEvb3ONHPLkxfuK+INLan3QYL7R3n7/qspRlbaArs9kpafctWCjk5kTZOAyIcW9YJ1ebcKTsFeyFNek5OIDtQfO3FxhELxmFU3a3wSEIPQkXDtxdDdzlaUhzun1RGgFip3yE/4SQXfc+gbQCAF8X36igXggylZyUPP5pn3drVYHOry4h6YtS+biRDwhhathgwwJuP3Zzi2dDr8fC0O1hVayf8E3GOYn3OLlNquJTXmaMQKHim2VJv/6sryt/S/fwaX+ftfdMu9bXm9zR7vv9Ev6oOpG3t5EC69kxfIcvhn46rlwxc9uNDShLVoFQjo1hN67QIlZvZKpNzgico4e3cy9w7816LRztDq29mQux09U2e6rT+0ow4b7Pd4k6aAlOLE37Rcc1Etbto8MVm/KKpocmUcfk/PKD5WAofgGJ5d7FH79521kXpqzb/rCTBn7gSLnYX28ZQWSnBikP/CiZ/S6Qlv/M83eThMtTX99mGm5jiL/2K/UkOaTFC20dHdD2URlnXSyFRMnt3YSBI/VOurbGRgkNNxCfqns1azu9sjQjIyTjE1/fwweKSZwfrrFb8a5h5XMrYZXtFFVWnFUq4f9crqt049LlGaYQWcOryH4pDmDXnq2BKxsf6iOrLObtJ7cpmc8zB0b26ONN5dDCDMkPnAkoCQZYuHUfEoduX9LEz6zaQreg7QBzyVyiTMg5tabrUNcflJLIAnlCmH0AviV607sU1DlUJ7yN+rE5c/QHDT2g8IYHImMbhq+oI9TumqLsH/4gpCR+ikuPRxmmEKx84zeqQRDMJs36WKNlr4+ojNeEia8PatAM5t3DuA8dlFmuzM8Oo9OSZy8jXiaXN/hAPXKTFuOgZKK2Hzk5R8xBFszdit7sFOnYaUYoHIBaOxuufWrlbE5bXE5i6hcpj94DeNkp7+M7E7E/xyUQxhoXeUynuLj/s7bDA7UPBB/u0lG/qPBORp7Kl657tMTgBBI5Ic2x+ql1xyb4h40BlTwR++OuW0XYjfmTuNtOBx9jkg/FGShn1qMwxh0dVuR3W62G+BLp9c/nAe1SVeS4Oc8BLEGan2gdsLHKlhLX3qtS1Px0rfcSiBUaTPjCjGRCN85sKP+GnPgmXS8jgG7IdGXiunShlasVvX8T5+7I6oYCiEX4cRCpHYL9sbxzigkaadFRgWvheHMsXEENAGZCco+bOMva6MFMI6679nb24Aq4O2yD3pvLT+DSeh/iH//9STdu8v2PWI+8c9ehkTELfyAKsb3ewjW+gPGNUuma6PE9f9+VRXay4eSUtn/sh86gABDg34+5yQHTulwweeJvyPjxgjv4o1a7RmrCXWwH0Xwi4H73T58/cQVMQdW5OsOr8ctD5NeoPhp9f6dcwBlJmj5IIu19YGe3cjnlo26rPPbrz8bGJzdybMBBoVB9CEvVT7rJelD0NbJEOYGLaoI5qtyTbfyXnYRtiuraQPCKZawbzhEk33tnB9ZQCb/dEiTVkz8hZBk43/YWDy64bFzReqzie7meVrN4OY0wOZ9ca2qrt1KLbgTGZEy2D82/uvK+7NtwkQ448lb5XJM88fxbxHVJ7KyrtOen8MyNxtGtt0hRV8YNRiAjuXpPfc5po5T71hteu6BN9V+G1yL9QSJ01iPj77hLgovAiE1oEQ4eobfjNzfDvAgvTfIvIgxfdOl2V0jjRRhGh1KFe5ixbmZTOs6J93/rly4QK3lJZsoaSd/ivr3qFWizqM5Ml0Zf/YOxiZEZ5WtWFwe5xMBhbIljj0O5WMXjpGr3NVVBh5Wcvo/7faPKp97wQRg909aeupHb1IhiUcFtRO+hGmH+TX871r/2vidcu5dbEXXr/61ol+HLQC416KVwGq5DRTWrBrWwBiXxiWkAa6IZXFVvDwl91xrEX9XJLAsAtZLQ83Z+ORT4GfeDyIN08nOI8BkT5ZljJBGvbuOfvB1quF5KQYSnzvZlrAjQTiCv2H++x6f+EvMr6o2mMwt+SJq8Oy4Qs12kAMqd6AXNt4xdf2MzR8GzOibuBLW/4Q4il696QXwWDCD+SHLUgzmVrEyvvFp4ckX47pSaOkyekxLrO1S5t9kU2EOIj4LkdlMx+XEJzvDY7Nmu1wxYS3/Xqsn09w9DHOZa+sHwTZzAlD/cH0MJJhgzsYYYw4pLj92uJht+mTlacEzQff+XvQu8kqLYlqPsJupOr3xwY6dkwBezAXliGSsM7T0lILRKLkBxY8URvRTj2MOKeRUkgK5EXz1ioH40Mcy99ErZRTnG4c5kqLztzeJEVvF5sYCAopxNbIuUNTur9oFvJrY3XiCDw9TYKYMAXLi3AITp2BEh6VjOGeiE9iuMQw122GlPLIvRTwfLf7njTENDEymMpKGlTeu3y293ziW0mxZOirNm/ROtmcF8E28kGCi5Hb7cLeFXf9WlmKYcZyXCLAH/UvvScvBUVJGXez0bi6NYhVMZU66UrNNHwB0wdlfZQdwKlOu10faMxopG5G03/reHJdmtVH59MpCOi8ave9SMi3NhgoymVJ5Z1XFP3khOcwAjjF3HsTQwlMZYg23/8U1PujFIiL1lz4SEhGS1zKEVz/Y7zjE5+zCgj5uv5LYXfb4QI9D/ePodXuLOuEl0WabZwI/+WRXQxFX3v/r0d9M56sibFX0IFx90rqbmRq/LYQFAmHTzYq1XB0YLKg0DoaHcr3nk7MBiPf2z32Wl9xqzbLAMc2SlTMPNbOzJ5uTeuEOQAyQxVKd6yhOw972LHvBZ9haWwryknE1rm2I+6utdxPIgXe2HLUhaKvWfL2+0Ry9AAp/xvfwNSgM0DpHHe8nD2+DL1x3dEpxnkiK5Xz67Y4UC5bYfAEUUTMle5bGk5+YYKHZ6pTXDUa1B3H2bi1AOa1khckpVlPj30wByiMDtTf1fVtH7TQG3v18pnowANndvqrqnn6hjsIY7yY3TfrEksa7+yLDjAzrgYC0x6YseZDiPj/Kj7U2IGluJQLBTtuKLHS6oJM/LZZLvychvHh4+64OADY4YSQdvpATWFkMUQvvZKR/jE8YiqxbVVvSHza5d71z7DBDpvm2tkROsIKs7+PK+XUVPS08ZspSs/JXmtIKWLES++04nR9BHUklmwuArjtEcj1sLZ5OHJ7/+5O0WjH0oPR8PlDN4GSUukPd99akvemet+FmNzsO8TfZPQovwAERlOo/NCx7w9S9yqdPPYqe51DFB8vk34Uqbg6JDoZMweLBwpPBw8+C2QU1VcXRuXp7tz6a2EmO/a8Rl/Pit5s2mD7W78qzli01pAzFRGJ9QcnF+dIlZs0Zbnmhb4JoYI3omyQx2k7Ngxzp6UAYDpzy7BXGr7KmEiobG+e4hkZzDk4BrjAJi6YpaVDkSYVxUEePfgumCiI01dIDywD5I267F0w0BTABfy/gFtfEMJXTrrkZg5uKUN7Pm6WufSP+5HcYYCzzi5w9xonE+jnZ2gSgldHEzhCymTkgn6TA0TOdjHrHDggLoyweXtZ7QuW9RgVEH7iNfkhIp9T18jc4OKeymRSl0fow8nni3u9rhq8TfBnzLTbP5Vr4zp5fqOE65s0D5/g8SujJOLiqHDM6wzZHku3bxhI7pmwoj/f1xxXceEFYzWgP3NBwcyur6p4smuaPquBi+JIaK1beyKVsMnqZCam2xMn5WDO3x6G8WouJvFuIZyLeX+asCwWWT6FMD/WjyP8njNWzlOUPAD4UOZi6s5cL86HVFgsk8Q/1PgvWsv3nmwK829ofE5FmyM/bAumkG4Inj5GQvLRcFBvDprLUU/MpadI+49ocUhbLBZO+9L6yf+aDBAmrWhWM9feUb3vFP3sgV9DX1KGFdynn5j6eMRkftyxv3PvuAYBKZlP2z2k0S7AgnCPnFragmxO7JKTqFTDVjZ8mCFZwJPv7HFW8wZqrH98KCc9/25akzuMiiVPlV15ndukj11JPknTEumre02qbMxi1CqF30J6T1aEWbj4jHKqVUM+Rc+Py83Vijj8xWLaLulbVOq0o2vUv6RLu5htywsBKcadDFxUVXM0ODqUZM08qO5HRoS0UZHukO/RChpULlhfdNpextQ7WN9nujRwXeMp/pMU8Zh69tiVik9ZJufxRIRNkkf7CKuGnYJeue/w4cxT6Evywzbq5v8yV2OyV/23YJu+H1a9SsOPhSQlUs+PdT5H+O0m1PlWB1aPJJ4VggMpcwstoLJX8IbmnvPZNnqt8byoukuH4qN1LefOHWFhKPa9Tk9tUH1lHpApPj5/pMfmiLYoc3XhC9j21oWc0rhpZMsXwRiiYqJqNrxwtmxP92WDR1cYvppkEurtMeUpkDXDJ3bEYfNKQXhwRW+CUSlJ6l2efFlvhZZM2HaPtuhNLUBsBwYmk/XtRj3q+tmpvpvmwTav7CO3fRw/ZagY4VsvIaOqKymCXk63VM+0n1U6EL1Ct/KyWUA3+fGFmio2WmXGM2r5o6SFoCgjxfiwmlDXj/lpo4TBarcImv6X/i3fSKmA/NPbXrf71mJHvpw0S2twLWCljNLLjLxvTsumXqeh4oyN/rPQ05tGK7yd0EprrADxdymXpavi2y0wuzT/060kck3k2g16iY5HRjBVSGk1dwisbkrKJuDUD580LAxy4P2kGLv60Y6nbi5C0hO/It07QTFBp7Z+UhHpTEpLiisfVVzqTqWN5tS/LKzzuRR8qrh3VggKeQAi05V1oiqeqpNXHiXQkK1jOpsL3a/GMZbXNVrWukpyBJ4FPr2ndrDYi5v7PQytujpL8QFu2mOHYkO9WAAMKYvM8VRY6HAzVdjTEd3Xlkv6PQsOCKy2KmApGZpck/TCwJOrTjX6LStaDUpb2PyKQaSNRaXEqLJA0Ig6vuYdnHbBjV5yD5a6oU69sna+NSHP0LTzRQNW/9+K/3W2UPQB583T1KE/29ayFLyTbjsQhx7QO0tmWnuc1aauApTFg8w8O25kp0VxNI8aMd9MAj9c85MvZg1nFRWeI8x/PgMCFNl0ype7LLkdOAmOOmxVXrycotbMQItlsCZf9oa3XCCJilW7VeaivrVNC/+/boUQkySXpJ0BuRTpP4pX0wpasjMb6YrIcJD1023gCmqntVicVimCjECQMihS9VTCneh4eArU725bzGU3Xalh7+egcHzgMAwqib/wXw29C6viYwHTUlFIdTfBs3Rh6rKpSCyL+yU+7he8+yxc7L2Teey3ewRWuhTCzIyWuRsK6OlN1gNBL82/Hq8ISg1gJ569uTwnNgjhfvzlTF1ecQwagnpkI/WxJWYNMonsrxBpul5z4wz8uA15Zthn0GbkmdHjWI9/76sHDoIexte5r3FXDR675nWVlaaRCbcHaQfOOdUbknpOJUR0T48p1qQQJeyiS9rThERHhG5jInA09VoiulQEdIlJeZT9jp6xhoKqnnqCqEr2BKWo+p8oqGHopqGZvKg2kedlHZ7fjBiqsaQ0YYx4foNQ0CO56ptxfECFJWRccLvyuqZ8uaV+Nr6Wspoq/P+H3yYgrQWzv9/eyxj3+I/BPq4g4OgZS2sJt8iwWuE9qScwHd8aHWz/weCPZVAwwNhFi3O9N+TUyayAIripxeXk6ePsGQBG1Uzj0TCGgO03t3HNTxdMfSBIXxdDVGAAOTXc6XzSy/Q6J16ch2OSD0ZWIvcqOF0zzgAnzfxXe3kggIrBNexA5mx4fcquyok9NfUF72xErReeJsYybHOx/k2DU18EeA1dTl6QHk5YuJz/aNH0wcn0PFjReRXVa1tmq01xd7DAGiw4bNEWZIX9FyIfNT3ieEKI8cEiyPBOe2BJSNkaS4mp2+kT6nO619sra9UIvaxXMfCd0vNu1434ZOw1yM7G9skLkdb00NljpeIsK118gVz5KoPqn2Ti2xiTuQgm8sM+D2Cg6xCH1wIuea/hX7oc2JmvbadtiPxdD1x3jenUul5hX5Uy0z9IeXdJodtMAtgTxjrQ7yKeh1eTPZpZ0aochm3WOpQ/PFNXr/U2t+gdyUxodpJ4pelhXBO05e9NQ4fJePI/ivNJA7MQa4wU/uLz/i5+dS/3yPJ2NWQcfg82Z80t2kFZ/8Bx82N2cuOKf+y33SGWnny7rwA/i1kQZKXcHr+YqMxhXQSO38G9LbWf2VCy++9yt4plcvLwoc7/ply7BRapD+qKEbte6E8i90VJMeW+dI056h1eren3wtERfVET9XBOzhttHkmZKpwPstfvTz+Y6bEu4DsciKFeFduMyLuywP807E+B89PJGlaHe7pp3IkLV5QH+2GRkr97mci0GTHYKqrq6vjt7wHt+jGHeTgP1SDstXHAyONb5JzrrJui0z346/f1qCFYera2uvI8FYWXiRnTAAV+D8utL50krE1yAcvJrPFX9/qudUNNJulouHqDxa5P/19pVhVUbBupsQEBBQUBqUlBakYwMiLSEiJSUh3SW9ERTpBumWLmkQBCnp7u5uNh33+4hzzo1f97nn+vjDefb61rf3mlnvzKw181q2EdTMeXd62qEWq/g61U9OiQjLXs9NtSJM89j0cd54dSC9tL+km5p9Ro54afplvN6fc1xck6n8bdYtqrYVnn/bb2Cl1oZ/vUyeXPY2xD9rL7BOBQIl+q4ac8ncy1oYSuJ0eIxeGsGoaM76h+lh+KqeaTwvWAhPS+lXvlw2FmkSYAYR3kI2CkWQNLOhLxvTE7IfbXULRIUFdoIF4dLavWs1B3/rEdxdwDM/jZc/44/npnGQo9xtYAlQA408mf35L0oCvw044/aNI0/E6lR0vwnj4aApg6ZlojeLj8YtNMMFgLj4dIJB4GFB1bFD8i5lzLtpLA73VpPOrOccWyUR4T8c6tPMzg/FIllPcHpjoisRRHH7dZ4zp58qMQ9f7Rxusp8uKYUPKRwesyY9ckFAFJaAODvgoNyj8RB6jjOD4N7tqgjvSGS5jOs7yg5Te1ceCTe5+Os8zY4i/Vu6RcCGL51kSGrg+J8kUyejR54zDXI439xvmhI3e/j+FcslS1ifWuVK91Wu6452HcruPTQvNLCKjkr+BS7yF5+60NLRgwNvP/9GKEPn0bkzYfkFrb+5WeDhMOax8fNpXvJNwfrgjFoMMgs5B1s1KRP31tT30+Qv+iHTNf/c8pG6ywox/8uxjsfTpzvrcxoHGy1+nX5/7Mbdv7DOdK8qX1A9t899zk57vPo57Qqr3pMYKyM+x3clbf3Dwwx8MQQPiaeAw6J6TQj9IiTwcOY+LOFKL+ssWvDPsvtgptISpjs8LWUqVvW81S0G3qdpPEWdgGgHOb644nRFFKTJdEjteYqGjICAiYj55WkdMti7jqG0md7ij+DeXs6SupfyyubbHwY59C0/QRuBc38avMRHZr/XyU2w0DDvDRU+RI57CuODyMp/w8JBs0K0idzKJ9ox4N8Hcua4keW2sDHuvWytAhPZjAOsmudZ+i8fNxnOB15qTW2lhH7FQfN6CSzlX56qnQfIax51v+BnK9GrfoZNgdnhtTUxDul//K2X5rqXaMsthFcFeI/j0STWNFV9XCK5s7AGUD3kQTed+Rrd/osQJc6PcrZYzQunB3lzvZfNLFdkYjqJObKh3TvDaD+dsxjbmDJWy81TL6LTbXsZqvjuaz7QuYf8FwZkhF7COLOIaOzICAnjV4ssSpXiZ0ccKOcaAlVMDzYYY6eO5M0Y2llS+pBOFI3eTOcwsRa9LXGvDilG4oKUoD3D/4asHuNgz5OM0M3ByBRn6QHAVfjv39kuE48IE5b5aGnFKV6pXWcBpp5rv1RMgD3EY3jdtFemnjhPxaTGQ4daFrCsHPVieENkbMTYbnSx8K8EwvmCs7ThYldEHH5UFeeqBD5ZrJhGrPKQY65Kq31akT6P+1njr7NpDhpE696ZlvzUNHum3zERlty7S+zG44IJ1jbYP86yNZ7kyqe0sAXr5kiipIHhvqB9JJMHlBqz4sGqlfnqYJwyO8ln8jnH03TR//18+Dt+Il56K64KAl9oTAx9X1NE+qFVH26Wvxg22JsXBsDrmo6AC2uoPXuaU0Xtu6JFgvS0tEKpt71tkkeo4LGwmcLDbzSDTUWllPvVrDXbNom5yXBVMR3y+bnpFokNm8F/nOyZflcOK65I/1oqaZ5htAxpsSWpvCTO9FPoV6Br6q5sC5SjKZh9H6mkGB6LL11tHenedyrbPJQ1/+o/bi6zRo/SnOJ88zgJgyv/nPEo5bqOnjZSQvyQwbJNwyYxsp3yH5GEBacRuzlunbx+zJf1/oUpuaHoh1H9Q1/tMmiwPglqE6MN4AMWuOZHNYbggUP7wn1R0+m4tLKE3+Yrz0yBsxhtla2Kzfifc8/Ji47Hogdt+sRGXdK/Wjzrt/A3kBCYEH7fS6JG5cEBQjiMgPDaqzPtopNs8qDDbV3MY2LBX4/tLJ4RF+N2SKgRvESwlpRqMWL5MY88HkcJo6znGepBRhSW1B78eXFwzIfh33v5fvtkP8dNOW6kjWTTwuHFWxFezuP63fBem55n3Uh1xQiXF/3m4fHxeUQvhoOhKWY4hpgAhLKj2K0gezx99gIImk+nz5R0El/NC7jPR//ZTbiqY+6qXcOn4tZGQXuNB+bBmZYNqjiZNIVjx4tXBwfUJYxhqOe7ZlfmzbA/ZbZenb2Xg0xfu/qqw8p71qM2Ui6ILdcw4jPr+75RwFKTkv+laOCg6SBKD+v3IxVcfAw8XG6E8s9sW8+XXT3U2KxZW04z4TNgTA4gPxeirjzxzi1bERqcwNW/5z2BRwTxCKaFiUOekb/ihiCyIQKPCfYDPgv40efZWhbzUVfHuP5+7sGXe9OXm4i6FEnUA9aG7r7SAiVYbYRHJelIjjzhI+wuKFxtKUFjT+VlIA6QbFXt+8h/E+rG8ngOGqsxjFrLqBU1O+hc+foXXXax/2pdPnjov9QcC5HBsglhSH6mJh77FGcCmcz14G1B9KNeBLStz/W8OJmCiJKbqx5Ry2Qug90Vm9O5uBpG/PvkJM/w3VGM3mQaTcAz2a/WsBnVEEEqipcEgaR8QwJmfLyQ5IdcKvcRvV/DtOAfYvvLjOu0PkLhZr2TFcflfj+0Lp/8/kq+kVpeMCaDMcX4i3MWF5niWx2pw4SwOIkV8nXLfXe+lttJkwuL0ivYMu1Qh3u3D9GjR3pHU9w+Xn9O6k2KHfp/ypS4nOyNZQ8cy6Z670N2SC74scvTJOyM4Llqpb6cFlOMrAG+fn500Frjv3g+SexlagJEPA9QH9gYqUez0FrOerHyt1QB42s4cuJ8da3pTE54jd8HvzHDyOgfKrA2ykkesGvCKV/bI2w+vyIunlzveEbDk6DTjGVtwcClSZzbwWBfsP/KdhjTl2xqcEJJQf1VrwXL3mmCsk7g+pu3AtmhJ4pbXKHDLiUDOVjrvmf9kJSfPkuHvhHVmXPlnYoh0SIRpB/6Alf76aqyHC7R6Lxz77WWmDmzkf7QklAt6N7p0xr3VcQuX1Q0VQWPAQ+4DcYrDO/4usCOYx7J4wn5xGVvkDho+FyaudJLvTOWC592pXcLOU5SKibT4yWaFC72btdQkW5LtkD30hWRgMvfOUFZ22d0ChpKkSxG4rpXlxc6WGMjUthKRq5+BWKyDR7R0sUIUvNtUCNdhsGIMj2WJqWqTF1RG6KwX0G81M2EYy/f5G4GEn1023MfXWfItNQKd67pmMGZpvqq5jwwyqOLH1H9u9iWxm5FSgZ7aJZA371VA2O7r8TAXYr3l3dlkemR2xj83U/lr9bzKsZ8M4bWpcbG5t0MH4R5cjAFZNgyNwfan9PnL1nRLNZlfN1KtQ6Yj8T9NiooLgEvRLo+Lx1YsV4DrxOegLfunmJBS+fgGT0prGpk127knNWc2yJuWa/pB+kbGsvCQw8k1EZrorOlztjl72TzjlLpvaOGlhq5LSujvvCOH6wijPOKpjY2Nk1lRTDgpws3qBcXWKurRpVI2G120GduvJyziPBT2n6cFbo228zB2tdmtNHe1pY9UePs4GljSV+mdntlcMa+/FjGNAAj8Xm+umoD4Cn24E0ryN9d0Ru7juvVjPhRMn/PfMHo4DSxz6qs/gZWISkhq3FIpy1FtBsKBWVrXXxcJIlDQMApFNVuC6cLj46W3P69H6iqdTjTG4mueNeUj4PGiThZ7ZChrOnYJPAT2N9fmF4VooLF9yRPn5IGLcIkA1IoQrnVHRcjNP0bLrHqquzWn9tg77lbjJcERjJrSepHaO4G/YuY5KtBTfyK5JUyUVNiN35oaZle7zqD69Y/8Hks9t3ITNUXVeFqJcIS21F0+oAdm1dRh2aGnVUB9STZKMVMpnwQaucn8v67ToO9LSTSAuNcGU82xTgIGpPMK375i5X5pkwhfeu53xCYPrLynD5gQFU+QJkT48gsKcFsrDe5m56c/22i25lPeF8u0wOiz9yilXNGcqppKCOab3Q0WA/qEmUzmHGnS2lx7EvpV0s43w6kTmvzWh1b4RnE4e5RHqS9oOLwGbNflxYJgv79rEsZGfkp4ptGiYnNc0Y+iNl4Wbq73Akq/PcqbttYtbf1vkPPlCtG61nn4lihV4sDHR1deLUuwoUij9evQ/23SoduvLqG2FWndqMk92eFPVdGQ3rDB3IZNrHxvUeCfNWr1nYJ2U3a3caquWLNSeLEhmVPtbgdVbprF0IYOAN/unH4PS9PiMUKWWqzHu5RpkYylLPn3OycEscPgsZd7vX9xSmqvezI/fdBGmMPna7l34MiLhPyybFi/6lza6sLOjSPyL+/kOq22+Uuj0zYj+pfZj5nCdKsXVaQkvWnddvnCZWK5TR7nuWopXbezl4cZfE7+s2ZvXlypAwFw7Bztf1mR+EVXFx36lvoJOq0G3aQQLJpN1Eb3TPenzLv31g4bpCd6nZuz86qe34O0emokRe+ghroeZaQoQc0U9mliRYLpr3G7z6Sl3Uuy/OrsEFsoeOtfLk5URo1539CZJ+rJSFXQEM0kbOm/hDdKhebWMQNR9t+ySnmj4vIgyeegQIeMS2i/dS0/j8F/Un5q9hifAPeJN3/JXlmXEKcS0yPXrzqrbpv8qP8z7T82nQxpZkmq58YnFQsvQall3LZNv8tizqnYiR+P7bg5ZvaII5i39TVlVoN/1Sob3ZG5X3XgiPzD9uLr20V0JvCz5y231PPKacEuMLIscWMfXwlkxBqqnqG3v/pXZlqUXA7QjiAf62hxBYP94sp70gSna4RVZ6+ZJnLyktBPymTj/3+8eEeKWziYtxocG2QQ3Ytlq14xR+jO6PcLV4Xubjvcy2ujBbjmKg72QQ9o+eHt1VBSkQkLhareCUkrxMnuGJ+8nN5a15Z0gTwOXxWOd22b1uBvECd7neWMlJ7j1qnwcr/JKZ0fWFRYnbLbCBXVW9ydhcR9vgDckKFOcHkzgCjk/5YkO3JCe0Qt6Ypx6O5iD2HvFSj6bwMhmaNV+7+nJdTFR7558t/Vs9nu79VGBx2Yp9gCbVAG4xCd5c/MBNPxtZOD3+bqnZDfh+uTH1pOM47Wa42klvdV9W3YFIA32v5seaDZXoW0OWY/x29vY2uq4Jeiag/T4ko/tm8uLxKZGYFe8yCN5JcK1Egdb6du7Fxmvmxse6z+DE6j+Z4J2xS+EH37Aq9SzYtYnple4DfFbSNgSnj9Ty+SEUiLuLH9mneqPzfBzpVGzqU01/rP7b6F2TJrBcjvsdu36wVJGj4Kym7tZO4rJVOyBWUG+M3HvFal3bqdUtZGDUFuh0B//y7pbkCTK4nRN1nOh1povzQuRIDCpO67m5BZ3WcAHHhWTtPfo9a501e3jzpR6KEkiJlJI7oAo6dhfLTdsQTl+i6gkkLIyjDuP7oh4zHAy9MeNIblRJxxaBdDlKUOSX8exRE7oeW1oTK7rUi/KsZek6HiUfieF7P5yne0sWjCcfipOQj37v8R2wjgRv0KSp8sJlJW9Zxhc1tiewfC7v4BQ1148Nxs47ZjXK6ADLZhb+tMB1dfdpHWcyyaPWiMbaGq33ptue7LagHw1Kv51WFiT8YpdCC7E4BEymifoYzdR6fcONWg8O0fziIY5EUxuTtZl3fIDIMNg/1Nan+VXj3Lquk4ruvVJv/It7LfSJa0V2mjHh7IteTxajaWBON8oMBtdZlaRn/ivudtRlFQd0IccgB2+iEjhgTlb2uEztqgyuzI8ZgCdiofsthWnq648mnvJGPDzeKIqiQ5vzL4ow4BbtHUOlTO3S9dZuHs6p83xVx7Hn9KD6UxK0fnfRWd9mVk4uaJFrGCGEXIUcqlVWyq4dfaBqTN/xijjRqvt9LT7IxpoAhDd+RrSOFyBjrXZdRMhqJN0zfm9J5mPkOCDh4xDxMC00PjrIi9hc7aV9NmyWclnSYmPQXu5pEY1uWNY1p5W3aQwXsmppp5XN3n8MwHKzP020/kSawvS9dv9SiTw0IUZI11uL8c1nmwFJq2Psq/P777IxanQ8OXv6K7b7zRksd0ew6jWgO+xMhuTmG3bNj+77hPImU6YKGw7ka4p4o79Ory5z9uR9/2Mq1Nj2RU9soxGgXhyjIZ2LfNJVHI/tjeoesVsxy8EIeEjtx33BG4bQi9yB8X1hpcLhhEvpODSS/jd/IAj1vaamANIaC53SH9KaDVQe8OSxK9p6/7dt/AyGFSQRCXW4I2L6rAKnDa9OVWe6bzlUqXCHUp6Gp2NM3DW8UkZD/5dXq/OkfrayKviPIyWVGV4UvslelsfELFID0pREYsIzVCLws5oTISIwzaiYirJgRqSNKOqR2C5AC7fGHGwpTk9iy0b3J/Xep7UAOp0t8RjPprSKW6geXUL0mQQtw7ZBzUjfU0WuSdNLsPW5kizGy4SD5rLbwDWrO5J4GHdVfnvuWbtNh8zSd7Teu4QI7G3951QA0jwwawp3sdyIJN9RjJ1EJshXsWwOS0NFdKFXnn6TSzixuHjt1Qx6Jo2dLP3ub2FzmIjdwiQRME2m6PcWixL5thr2YOY0hKIprhgZgOsinQYtcFQihXJaoMLNwPLA5Kno/MW1kcF6H9lhOgZ1tolZMlrIBNR4XqsqbK1eNmQ612skj7crHytEUGV/KEk+RbhPKaXnXpVa233ooGW3iokOYcaiccWDIbfkGJJpdeqqf3NvApiAga10IttOnzNaZmSzEhRZtFLfD1nd5Rn1G1DCP3ykwVlbp6eUNDRLVov5HfHQfPC97/UOb6o70gAXMW6No1ZBuyfKCwTMgCa/7Yl9u25EvHgmhJrk5pnD8/51ihv75cxxiNRUVCaWfDD6Wbm613KkampoJr7+qRuG6LiwsfF8ocWJvhaVVMQ71NeJiCrUBMDQV9OPr5F94enRJhdkNiVElhAqkeHQvKRpcC7XuWelt/n4ksApNLuoGLzMNDyxNHWs/G1HgBzA/4Ju6osKsqAWSCEJmMHIWefs2JPr9L6mci6t0sJyGlmvO8YLYmNyampr6rZNTlauCmAyUjoaW1k+vgu36hCMybKylMI3NNPymBLupqYkjE2yNE+vzf39X+EYt2iHDkOGT961SEZt8xiUk2Jad0VBBI1eEPnXU+5rkw3VOFuU/ST5Kli03HUXb4rpDXGTCntxcNqKaHlU0UUJWE76c3lFI/E+NrsDfYGDJhaQw4XfMpphowEZGEfbCuH3gKbDvYeEEa553/B9grQnmFzT82951ISkIKYSOdeCOafEFqFMC5IfMd2Qw4cDeVtC5BRsI4kfw8ucFGoX6HdEGHaDRtEb5Ox6PvzBYNOb3ZYooPLO5+7J/6/dzpNHbJTbnHWuxSbifvhVRE4jRE1LQiMOA+X2q/UyHambKI/mdqmwsjQI83JryHiEH65B5foKnW5I8htedc3x4Z1SM7E7cZcCjRBpE9r+vJAksXYcoee7op5IkeRsPgdFhBxKuxoYCn9iSSkrYP2lXW87dl2F7CiVIa4V9NugfoY3oXZlFUr8pOS70v75gf1e0mEYBzAu+HUyGT2D9NI97V+a6/G9K/EH2RZIPWzOgDd2QL14bzxJJlzVoOzzb8+DvwZJtUorF/MfMx484U4N7RwxA4Q/i4hPLijsmPQBVIRuod6r0ABG6LqnhjrtGCBVydsdR8hrsvKZ4sWJ519Pfg1AHveNLeigB6AANk1cd5VYHtpBCbFC1mYwsyWbVvxrqcCuAgTN4wAdxFA8EQOJaKWn01TD3aeFNuT81HuxgbbOHJdUMunjGDC7aaywPHeOIyNdnvg+CzmS/k5TP0qHB6p7h1pHCRMS+owjVOSLev38fNaL5nf1baqbLBcEuYoQA4Hs8k6mLABI0PvsNHJ76VnYq1b7cZg7rOFvhGOSMzMzvSnlq+lml8Kzehvv4x0G4VhK+ivjrMxSdvv/YFLAAFCUDuRXXZ8AC9J5GJSLlwxZZU8crOb/WdLnaxBYKTf0Fnih5J0WP4TKhWmuacqFpR46ZopRdyZs6cNnCvxPFv37x8+e5vk31xv26y8ogYM5Rxb4xort6mTSYGoLx6ba8bqzMTHp1dbV0+wI6bS0lIsNZNzGhbsVZFxdHUnRdYadZ43SionGWLYa9PRBjCgM2qbecSl9gT75GdSMACGYAGtETJ91kxUTKU5QnVC3436t4B87sc5OWmD44dU52nr1+o8iur6fXYx80yddT263ah+i6/GyrNHNXEpdSKTukfyQqPt7sE5Jpb8rXqILSPMODkYC8Ak0QjTqtknm8i92vK+cdrNNHIRwk4CklcZKGPQHSjPPBR1LkL8FC/ZyuK3HFpxVM84EZB11hHwgwbphwgierUdhAht6ndX8afwNZfZScL3zM67pDIjMAsgUZEEOd9hD6jrPifvg6m5/vtf/XaJA1GexA+EKB+QLx1r7iQJwo/p384Fb+8gWwXeGHKAQIt0bXBQBNnepUzx0LDzADKSwJDcn/DggMgBBihtHC8A7ak0AkkfdCELtDkmLAiC0juUruuCvlcSx3J1r7FFRUDGn0Q9p2XXr4v2al7ZAM9yloasry4/9KM4QXCvbxnKZXfdZ0YrbrxN/u1gIRw1EwtDFy3Sbu5vqubFf97BFYJbb9n5jReJZgm9gzYhDvfQMcccvPaGjESmJNbI1Git4+QlHXcoorxsQHaUw8cOt3+S/6KBZnFE1H8RP2qpMPn+HTygbTIM2FpdAQ4lwDfkDACPQbNdwnIPSu2rkCBe4DB7sTXKzpy0hSw8PxWLny4mjKiJYUBAAwGxuR2uxfue6vgfvc9tc4sJXJFvUnNJSCRE6XrDQyKGmjLfHXB9IaiTtoMMnkO2YQL5Cr9i8ilsMtqsOApYVd3FFeUQSDsEJgW4F5q4oViDb5HQcuGhWw6sgopMy3gA8MHUG94ylBBKM8D2Hv9lsIgmFBDu8ihu9CYIyn3TPLdgtBFZ+r3l2G4QVglheE9P9TqSIao23nvAzHk+7K9+doqVXELmeJCBn7dPIU+GnDa7vhxZGbLmqHQJii0zqS1w1nJPQhk+MJyNcON8JMpOrI45YcVXWTXKQqU+VNY4tUFxn/tBoylqHPPKg7SruYFBE2YsAa/E/FfjTEajeuGfCoj6Z9lmkjqSSLiAx4ApizMLVepLMRjX+GWg3po5Y0wQTTq7jxIfg5lNr5I+dyDfU+I1KEKltgq1PuOcajJaJxIBBYmbtQBBtHKBt2B/8Z6R3H0QqIG8It2Z4afaEMYA7AfBm1hR2z/vXxh1VH0XncdktNZKNekFA/UlAhceOXjRVFrJcX8kOWqeR0tlyYA1kWCaoYfcN9hpt6xSVcH7AWN8qONfQTqd4S6AIiXUPHmi8NF9Oqrumm3Y9/zOD6HCxk9i4uWD1bzAvI+4ZEXtB3oZ7qLTcm2wReTthgLY9IORb7f+NJDOI4S1BszhrZY22vSiOFVG+8QIEdsePdaecFqB2VulnmW+28AEGf/M5gkMFaO0RPrKg7XjTAzyD9N/mZ/3il5WM8vO+TzIyM7Mt6UehLy8t+uGNgh8Pgj33Z5Zu47fOW0SrJhYIGfzcX/qrLLiPlf2lx4BesuK1pk02g/C81bf+nzrr2SkjafgW4e3cvaT/WwdNbeNote8E7KV9f345zovdqm1cs4A7mOfoSgV1Xy/XafrQyNDT0uYCARl+Zo8IbWdknrZb69lQcHO8ZaW86F0D+hXihm2Bgp0lirDd4bfbUpd6G4mY7K9HRq/HImHJ0pDuVtMIIYhTxW5ToU3uv07k1XAe2Yzl62rxtPdN4nQwBb5NuBoUll8AFk5gSCThY78opVN3CRIYgcYT3/zCyAsZ/xHldcaceBcAkJHySoXcESAjg3ejT75YItw+kAcoMfiTieheMI4HB+PX15c0DEgAaPKX+9l9CS3X+4Re08sHhc3M6Vl8wzBMX0/ZHi/0HoAmJQB7bq9vZyyL5RdcqeMRgNYRG6qPGI4FIM/6Ciuv6YWofe0ZnkC2PaCnGVYIvo+K6XheMnx6gQpclVW+jJ+jPtJavS071TiPQVZBD6mLW5d4N9pZWQoA9t3uPBwrVbPSTC1QgVKrrklicm5trZPuNxJ3szaLtQhFib5wG3fVIlXzT0JFOUlK0iPQqRVot0XUWCCJGmp2ut53A7wLrZm4jA4kc6ZuuseujBf2Kl0wfk3vn9UxLO6gxP7sPrlnRlwFOQ049DXS3ETIKY46/qv9UqSD85/8Ywaghe7zgxcwAHcoYYDFyDSCj8xbxyWmGZvm/nz4taQhel1r70IEl2hWokWBXcuhY9nb4VzN5jgztRv9K37FeK/dsWHr7XudWxgCnuTIDZcQ1Vcax3qqRYSKRWtC5ps+kAAZc6A7Vv4BJtDAxEcGdg30gDtjDN12/yLvQ3hBE75Jkj4+3CuuxBBRsyF1mSPh/MYV2Yt2WBp/hWZfA4T15XOCjOfCj3WoIxIvTw2VtvOx+qfFkJJZJfyjvISHe2FadE2BDLUByMUNg/BzR45T2xXq6s6For6VNtmECs1N7m+hB10Q8vACNCUgGvUgRlEcckb0Bj4OmckKrdgiNnzrNbRcO8uvsfi7KFo4mhIIPfCUXtBumppeooOdDpvUA4JbIeuCtoilV/x7T1fTUhdeIY4QViaCV+ae4jjfiITNiaZN4jJl+przffUZODWifiz6vsHz23Q/hHhfFVyX0tVcb35n9ezT8K2ySJShpFPZQH+G8xXlcck/Di0ZvqesNzQz31LTMNFv0z5fb21w1GVxHF2Q/Jn7uQ494upd5jjol+2oNnHYELlg1r0qV21K33dNyu7qiOuGl1ptjpb9qa926F9vb2/X19RcT3FvkcISSqhHFzSerSosrKgzvm5qahldW8jCI4R+dXfwqLNRuysvPf+LgVnUSz22ls+Bo+zdCFmmgVbrL0kQhYYG9DgOfOfxdhizPO4ZSpxytOAA+8ypX56iYrSQ6JmuzyeLzfyIJYvE+2LGI6go30DyOeLVFqOV42LI4UZzIJHgKF2w/ct/z2yvNSKxYPqt2Tyl9imFAxUEt0BcXWqHnxmUxRZKA8hkkbMeTTCirmeTzoVLWiXJhDnWuWJPeYvBpWFr6xGO/MaLZ5+qewD2wZuZEYvmB7xJdezq9nushAbaHeFQCswGqlnPLCX4C3fTFibP0VVcEm2Cne1/xsHciJazuA4Zly2cT4lrPuhF0Q++C3VhBLcLzYwXVH9FO8xWb5HHudscCbkIQBORDXG2kwkSsaRHtjnu1DqflFufkkFMTClfXeVxykn3VlwLIOASoCJyQ+Z57SzAIItpTmJA+IEbUS3pfS3XcaIUIBN7ML1CRccD56DEKETIxZXlvpGZh2IM6blsPnqfXz8qrwh4I6Vv2PbuWXmAhmkKakzE6EBDA0R5tiALxSEkttLy6JiYmxz0PacWG6c2NjdvQ372x2ZEcFxJI4Mpq7SyuqRFwwSskI17HLP+5dAE/lnSgwnJiz+ObmFo7dbjnPZRmle/plHOVMcXqcEY//7McVQD9s4VgqHp2q53EXpqA1FzoyWrC8hRlKP9aHlSaY8vRJ35I8RfZhFP48uHm+GEBoU/lHlWDXebgh0n3hc0phvEx+KdPnzYMaNebV6AWqjK9D/4F4MKIKRm5aFpTRYPj7s0sKZBFVxPjRVdjlToxcDmYZ9MG2NhQL5lrvQ/SGYu7l2TiEEQVZU6VLxJKu/n2UVSHHJf4kWb4wts3eVM09mpyW06WkzEcsRPKj/bTflL9VLLoL4aMje0qQzrRFx/fLA3W/UKEFfliultRDFhGynoGcP0hiKSPtJF6kj+9o78RAyDEMLEvyjfrz0tVh5Ik0vlT8UYkAjQXgFxwPQ7NVgKYxTtGJutGZAeUQXRf8yE4DmclGJgFiyNW4UY0RB6BsD/6TXGttp5nwCykxpzptyLaQei7GLxUkWcTRUVFrPwDzT7v/J1545fKysqW0uiMyEitZxMWlpb66tmJzMmaKrCMLOxiIhnHsDQKkN5MhDRXYqmJmFsAm+qwpNyIefhZa9ihhEm3Bq53dKJuyMJbi7B+SJg6oKJE3aT1ye5KHfbQM6m5JM+EmO6ByY5NN1tSY7iS9ylVBkJZ6XpnV1ekN5lTZnzfurLRKxsbGyV/5iHvjx8/1vv2y3BHJxKHLfRS9DwwcXJySuOgC1qKikNZ4R1Yy5M1jUtMLJ+wtSJa5qVaaiWSknssbW2EtMpL5dhHRm3Vej7wwWyHxI9KrhNQ2Ck/D0oVtNKbgbmygGpJbX7rvndM4rp8TvmqEcxs85gBYpfocbNRWoBls82dDfe7FmGAVIiFc7O+SeD6csaH3qyvNlId0fU/hJI8ATXwjthKiV2LsAcQE8itCaADJqDdT0d4LQIqyEW41fFzQMeeMhl6NyLwEpS7lwjJefNViKDyrTiObfUg3Nv6SL6+vm4eM2tTq6ysnFfBHpRQuUHWN79knDvJj2hJ2jdVhVq3croglMXeZSeBhWR9ln2ZViNjU4klmVAi3urB2noecxZ8o5UGWKi6j4a1u6XgcltyCGC7qKQoacsH1Q2qGE7bajXvmXiXhQ3ouFPzNORQaRmZmXm3Bc7vsfBgSnyt8EYw0tk691pspnIsQQttgE2ot6UGiFkTBK8nyYgFTGWLt34UMopbD46pduwTEZ3gSFsf/ODgBn4Fo7gPYZwOp9bzE+smh2WRGk7YctsM+ZVbIcAT5glqNTLYSMZnlYsFEKqm2+XUBZdzyLJN9FoUAiQV1FvzlwfMHzVWtvVGLETQfnADUj3CoH2b8lJTXot1KDCOW+BDvQfos3A13P9ahBBDhiC3WxAP2IL1P9+x3YjAS5DuXpKUYTBa7f2LivWXzFwogQDXHzGZmA9sglw/s9Id7YayiOe78qL9iJc1E5eM5h6hdr7thloAOS1earKIblzIQvZJpWXNe2eFd0TzammOisoBUzT7/RMT0HTH0AoqMXgFBzVWW7KzZ2x7nuT2dj75YD+nnNx92x2fl7Y1xKlHciDUNWx2sbSmEghXP7C2qSzd2Eonjj2QOAmNmaikdiTC4Ipa6+W1dlonMjTf+h3jFKpeCsUvHQpmfITatGKzPkLreHDZdcrzxiLshA7cgNhhGWL5DR1y8FLSbFFiiOvPuNkuuWET7j8Pe+Fm+we2Tn0dHf46atiPt5tXOu1+hRBnzJNOEnysZJUk41kRYxMQmByTxa+qtO2LeqJmTZQE7hpemfb42I0+bq4oEnWb8oJscZP3fyQqxAUyKtW9qeTGUqmYoxOh4lXCvE8kE2JqZDZKLkmQNBInGh5HAUYHPBAT+w2cpopr5bKAxaCk1ORqOVbWPl/0dtsUPdFGejG0+dH6RiH+AHTWGao73uJjsxogqpR7t956JBAgU60WNW62JYEJIOpOktHcujoQBv2bIxxu5tbOBczB3875BwkiKtrGeYHLxcLZt+5HPVcvz6v/aAo+hVmABuQADCp8hPMZdk8TetG516AsqJh/xNXebpDAIue2MGu9jJDvnUUJExIAXgXPbfWqG2fuETyUcNztGLV2cStz2djXovdswY+szsD4vZn1WQQZ58cLD8ooZo2xShvv6ZdvVLToDL4gH4z0pqenn0Gx1af7P7vstFwpoxTsk1/Irl7tji/uLbbrby87OTpW5ubmYhYsg546aaLKzmTnn41FbkGBM0KBVk1Z2acW6v1qJ7hxBL/nsoXJPvueetvLzyjl8cvnJ/v2gKs6PHN1NkeXmHvk3civmrOWY03KzpxwESlL6tg/7Id/9RJzckx9+qR90MZzm3yU5yho8Kx2sjWMdvts++gsSlJGMojdG9DldAAxl32+1p9arVGLU83QDDOLIxppa2r71mWKSu+c2oXaUG4gVttY5N9Z9hHJm1objqrnQUcdKjz6UmZ//hgSK4RMBvFIZ4EK8scZ7gj/WYbP2SQyCtbsSUVHxzMExX/WqRA/Hv4zTAbz/UGBiPa7H7HGwQlxQRyCoKW9w1cLRH7yCacT7xb7gwTVL51PDRa3/1zpR6/VVLt8XEW6JH4EIHc87owchHfYGj0BGYFUJG83qwCqfHngHdU2tKTkuPFSi7kgD97uHCj76usAoN/UJ8BoMaaOM8oXOh7Y7uetKSwogLHMoO3QQgKfDN/4yBfPAEygSQsJv3EMMMo6lLreObUbG6R5DdigF3so2w1qWwJGxq2mfuMz6rwAYPoRLKaAcoMsgQDoiPuZ3JjnDxyPqcNDe0J+kxF6Wp+3X6iYfXrnFUwG0ipXPRhTuwZ//+lydhDHHeeUHa5mdbIol9uCd573c5fTaBTYLfygtXlHRtvykcTYlVkxbD0+N7mhsSJx3GrbjE3wpNKV900CUWVf6XKXzPDw8QS93P6xBTxn03A8Q3NvSdKfJqGRmlUmv0ylE962iDn5OvSkkrUTvnUeUxUXE6Mu+yasJVorj1+GtGvCfni+kZqDY3v6N+e+wApHJ0e+o30sueLZ1jZ5c32RxD673WbCvvBV3JL1mtHSTDOVjQmm8QJvxzlRGGeBDFOnSeyPxuQsA0VHJ8+3gWTwBNKYxsxj8Hd25YrJtKXK1GXa0AU9gcWKL5n75zvohk7RuMrsPRdryKNjlTNRYG7ipWxNT/XUAIAkj4OqQM/Nt3c7GPDjQUvKgLs+J/Ec+GDVGiNhLLXHvt+f89L4QOnoIisrazbqZTlgbVyQQPr/ZgUCmx63T6LtxlnVPdRGmjG0YroZiUtRh6L9ikiS8Aa4uIEvQMJ0C2var5AvsXfupX+6pAjRhwB/JEVlX+e/0v7yPwBQSwcIJbVJvY1tAABOdAAAUEsBAhQAFAAICAgArLb+TAatzOTLBAAADyYAABcAAAAAAAAAAAAAAAAAAAAAAGdlb2dlYnJhX2RlZmF1bHRzMmQueG1sUEsBAhQAFAAICAgArLb+TMQHrm/gAgAAwgwAABcAAAAAAAAAAAAAAAAAEAUAAGdlb2dlYnJhX2RlZmF1bHRzM2QueG1sUEsBAhQAFAAICAgArLb+TEXM3l0aAAAAGAAAABYAAAAAAAAAAAAAAAAANQgAAGdlb2dlYnJhX2phdmFzY3JpcHQuanNQSwECFAAUAAgICACstv5M9VjjHEoKAACjKgAADAAAAAAAAAAAAAAAAACTCAAAZ2VvZ2VicmEueG1sUEsBAhQAFAAICAgArLb+TCW1Sb2NbQAATnQAABYAAAAAAAAAAAAAAAAAFxMAAGdlb2dlYnJhX3RodW1ibmFpbC5wbmdQSwUGAAAAAAUABQBMAQAA6IAAAAAA"
};
$result;
}

ENDDOCUMENT();        
