##DESCRIPTION
##  Algebra problem: true or false for inequality 
##ENDDESCRIPTION

##KEYWORDS('trigonometry', 'trig functions', 'trigonometric functions','fraction', 'analytic trigonometry', 'angle', 'angles', 'radians', 'magnitude', 'direction', 'vectors')

## DBsubject('Trigonometry')
## DBchapter('Polar coordinates & vectors')
## DBsection('Polar and rectangular coordinates')
## Date('7/31/2018')
## Author('K. Andrew Parker')
## Institution('City Tech')

########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "AppletObjects.pl",
   "PGML.pl",
   "PGgraphmacros.pl",
   "parserMultiAnswer.pl"
);

# Print problem number and point value (weight) for the problem
TEXT(beginproblem());

# Show which answers are correct and which ones are incorrect
$showPartialCorrectAnswers = 1;

##############################################################
#
#  Setup
#
#

Context("Vector");

$answer = Vector("<2,2>");

##############################################
#  Begin trig functions in degrees

Context()->functions->remove("cos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cos {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub cos {Parser::Function->call('cos',@_)} 
#  Add the new functions to the Context
Context()->functions->add( cos => {class => 'NewFunc', TeX => '\cos'}, );


Context()->functions->remove("sin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sin {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
# if uncommented, this next line will generate error messages
#sub sin {Parser::Function->call('sin',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sin => {class => 'NewFunc', TeX => '\sin'}, );


Context()->functions->remove("tan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub tan {
  shift; my $x = shift;
  return CORE::sin($x*3.14159265358979/180)/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub tan {Parser::Function->call('tan',@_)}
#  Add the new functions to the Context
Context()->functions->add( tan => {class => 'NewFunc', TeX => '\tan'}, );


Context()->functions->remove("sec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub sec {
  shift; my $x = shift;
  return 1.0/CORE::cos($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub sec {Parser::Function->call('sec',@_)} 
#  Add the new functions to the Context
Context()->functions->add( sec => {class => 'NewFunc', TeX => '\sec'}, );


Context()->functions->remove("csc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub csc {
  shift; my $x = shift;
  return 1.0/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub csc {Parser::Function->call('csc',@_)} 
#  Add the new functions to the Context
Context()->functions->add( csc => {class => 'NewFunc', TeX => '\csc'}, );


Context()->functions->remove("cot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub cot {
  shift; my $x = shift;
  return CORE::cos($x*3.14159265358979/180)/CORE::sin($x*3.14159265358979/180);
}
package main;
# Make it work on formulas as well as numbers
sub cot {Parser::Function->call('cot',@_)}
#  Add the new functions to the Context
Context()->functions->add( cot => {class => 'NewFunc', TeX => '\cot'}, );



#sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))}
#sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])}
#sub atan {CORE::atan2($_[1],1)}
#sub acot {CORE::atan2(1,$_[1])}
#sub asec {acos($_[0],1.0/$_[1])}
#sub acsc {asin($_[0],1.0/$_[1])}


Context()->functions->remove("acos");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acos {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acos {Parser::Function->call('acos',@_)}
#  Add the new functions to the Context
Context()->functions->add( acos => {class => 'NewFunc', TeX => '\arccos'}, );


Context()->functions->remove("asin");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asin {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub asin {Parser::Function->call('asin',@_)}
#  Add the new functions to the Context
Context()->functions->add( asin => {class => 'NewFunc', TeX => '\arcsin'}, );


Context()->functions->remove("atan");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub atan {CORE::atan2($_[1],1)*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub atan {Parser::Function->call('atan',@_)}
#  Add the new functions to the Context
Context()->functions->add( atan => {class => 'NewFunc', TeX => '\arctan'}, );


Context()->functions->remove("asec");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub asec {CORE::atan2($_[1],CORE::sqrt(1-$_[1]*$_[1]))*180/3.14159265358979}
#sub asec {acos($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub asec {Parser::Function->call('asec',@_)}
#  Add the new functions to the Context
Context()->functions->add( asec => {class => 'NewFunc', TeX => '\arcsec'}, );


Context()->functions->remove("acsc");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acsc {CORE::atan2(CORE::sqrt(1-$_[1]*$_[1]),$_[1])*180/3.14159265358979}
#sub acsc {asin($_[0],1.0/$_[1])}
package main;
# Make it work on formulas as well as numbers
sub acsc {Parser::Function->call('acsc',@_)}
#  Add the new functions to the Context
Context()->functions->add( acsc => {class => 'NewFunc', TeX => '\arccsc'}, );


Context()->functions->remove("acot");
package NewFunc;
# this next line makes the function a 
# function from reals to reals
our @ISA = qw(Parser::Function::numeric);
sub acot {CORE::atan2(1,$_[1])*180/3.14159265358979}
package main;
# Make it work on formulas as well as numbers
sub acot {Parser::Function->call('acot',@_)}
#  Add the new functions to the Context
Context()->functions->add( acot => {class => 'NewFunc', TeX => '\arccot'}, );


#  End trig functions in degrees
###################################################

###################################################
#  Make a subclass of MultiAnswer
#

package myMultiAnswer;
our @ISA = ('parser::MultiAnswer');

sub new {
  my $self = shift;
  my $ma = $self->SUPER::new(@_);
  $ma->{part} = 1;
  $ma->{answerName} = 'answerBox';
  $ma->{answerNames}{0} = $ma->{answerName};
  $ma->{id} = $MultiAnswer::answerPrefix.$ma->{answerName};
  $ma->{singleResult} = 1;
  $ma->{namedRules}  = 1;
  return $ma;
}

sub ANS_NAME {
  my $self = shift;
  my $i = shift;
  $self->{answerNames}{$i} = ($i == 0 ? $self->{answerName} : $self->{id}.'_'.$i);
  return $self->{answerNames}{$i};
}

package main;

#  End subclass of MultiAnswer
###################################################

$multiAns = myMultiAnswer->new(List(1,2), $answer)->with(
  format => '%2$s',
  tex_format => '%2$s',
  checkTypes=>0,
  allowBlankAnswers=>1,
  checker => sub {
      my ( $correct, $student, $self, $ans_hash ) = @_;
      my ( $ggbStu, $vecStu ) = @{$student};
      my ( $r, $theta ) = split(/,/,$ggbStu);
#      $self->{ans}[0]->{preview_text_string} = "hidden";
#      $self->{ans}[0]->{preview_latex_string} = "hidden";
      return [0,0] if $self->{ans}[0]->{isPreview};
      Value::Error("Very clever, but you must place \(\vec{v}\) in a quadrant to receive credit.") if ( $vecStu.Vector(1,0) == 0 || $vecStu.Vector(0,1) == 0 );
      my $vecGGB = Vector("<$r*cos($theta), $r*sin($theta)>");
      Value->Error("Your vector doesn't match the image.") if $vecGGB != $vecStu;
      return [1,1];
  }
);

######################################
# Create  link to applet:
###################################
# You can name your applet (anything reasonable :-) )
# Adjust the height and width as desired
# Paste the geogebra parameters in at the bottom of the page just above the
# command end command
# so that they don't get in the way
###################################


$appletName = "ggbfoo";

$applet =  GeogebraWebApplet(
   code       => "geogebra.GeoGebraApplet",
   archive    => "",
   codebase                 => "",
   appletName              => $appletName,
   appletId                    => $appletName,
   submitActionAlias     => 'getAppletValues',  # default actionAlias
   initializeActionAlias   => 'configure',  # default actionAlias
   setStateAlias             => 'setXML',
   getStateAlias             => 'getXML',
   setConfigAlias           => '',
   getConfigAlias           => '',
   returnFieldName       => '',
   width                 => 613,              # may want to modify width
   height                => 525,              # may want to modify height
  # mayscript             => "true",
   debugMode             => 0,   # set debugMode to 0 for no debug
                                              # to 1 to make xml representation visible
                                              # to 2 to add alerts detailing progression
                                              # through the applet
   onInit                => 'ggbOnInit',
   type                  => 'geogebraweb',
   submitActionScript    => qq{ getQE('answerBox').value = getAppletValues() },
   #submitActionScript => '',
   selfLoading           => 1,
   params     => GEOGEBRA_PARAMS(),# paste parameters in

);


###############################################
#
#  PNG picture of GeoGebra applet for hardcopy

 $im = init_graph(-4,-4,4,4,axes=>[0,0],grid=>[8,8],size=>[400,400]);
 $im->moveTo(0,0);
 $im->arrowTo(2,2,'blue',2);

##################################
# Setup GeogebraWebApplet --
###################################

###############################################
#
#  For persistence of GeoGebra applet state 
#  when answers are submitted or page is refreshed.
#  Uses JavaScript interface to GeoGebra.

HEADER_TEXT(
qq! 
<script language="javascript">
   function ggbOnInit(param) {
       if (param == "$appletName") {
           applet_loaded(param,1);  // report that applet is ready. 
           ww_applet_list[param].safe_applet_initialize(2);
       }
   } 
</script>
<noscript>For this question to work properly, 
you must enable JavaScript, or use another browser 
(or computer) on which JavaScript works.</noscript>
!);


#######################################################
#
#  For setting initial values in GeoGebra from WeBWorK
#  Uses JavaScript interface to GeoGebra.

TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );
<script language="javascript">
function setAppletCoefficients() {
   var applet=getApplet("$appletName");   // alert("Updating coefficients");
   if (typeof(applet)!="undefined") {
     if (typeof(applet.setValue)!="undefined") {
     
     ////////////////////////////////////////////////////////////////
     //  List the values of the parameters to be set in the applet here
     
          // Initialize anything?
          applet.setFixed('text1',true,false);
          applet.setGridVisible(false);
          applet.setRounding(3);
     } else {
          setTimeout("setAppletCoefficients()", 1000);
     }
   } else {
      setTimeout("setAppletCoefficients()",1000);
   }
}
ww_applet_list["$appletName"].setConfig = function() {setAppletCoefficients()};
</script>
END_SCRIPT

HEADER_TEXT( MODES(TeX=>'', HTML=><<END_SCRIPT ) );

<script language="javascript">
function getAppletValues() {
    console.log("getAppletValues() entered");
	// get a handle for the applet object
	var applet = $appletName;
	var mag = applet.getValue("mag");
	var theta = applet.getValue("theta");

    return( mag + ',' + theta );	
}
 </script>
END_SCRIPT

@completed = split(/:/,$main::inputs_ref->{_data});

###################################
#insert applet into body  -- this does not need to be changed
# except to adjust the insertion of the reinitialize_button or
# a hidden AnswerBox
###################################
 TEXT($PAR,  MODES(TeX=>$im,
                      HTML=>$applet->insertAll(
                                debug =>0,
                                reinitialize_button => 1,
                                includeAnswerBox=>1,
 )));

##############################################################
#
#  Text
#
#

BEGIN_PGML

[$PAR]*

Using [`\lVert \vec{v} \rVert`] and [`\theta`], give the vector in [`\left<x,y\right>`]-form.

[`\vec{v}`] = [@$multiAns->ans_rule(30)@]*

You must complete this problem once for each quadrant:

    [`[@ ($completed[1]==1)?'\boxtimes':'\square'@]*`] Quadrant I

    [`[@ ($completed[2]==1)?'\boxtimes':'\square'@]*`] Quadrant II

    [`[@ ($completed[3]==1)?'\boxtimes':'\square'@]*`] Quadrant III

    [`[@ ($completed[4]==1)?'\boxtimes':'\square'@]*`] Quadrant IV

END_PGML

##############################################################
#
#  Answers
#
#

NAMED_ANS($multiAns->cmp);

$needed = 4;

install_problem_grader(sub {
  my ($answers, $state) = (@_);
  my $result = {
    score => 0,
    errors => '',
    type => 'persistent grader',
    msg => '',
  };

  my $name = $main::PG->{flags}{ANSWER_ENTRY_ORDER}[0];
  my $ans = $answers->{$name};
  my $score = $answers->{$name}{score};
  $result->{score} = $score;
  my $data = $main::inputs_ref->{_data} // "0:0:0:0:0";
  my ($correct,@quadrants) = split(/:/,$data);
  my ($MAG,$THETA) = split(/,/,$ans->{original_student_ans});
  if ($main::inputs_ref->{submitAnswers}) {
    if ($score == 1) {
      # Check whether credit has been given for this quadrant already
      foreach my $Q (0..3) {
        if ($THETA > $Q*90 && $THETA < ($Q+1)*90 && $quadrants[$Q]==1) {
           $ans->{score} = $result->{score} = $score = 0;
           $ans->{ans_message} = "You have already received credit for a vector in this quadrant. <BR> Reposition the vector to another quadrant and try again.";
           last;
        }
      }
    }
    if ($score == 1) {
      # if we're still correct, record credit for the quadrant
      foreach my $Q (0..3) {
        if ($THETA > $Q*90 && $THETA < ($Q+1)*90) {
           $quadrants[$Q] = 1;
           last;
        }
      }
      $correct++ if $correct < $needed;
      $state->{num_of_correct_ans}++;
      my $need = $needed - $correct;
      $ans->{ans_message} = "Good job! <BR>".
        "You need to get the problem correct $need more time".($need == 1 ? "" : "s")." for full credit."
            if $need && $ans->{ans_message} eq '';
    } else {
      $state->{num_of_incorrect_ans}++;
    }
    $state->{recorded_score} = $state->{sub_recorded_score} = $correct/$needed;
    $data = join(':',$correct,@quadrants);
  }

  INSERT_RESPONSE($name,'_data',$data);
  $main::inputs_ref->{_data} = $data;
  push @{ $main::PG->{flags}{KEPT_EXTRA_ANSWERS} }, '_data';
  $result->{msg} = '<input type="hidden" name="_data" value="'.$data.'" />' .
    ($correct ?
      ($correct == $needed ?
        'You have completed the problem.' :
        'You have '.$correct.' out of '.$needed.' correct so far.') :
      'You must get this problem correct '.$needed.' times for full credit.');

  return ($result, $state);
});

sub GEOGEBRA_PARAMS {
    my $result = {
showResetIcon=>"true" ,
enableRightClick=>"false" ,
enableLabelDrags=>"false",
showMenuBar=>"false" ,
showToolBar=>"false",
showAlgebraInput=>"false",
enableShiftDragZoom=>"false",
useBrowserForJS=>"true",
ggbbase64=>"UEsDBBQACAgIAKqy/kwAAAAAAAAAAAAAAAAXAAAAZ2VvZ2VicmFfZGVmYXVsdHMyZC54bWztmltz4jYUgJ+7v0Kjp/YhYBkMbCbOTnZnOs1MNptpMjt9FbYwaoTkWnIw/PqVJfBlgRQMKSRNHiIfWTd/5+joSOLiUzZh4IkkkgruQ9RyICA8ECHlkQ9TNTobwE+XHy4iIiIyTDAYiWSClQ+9vGRRT0ut7sDUxnHsw4BhKWkAQcywyqv4cAoByCQ95+IWT4iMcUDugzGZ4BsRYGVaGSsVn7fb0+m0teyvJZKoHUWqlckQAj1WLn24eDjXzdUqTTumuOs4qP3X1xvb/BnlUmEeEAj0d4RkhFOmpH4kjEwIV0DNYuLDWFCuIGB4SJgP73IJ/DpKCPkNgkUljceBlx9+uZBjMQVi+DcJdJ5KUlLUM0I7L6NffxFMJCDxYb8PQWSToQ9dz9OUWDzGPnRsYYZnJAFPmBU5OFUiMPVN7ggzSZZldU9fRUjsm+6iPKcTQxFIRbQCEAQyJiQ0T/YLkdHGzCi20l4gRBJKkPnwFt9CMFukc5uaIobNPZ0vuvSquWrGSGXkF+0F1u0AhyQmPNSFapRRI8q9gaGcJ0ObvGbI3ZeG3HuHvAky2p3yN15l6zZii1zPwDXpu6Oo0L3mf5JIj7nKuPPO+KCM6xbcbUTXMWydV0rWFLEMZf5fhzFiEjOSHRA8o7yEeGOEArrbLLqoQneOgtxpjDzHYeGpMQ0eOZEyJ1u2mz/8QUO9epn+hA4bqdItof7AtkD+4TWVUa0xqss8r4ZRygNlHMoC7Zc0earqotN1jqGNss1DK2Nf0ptZShLlUsHlfimXht0soPt/G7ZIFct7vuZK77KIMVe58mmPhMQPuqlv/CHBXOZbrbolbdZbgmfP6cx719np6Wzpt26/46TQRKqD+5Eee1hVX7PwaOMC3nK9Y+twB2++lsj+Ic1JGfSO1noYs+o18wqu012PsdU/YbN60p8nSh7fF2IZI7zHa7u7yDVBNk4UkRTzf9uysFlUmeN3S7nQR9/qY/8x7ryp9DpGpx5asW/k2D/U/egg1EPusdX8PODa9uSuyCgRoyMhPtFJs5lmIHh+/r3cXlip4Nh9Y67jAPs4GhFuPa4EIHNMsZljKs+dxc1Ehow8Q+btHNlsU18PPKEZuLI1rmzBK9cmHZt0beIVgJptHo1qY+21KuHzT0tDt9mO5zU5kjep9P8ghOfphCQV13C7lAvj8axz0O2lpKbaLVzBJjvZbBWS0VCb0IRqJZ1p7U1wZrSIh1KwVJH7ICGEl5d01oynNFTjPLDTfY9olpuLfTEWCZ0LrgoYIJ8EV8zc5tUOONZZj/tc/Fqz1f28M+YRKyfjlZVKBdjze1Po56O9dXqpInQWBHstd9BBA6/j9FH/ozfobUkUDRoTrXkbC2Gb5QQ5W1vR/t5mJ5/hrvMZOAnKQ9iOc2DDWNlK/l5klJugUzwZNCazUvTFDv2YCFJZHmRbqSA0eGPhDU4zyihOZqs9vRhhRbIywHgwQuV3CCcIePOnaOxRObRrK1Wu++3HjKimyPFEV7CdUP4ZB49RIlIeri5bB/l0dGzb2gxtKAQjuHREn5dy5Zp5JVDYBGj71eDFZl8wJsHjUGS1xe15H0NlOQNujFC5/l0zA/ZZ886ObgpNjvS2vZU8WzkAbFd+CdVe/trq8gdQSwcIBq3M5MsEAAAPJgAAUEsDBBQACAgIAKqy/kwAAAAAAAAAAAAAAAAXAAAAZ2VvZ2VicmFfZGVmYXVsdHMzZC54bWztl91u2yAUx6/Xp0Dc14bETpsqbhV1F5vUTpt6s1uCTxI2B1wgH+6r7R32TMPguE7bVGvUbZq2XPjwcc4Bfn8byOhisyjQCrQRSmaYRgQjkFzlQs4yvLTT41N8cX40moGawUQzNFV6wWyG09qzjXO1KDn10awsM8wLZozgGJUFs3VIhtcYoY0RZ1J9YAswJeNww+ewYFeKM+uzzK0tz+J4vV5H2/EipWfxbGajjckxcnOVJsNN4cyl2wla9717jxAaf76+CumPhTSWSQ4YuXXkMGXLwhpXhAIWIC2yVQluxkoK3ndjFGwCRYbfS+sWB7yeGeJLvXLxTXCG+zQl+PzozcjM1RqpyRfnl2Grl9DG+0pc+7juS1UojXSGe2mKkQNLeycYTTJc8yrKOXOliJLwo8mQUDqgvRBesAo0WrGi9vYtbGkV9xl965QVBra+buxrlUPoSRp/KRaeMDIWnDgUI1MC5L4UVk+8UpUXvZtPSLixVQHIzgX/KsE4+mknqC68E3kO9bsTYkDMQK4cEKWNU5z4USri3e9I84ptqK9X1Pfe0dDs491UtdigcYgYB8dxL5h+MEkwaYsEbmWYp6mfGS6Zdi+ZS8Tr/lHcaP1IdbYRpiP6uK6+3RGa9A8SmniZyUOR/1JJ99NFTRncmr9/ex62/4o40xaMYLKD/bLueMh98K9z3w/S5ZfQ4ffR13f4uT3wIH7DoQfYo0OP0Nt2h0pfCyNXSucGbcImELYG/1y3KaesPoKaUSL6DERyIERVVHPItZL3HDtN9yj7DcpDvpyX4qdp3/NP6aPjIWmOh3Q4IMkgeTUtDn2l95OVYNtlfqjLXZbpf5YvYXm7ZLnfZ5ulftrWu0zpYXcRkjy9WUYnrwbkV1wgnrw+1I3hjlAFc9drE770RoHGg2BOgjkNZrj3tiEWZSG4sM9LaZZ66i69Tx2ATdeuqsmfUvU+8W85BOnP7t9x5/Ieb/8gnP8AUEsHCMQHrm/gAgAAwgwAAFBLAwQUAAgICACqsv5MAAAAAAAAAAAAAAAAFgAAAGdlb2dlYnJhX2phdmFzY3JpcHQuanNLK81LLsnMz1NIT0/yz/PMyyzR0FSoruUCAFBLBwhFzN5dGgAAABgAAABQSwMEFAAICAgAqrL+TAAAAAAAAAAAAAAAAAwAAABnZW9nZWJyYS54bWzVGdty28b12fmKM6gnI7USuYs7HdIZ2S/NjJxkIteTaZSHBbAkEYEAg13e3HSmn9Rv6A/0l3rOLkCCpKRKluJObYOL3T177reFh1+vZwUsZa3yqhw5vMcckGVaZXk5GTkLPT6Pna9ffzGcyGoik1rAuKpnQo+cgCC353DW82NzWsznIycthFJ56sC8EJqOjJyVA3k2coQfRCLKkvM0iINz30+T83iQsHMex36UBCzyYukArFX+qqy+FTOp5iKVV+lUzsRllQpt6E21nr/q91erVa/lrFfVk/5kkvTWKnMApSrVyGleXiG6vUMrz4C7jPH+j+8uLfrzvFRalCnSJ4kX+esvXgxXeZlVK1jlmZ6OnJB7DkxlPpmSCtzAgT4BzVEPc5nqfCkVHu1Mjcx6NncMmChp/4V9g2IrjgNZvswzWY8c1vPc0A19N/JZyAcD5jpQ1bksdQPLG5r9FttwmcuVRUtvhqLPBhGaJ1d5UsiRMxaFQqnyclyjRrdzpTeFTARS1fUC5zuG+Jn5iyD5R0no0LBWEyMn9tlZwNhZhE8QMMtOh3bAkWVdVYXBzOA34BAwfIAP4AzCCFdc4AH4uBLjSgQerQXcBw8IhHvg+zj6tMxD2gvwfMCAc1wGl4HrgsvB9XAaBBCEEER00EXYcGCQMXwIGtnBx6M1z8PHrHk+Pi69IaLAokEmAi80bwFBI/7AJfbNoheDP0BCtBBEHDzkAecRA8ToEXpuhPAZ0D8OPqF3I3BjQHwoN2Emaz7cKs3CgVlaowRdo3A0Bj0hPsZaB0bx902CFmAo2xkN3A7EbhjaLWbXmGcH1w6+HQIL49vjvgW10jLfwvjeU8VshfQeI2TcEZKTEGgU4t4MHhDf3PBPg99MQzs1rsY4a1Zj+hnQBHUSxubliTJ5nyQT71C1UfqIKG4phpw/nOLTXHQrpfsoKZ+o3JYoDzpEA8pTJlUFt5D0HiXnnap9BMVwLwg/N3Xuxs+jYz9+MM2I3Zp27Mib8bPb4anZaauI/0Jy2G8r87BRAqgpwbY0tZwp0os3gMiD0N2WypCKWVMvIxeiAKKwUzXPqG6Gwa50UuGM90pnEHfqJxbPkBYjU4yRHlU/W0tdvy2nZ01B/e2ooGL983clEBkkVBwA67VJjW0tRC7cbTV0AyqILiZPLMQuhJSM7yiM2CVWKt8qdyqL+VZHRo95OV/ofd2ls6x1Bl0huChME9gcyKr05s2huqVQuosXO6hdn2Y7qr027sWwEIkssBG+ImcAWIqC4tZQGFelhtb1QsegMx3jUC7SIs9yUX5A47ft2beLWSJrMK8VSWmQ0HHYtZa821p6FiStqjq72ih0FVj/VdYVlQ+/xzp/PEwrG7uFnWNv0P2DWyoV5Of+YH8nRrVv7tqztOXySmqN8isQa6lafU/qPOu+f6PeVEW21ey8ykv9Vsz1ojZ3CAy2moS6KCeFNJo0VsaOO71JqvWVLRahxfV+M6cCaeknk7dVUdWAQegG2GpPmjGxo4EhxrZQzMAwA9HYiZBu9/nANRBmTOxooNDIlrVGUN5KyVoquQI7b3yqyUPkINS+L8pcX7YTnac3O0EJ3pq/1eA+Sv5MKIf9A88b3si6lIX1ohLtuKgWyjqxNZXhY6Hk90JPL8rsBznBEPxeUB7UiNqC7jjOZJrP8KBdbzQnyKp/QVbtaiYntWwltCFp9drEDqh5LUWmplLqrXatj3fBjDgt+0OV1vmc/BASzMI3cudrWa4EZvGsIxHJqpDplPIJ6k2Tzj5gUKMLnMM7MUFeFxkauszAuCTeWBd6WtXmFiY0Qpu72qqqbwyb7+Vag0iqJeERdU4EFSYdmC+KAsZ1NaMrbyH1l3/g7CvzMxMTGFF1W+KPWdFTqQWuXZsXdD5ZkMkbhmcz4qY09cey6uzSnmAj58TFjHlqtVUtdLu8tIpqEFDuKeQMb4qgTRgtLarWpZYGKfkOVMkvuHfgcnYil3RZsJSSXw5Ci7M2uEAU86nYmrAQG0pvnQRpkH43HiupYT1yzind4BB1dt9VWWPxbrJTBO4aaPz9uEO4jVE9xWjAuy86P2YDKyuzL3/Os0zaMlLNRZprxMKj+MjPlI2lVNRaKoyYBgCzFnmN4enl9aWoJxKuUY1/W/79ZRNjVsNHZnub1yn60oHZ2BmwU9QV36b8jvXS+62HfkwfUBrjpMfGaxuTjvXuMNvjjLYzy7NpHquaLK07KoA1M2AbZg5/ZM1npzU3841RFnzkdtmcnwld52u4sAcvLMSFi/5Eo2fOXfh2CLZ8y1/LPWtTbc/HeXpoSbnGpKToQ1arzAuMgzV2FSfrk+XpGWzw97SV19Q320DsG6zZ2KL4JIOFsbEYDYkdPs1mt4dS44SGUyq89ojfXT1Kw3d5/A9ic4+7Xxy5+/h+d68RXaua8f/Q2XeKO/9dkhBqpNhgLfumpBorTZk6rso3Us6pF/qufF+LUtEX1P1yfLddtogPrJNaw4yPDPPGbtxvnn3nfvMgA91VS57fx1kvYhFnYRRzHoeBHwXGdLctP3MUXNTp7Xo+4U0wuKSuQ51nD8j9cyxOW4VmTynenLmd0HCZ//jybRXK+X3l+/7gcP8/C8WzRexd7YXpCI+6i+OapLEP5U1d+vLXRaW/enldfJC1bjsUuK5p+vIV2G34E5xcivfyx59EorCC/XyKK3bL9KQvG9LH8CeXspzo6U/Zz9CHf/8D/gg8Zt3z//qnHW8rg8Sns8/0Q103/p1yeq6MXIf2MHd4Jet83DJBF/J35E5++40nclqnaFwfdVkVCy2vUrzolO3/SJlCO7D1YnB/xrC6PUgay6MMgXeI+3NEuZgh67sOkQ6Yzw2iWDTiu73YjX034q4f+SEfsCZQ1WaWVEWedhQCD9V8jzefHbg/YJyH3H2A69pLj22pdr5l3Mp42OltjnQooMVyIKLv73+6GHwOCfvdO6b58tP8H+Tr/wBQSwcI50AIx7IIAABPHQAAUEsDBBQACAgIAKqy/kwAAAAAAAAAAAAAAAAWAAAAZ2VvZ2VicmFfdGh1bWJuYWlsLnBuZ4y8d1wTbbcuHIgEEEgUVEC6dFG6IfQuHQRRVHoHBQm9dyQgvQlKF6RYALHQBEJvEjrSSei99/pNfB5999nf3ud3/iHMkGRm7llrXWWtIfyephLFxesXQSAQhYqygg4IRAj8SvCdBAK8kIKlYPgtZx0lOVARhn4RBJIBqSjI6noQrBgSp10d2JUqcfDOUQ3nelDi3jTkLUygdovEbH3cX/iminA4UVQzTdTV/gJmE+Hmie6At8oCAdrMvRdyVQzdAsctk/czHV4du02MvnbmFXdwnxyVTJ7cmuus2E9b99sf27KaG8su4NR4Ldi03XmO3J7rPAtatRZ5Ptba3tHRkTP63d58W1ksTzO9/VVxUVH0Q13dEk82ulS4vVmHPNdsHJKQtcRJpZ67re4OYj6+f00nVd798WNa6xs3cWOVrs+u+BvEI1+9euXUZDj86rb+kASlZqFEov/ugF62ne6DB4o8PDzyi7EjwQ9qDLkce2jOyCb1P5sRkenV207X6hi/svv45NtNF6ZYu6EdS5WWICKyR7wsiT5KpaZ9yewMQtaYK46qvtqlpkLUKuYvhZs5c/jUakjfsjHB6offuPGLfh+ijE4odSqUuKgmltYOL6egS9NhJOW8ybIYWx47A4enRfdjfVQR7Iyge/XcwyU4ud4eGoS9XUhdeGl+pp+Y62Lvi7RwnpLVhiF+VXaGeHoDBNVL7flj7sXYmXs3Z+NmMHJSX6FtEKRqDcsi8pr9vQYOFRQbcXMJ4l1N9EVlBFtFyGV4d5GBapKA0ftnlzBhv1hkwvarjJvSjTkS4+JwnE5dspjFAPrCBx8eruVP+sdDE3F2WT0bLXdKLMBbymxTxHxuGA4Yimq2zvKRlCgdFuSGqGGXosA8avKFxqIaWNoEWsE5/bj6IZAd1mZczGOcB5fY3hDaH8MGvJfhzXJ/fDSzqIuxUxv71xZG7+6xCufFgpzot5vTLX2a4qsXFhy5lVFtId9BzSU+qvYXf+owqjzHr2sa/RnPJntbdvAX/DLT2fBhICMPe1na0o2riyJCNo0LYbRpps2aDWARuKo9If3dGP3FkHN147u+QfGN+Wwooxmczw/DKre1oacjzSbhFlZWb8zdQXt+MaQSxMwhHeBLafQhMpg43kUQPQYCYTtxBMIpnp7GbgtJg73SSyhTytLWQ7rZgos9SBwZCyUj5XRBc8bGqBlIMk/PdEdA0UDMvZXYQp6RPaah44JNkuSwFjDOZ3jzrtaen+2MTp69qjmco6vcxojbZvruOcY4KqJTg1xSJJqLy4kCcuXjG/Z1svnUED0x7yRMZzdIGypsgxu12kIb2XdUxNgq6q40+7+TukBneuxk3qnYHI1j37w/dT9ePjMx33t/bc0gAfFckVLcYWpfoSTW2dNuNlzmhxsbQ2uHqYxSHVOyEnDuFBBMaFuwVq2FHVJMnbF2C8HZVvclsWWhsJjivpSdJPkgM32zZuzYd59HaWPF0ODbH6NCN4znhy+MOHdJLhZVrTmZBhUHQ8XCpA5nqLywgcQLPUpMdzbGtdocIiXsDN9t6VRdp3fEShqziZaqAifQSp3IxqAjGaUV8tl06XpITysQrJad+Zm3WxEVdTi9wbSEg3g7GtGdC+xtkHbHXZuPFVw8PB3p25GS2GbEC61b1x3JmKWla6EPz+YfXuLxMqK9B+mNDgTDJGI3klOQiiD/N4kElszVNOiQdQk25We/it41RVD3Vbouj72QkZUlcHX323VfGYpLQ9nYDn587NJMa33Pzc3NgMiV4JISyk8b8oYYTAMS5ga3y2xbgsqu1ZyfeR+dbHe1hb9B+o+nkjKWYTbKgHtVFREEtx8pJwtWgLsQxLaG+elBelMC37+kSHzdB38YXGM9KetsXMG3ZwjDn84pLXpG2ryZ3Ot4tWwcdQwVnbm0T+1vaprtPcykR5TWLX10tFw83j7FImtwW1m9EqRsoVDLA7b9KfNgmP5okY/vXZmSxMLgW1e2s3jj63oa2TzgdrQbhP+wdndQv7/G11vscCbBNtnQLXtszEDvZThj8mm676Wpg/qrBAktJC4EJindWpZZtPwlS8ukusrji6si54pE6V59XVoSb2n7VT9kkXOBQwMHPuj1URv74Iy2WvwYxytdI8gUFVwuXANJeu/lMkp6lTdtxDJjODkMQWyCbyeTAoWtsH2cGdmvjWf7VFDGd7XGna+MP+FPmeuiSW++o7PQHiBoQoMuPj7cnq+CnVDQ8KPHf5JpHU2/ClAhIR3zXq8eD2NyxgyPx23Qy5g9vWNaH1IpAqF0XnS3lKFV9oSIuS4pVO0uTtsVBFzisb7Hpa8YV4/RfVAdZYU5gJle2sgzbQjt8d1f3CYw0tx5RKtsCbnUHBFVVRL+azbZeM1o7nyzhWn8+MwDdImM4KRX68dkZnLyvJGATOuNa+Tlktg5x7MIdttiV4vEPfaJElN59cxJgj0HtynKU4XPzvHQW0/KZy4x15ABub2KojgGinD+l2DdK48q2HCJM9jrac4AgHwTdVE7fq6GSEfpNroTKCPUVFXLn1YKcjK+6mQFypr8Ya5Kgtwwcm0gfJ2QpL6ihddbQTrly+Bsqcts2aMxpQ/gylF0IFge12i0hAmgUlW6JxmpczARz672U29SSwXVr5WtO7JtWPrZ53Cr21fmHwT5Yt21N9vtLUjrS3v6Wo0S05bA6dxk+FqYVwXea8m7dZRHtUbtu/K55W2M9JlHP+lhhCPb4UF7Ek8OdhrF5BeWCtu5JnOTb/9KqmBJMDl5rYnW0azWcU/d8brG9tk2abrjtNCMAXl3LhAOnQEEK0Vhx0fl+Q5yZ3bvXxif9uVr5Xtuz/UU5mcsSQmHOVWCSky1trnAtvVodJVx9oNaI7Uv835K6d6rGuJlfaeg3tzFK3mCzIIB4yfRUkc2+05+mprRKZMWaGeCZ4MfWeZVA44nAoY8nbMygVDplqBp5bHTOl71L6mS9tguiBuUAxIbDVrDHdtvNdkrRTp5wE163iquBQcMMXTnYqM+aEBPdSGJgV/N0RXSmW4pV0sEs8iHBAJIGWBbn4xrJ07l/d3XNlsUwTQaHtQg3tPpi6DnTFeoqeniJgjDBLKJgWPio75d5jp/gN6l2v0J7zVxEIh6/wComhrAWiFw4f4J/IfmVSbzBE/d+7LOTo/p4o4Lu0CNPdnGQErLAKeXGNhvBYKBpuKlzmJLqEyKjaqn3I4aHgSen+3XCjtOKexXg2oPmvhIsusnrjjKkY/9Ippie0imyZWkZzRPo/r5nck9p/rLQQprnisq5AyPPl9gVyFud0Rowwxu4CaUbg8j5RR83DW3zL63NeAuxHhYsrqf3wV4Szricll1idkMJer+U+eFpmWOOu978cgbblelWKqfJjcshhANhPtrNIIZscGyQJgmoYyQGFyiHdo9Gwxbwx9jBOxVfuGcLJhKgtBc/ojPB+ZvihzZuR/OF+4wJHOxj5iNboMUidm1BGJwrgB8pJ3ilcXH+mLTeGh+hP7onaPKFKHElO1RaFpEN7F2Mr+z9lsHs8jGVzUQEak4Yr4LTvUnY9sFx/UvyGwXe9+FoIg2OdB+p0cV9s9GWuxZf0ig2Vws1k8JnCeOw1968Al4WM6l9AXebVhqChCxpJQWBaCpTVbsrLvcYMoF7bnZyJ4M4Cy8+3VPSm9uyyJISjNaFhzNKynQvlhi63wQ/Sv5xGPazsQkSAsZHKqrl8h38+aWEIM4KlDMUXolBXlMgOrLk0AixUxCE3KC346Xy8r3EmRueT+y8b46DluYRyrDsXrddbdBbABbOurRqsd+XcyH+zjS8IOC1/l9ZaZuPPZ9/YGYny2HT16L5nyggoR8MJd5cqX3XcTxNLT2pPZ4TVzkZ1diHeeJHwDe2TnRYu3pXwTWOUKhkPPGjae95v8ECj5elGkhYi4LMuX5MZLnjgOQ2r3Ne8zXXh5sjlR7u7vjIokuTJB7LGSjBuuKyALEEhdI/3yS+RstYquVdTycUU8N+cK9uHrv4oVInnU/lQbQ6aMe57+BfFEoYFRjbi91v+YsWmaQCB1d8s4z6U+gWtUH8ILeqMg/vfM3dhlMZahB4vWv5h5A/t33hhR8uZFrqlA2IXfIdGZ+a6i9gl9GhbhkxKSkzXwp/SC81MAzuI5D6+3zbQZLUcZYrSm5+Bz7dxIlDQMNeT8SleLpeNmk3lzP14Uhw69AOO2wXWOTH3pozmtDwHui5N2QSezW0BCaqnlik5XR0uOcKBQgnl4ggGvcfZ7AST4J931cz+z4+hZvFqc6wBva0KQ+2Ag1ykxKTXNWcmo4QLgOTMNSUsx1L2BYJNVBosnEdxGiFjRD/N3TnQqYqr7Fj18wL7/V6TtUKxIzuMSW/wJZN8R+9tg0E5HaGjKPiVo8P548XxtJWLo5AL5ov4VEaAQjwdUW8vGpXmZZc5QoKkHN5w0vQqicbLdQDFCURUMiLKchlKxBdOWQnZRzZsM0GkbHYCAZyEOLMM1hcZSTLscu3XAMBBYFCRD0GQirtRJAOjW8HX2HHxIv0Sey1bDR0RM35+U29A/M7B8nuoQkNeTlgj3Lj976DNXc+qZEk3Z53Up0oTWcjkimQiRbzoDxS09MWeQ+DUXEVsnumHNXTwNMWl0+lDzc8nB83U9JHgV9B8BeE+aQgc0sRcKzVP/52cnh8xyC+XX/c293go0vmae38qYG9R1cRUjEtn8KDXsF7CjQCbCPwoAkEkMTiAn9ymaXYhc9uDuFVHYq7NRmTP2GhLOlWWlz2V9oHEY3C6fRI3poK36LHjvlU+VSglq5Kke6QmWxks+UelrYC/MGjG0QF2U6AoB7ivrATlnJNVP9bm6y40ImVddtkl2UnVJrbgH6iN7+TWSNnQ840w8oQbssC4nNAJ07SHQjTHQBEc+nosOk1yUiBCfYhmN9qCrtfMZg1glrYBP6ETUI5jBdKWN/0n9yfGtj1KE2qvpCXSJRG9m1W8wnI/774+7jsqzKWt2eOWrSf4N7DiOTwAtySzT48RNEJZm8MZgz/MW6b8J7vdPy/ut/sgNAmLe0yh+4QM62VxJ+VOgEaBmUyolNhbGOqbEAnMw44rnupx9/3sosEPuQLCCrHSs1U+44vfClS8p8PZ846df231xTM2devQaa+5YaTuG606PEmdG7/udvAYKX0DIjt0BGvrsP9j8pBRPD8syjRUONvBbZy1Q9OEx+Gbq/w6SLt7yggvSX1Wnq2d3MdrDuzHWw7sppvyn1/es1hcSHt/VXalo8Bz+e9x71vnF/+av9WfT7sld9J+edrhn2B0xJr4hrO6Q+pwGc89/MDaIJS3z1qjJbuzXgUn58fHzWjzXPkFvMt1i4kpKSiMLHtsuNYQwk2zJD19Zw3vaEYIx3122D8liM9LnPk89mL8Q3z5eNz50aLhBD6Rfm9Estwj/QlYHuQcVK+ANuS4hgB4i8y18y4xQYE2fj27bbN1hcmsDrjaVzX5flpqVOgl6mCxmNfibHyD5okxDbW+hbqJ207mudX8vNZzqT1Di2Ekh3bLLfGctcy3o4KlR7glJbfFX7426q+YvtxcF845zJtNFJSdYB8K6id+B9XiKcAm1hZxzxujzxOoBfpEUajyxioWnzyKmklAK+ZHSwp1UjX2bI/Zuw9HCpdBulUoLJF1BeNp3ewL1YxpQXxkgfVTiald/p+JHy7xAFwO2N35sh6MQ6pTGl5v8CcCmOQpHn9VK/5tDcUWqU1Fjti2B8oKKofC+54uS82OniHQiDo2KOcZHYKDmObFZqavORi5oFykZGRkiIUv2iTA4LY0+j51Rf6mRjoVfXuOd8urD9iPDT0e+fqtxBkdOB0jN1k28VI8Y+gD2kBWuCXU6yHXEhRoYlyampakRERI7zVDwhTh75gsmx86l6xS6c/kMjSy6a6tKzyw+7bYl3szbPM7AOrJXhT76JDj2QdFck3V01UEg28rmu+ORbzq9MzwBN8WOzyo2Qmp1+1/nu23y33vLNnyUXVWMNjuyW7PJ/cnhttGf2OPDG6GzuPQgq8wv3ELBB92xWKvM6itmN0K/O7bJL66dRkKvb0D3vDtDUJG6nvnr18HV0uryS7LFz+BfT+1I9+ad2jnD+ADyKgsVEGcLSUTHOpEeiDPEUB1ddkGQgFRQbyk/dC08mLOlqEOQ0uXXuKVLMjnEfa34OgGnie5KzOSPPCYNDqI4S7lATOACsBem0bnAll2lelAFeYzs97QfvnbqGno6RXqtZ9hh4/0D42eCtZCU7Ggud/bXxcR5CFL1TmOgAV2bPew3hAz2HXtSXa1UJK+2imU3Dp08rztlP2anOGDtEHAaaV1vsC11z80StbqVoPUj91UcV4dVdyiSVXX28ppipefgKeq7ykebThp5hqmzMH0QkCA38+gi0UtrBu1bc8MvIu/PTRb+61c8tjI4D1wXMLhh6Lb5znK/3e7wxLjFy75rPos90TMeCuzlHb/jxytSbkcLJyzs4KurGn2eXUsfM88eC7Qe+iRHTf6/dTmUsTHs7ypcNEzvjfznz7V3xpCezGw9INHo57PyZDtPZjci2vMy/9YMPbPsZxN/K5LXAsl+iFMM4l3gc3tECCFlXB7JGUjkfUVyYb6H75HiBvSn09q3C9JLqmxFVa3ep08tr96/sDp0WDO2febu8ZPCGVTLpVe0K2D+WJ/rktGB9zCjzIt39Xf5Ehuypd9m2lXuMHuY8cS3jsbPXCGw5fHK4VtNRMvufEpZF3oivgaCxZSWx1a99UzcD9yEeDqSNpMuH7+xW2x+1ES2fUoM+k2cEw4x/TJ4O7uZ+WOSULRvXd39LlGC0l/g+C1aTIV+1Ix4QmqHjTpXTpykVX7lTSOEs4TG/9rXk1kGgeIyZpO6f8uUDlGYiGYA/j6lis0kSLrHwO80oVy4XVf0szPWR2tN35vaHiVlRn9WVPgl+57K6vZ1ebGi76a3Vn7EnuuR54Cm8vq3tGP8484fmIcNejpOEyoLKuZe7RLp/91qGELtLe4+SR81y9p/6XGJOnjtAiE3Qksp2Wx3p2cA1LNYcrzn778Jo+AM2RiYt8oNREgtS4QxRfmqvH9560PPY90hHVtvhifaRpL1/Bh/oITvFwujLI3o/KUlxnP51v0HSb4eDXW9HM87K96qc+2+yphq9gacOnjmYrZy4RzxMWCO5lCMP+/da3WIDSer11H9y1qvnpiiEbL4+8ft+iWP29EVGCjGLJMcMWUs4fJPBSZo+5cSRDuRwvDz4ie6bbV/o86pV+K8QEekNlhNCr43FpNf3rgOkpLWD3SFjfq06yCh1Mu7Vp3QFIl3c0Mes4gk/1JHY/h7ZjU9T/eq+R2Jex+Lzuj0TwsjadO75w1+FDUXBi1vStbNz3i4z+i7rh6jKFqO0x4awLCUHieTnB8uI9rbnGrOYDDIjz7S4VWPyTA74Lz4YxlNU9SfrN9EtMHH9ZFB8SJ8bEwjKaafMYktiGKx1gZ00X5fRq64iJGGZFU3MYKCcG0J1+kQVK0GuuWMcYvbC1r2QcUNyiDSYCp5bcwtPGozR+9fjyfoJDJAA9wMUtQ/EIKQQpWH8CqnqS4IAtEZpIHy6PIRChXgLWRNeE3lIsnotkQ2oxnJ8LspSayihqagchRy1zw1IJ4fXjarZ9WlQ4TfDxGi7emwnjJxG98c8pb3LYu8LqPTJRmPf53V0gPSCSmR2iN+QiDdB40V2et2k5N6D/cls3ioGZdRan7jsk+2rlSV66CmtjLejRm+0lts57n8XJ75rO/6eyGVZ6PztMtMJlaXYsJ6vNIpD8/BNNGJz/qtJ+/vrRIbi2qlBPi8RZcK83jZXDHdLvqAIBuk/fRk7cnAUM7N6qElP3ZJZyhf3BTT07hnsptNUdA7hJPIiovIrDNRSBgAR+QpSjK2WP0LueW8wAwYidp1TyfMAXAYLwGMdnY50vJ3JPcLMjckstkUZ0XAyZOdsnB26OCBO3wXPxVzMkWTBfTt9zLydh/zwkuzrzU38Ymy8dkelufJPnZFxkTXkkqYnekvy3XxRanCDlAqSo5P4XLBlLhwF/ekqy/2RYuHbpqN0mu2uGJaZkG7qREWJQaFv7pb7Ay3jinDoygyTe9lCh3Xh/KZe1eGIaEKFyK4UFum+VTxSGii2P8T4/VF3sFGl1/GHl3ofsz9hQoz7SA7YZxPdaiZdDtt2Vqv2+4ykaxse7MfwUfc8Rrm5x1j/NJjyq9nEpa7ln+X/qTr9KIIPVaO72mlrlu0dT99ua8d80z8puzBRvqI/shrGQ/hhxHL9jc2LJdWuwTaM5XHm+Kmqxuv41zuNo3fo6HwPhr7N1Mhl6fl7YXunbQaCVgX8NTHOfiOpa/SbvJ3HW0wuPSX5Cx1NHwqNj7FXLZmZr/572EdCAcaDnaJ3Fq+CatUUyhVedAi2KXcneHhSUrQbRgu4tUgc8xu6HCfodp/dE/PTO+cKeJ1Zq5Sf7hfKrTO/3WiMSF5OTzQcCzwTXNG5ftxzS2wl17S/ru6Fhv+Z0/7tnZFTvR6effscZU2WU8sVj8rFselxwRZDkTehxWPhZ5trF+Zn5QZ21lq9bSYcrfus5hllN16YlhUAy1JeljdY/FhNcXXkfijv9gArUEPQqEAOQ9DJ5PkpU7rv3rBjiZH/qat713l5wu4JIBNjyukueEZMv5A9Pdp1Xufx6uJkdD9frBBe/aBGqdXiK72Upr+QGDHNd+wtpJ6GcDz+9dClr5o1XTqm/3GYgfcxshInQnsG0/jenjW67ptYdOZKWqX1VSmH8SZYAHLT7yBtZNWJzffUXHR1XECUbCnFSfK3uqzDE/ANEAOHucNSXz7s2I3dwWiRBibJJnVQ0iV1XN60YRLHYZklKljy6anQ7sS8ctnTzpOLVs52Up8SBLttyBea6u7mPKHgBkfVySSuIba5xFbODqrZrjeUjuC/f8cMkKqLA0YvH9Ka0D46a+rb6X1ZJLz2cnrI5+ZvEMkaGKtwzgmbV7zJHbcuwmoVN9p4QeCnwS0Rs2+bzIz7Lu9yzYH1WHEmDDNvMp4CB0zMvbJT3g3h5Y4JEaI1UN4twcnVcw9QGeB7IBTUlJptoRqpKmJ25UMlpsFVKlcBAYjfw74kkwmoMslxP9eDJPfuTUPnhRjM3hD1pAeP5O4GWNgBsyJGxzzuMR2186v2bAD21pu3567Z3VR4AqCLNqIkb5T73rJzQZpfW2L/9sECbM5lhzTPzJWP/9XCRzlq15ensQZ+Gz5fPclOH0/vF3jF6sQ+soLEHhpZKEWt4hP8rM0f3KLZUmhm257I/Xk6C9wSiPeOrDfMZzHV7vMEB/Z2jCotIaRUQ1tQ00f9+KYXpkjFzpPkHSDnTC485FVBsIuWQ+YN4pGq9tBxOpHnM+oK7GkYzcfa1lnWrsqIdNng27DsLx+upSrhbaMtlQSOF6UEjGji5CBktOf69EIL068T0Wm5nTYBB8UyiEtZIJjYiaMR7gzfddrLfg99x1rAeJfbx6Etnn3uw3p+rvxS/r0lu5npD6NpiXa3jGsk/vbFVoa/aCQRSif0JC8v10L6Z9oSOk2x+cK8KgZDPxV7U25EqbSlRLBNXpCNjyHG2fWiQJ+devCmZJu5SXfXFN8pk2N9eGIJletWZCEH3ll5BnEhiB3szMcmEfhDkLSDh6ZYvv4vv1MbQb4gH7AEaZcRskDPEd3oLBl5wg58ZIgA+EgJ2ML24R+bJauCP4ssYK43/gnszy4TYhrl10Ton6+Fxu8xhfSsAIxkRfyVe5bfmwd/TB2Au9CEYSUHPUxISDPDamfiqY6mm2chs47/GHucEiFZDY1G9XOL+ZmamgtXNFLhUV67S4rWGZ55VeenIpM+8oZVbrsDLbgQzUbFo2q7Vq4EpNkHk0ZjpzZ8I22zM4xhbsqSPS2eHhZUFz28zFIgRa9Z2JlNtZ8WxXQ4HIsMenYQcqW4E+dSH9RIPIkG+/BzqSDebXWTYUTTAGLO64IUmI2HfiXeQZL6XYZwelxoDt0xPmA8++10AYposiTwSh5MGY6lwqYtd8RDCzu1iXdy+Hf5p61ozl/4FnUG3r+ezMaLZjUit567saCu3k0wUdGIU2BSMKz/qQjIotTfsujKtWvNS6HRBgYG9Mdr7gEvZ2ZnD8PN5hEMNsI/3nXFRlefz3dlammOcTIg0QRF0fOZ/hPVF5Edt8u6xdd/UPUtS2LtAlxOVBxOp/GuXxEh7/AMzi5LyLw+gUNSk3jD0QnBR58+qqzYa9o3t8FQzHe42YxrRF/IxEIoSDkBrXJYW9GDlVVq6+y8Y+wrLZoGnw0M7mMZAIO5RqLohNkYNgob903YHl47asCGGMsSl4FikapOC0rSpMF9eZkjG02maGaKmoYdgvAIi13LIkZ07D94WcY9d8NEUx1ufrHFWnDuBt3kVzxGaniRYedSHdw4zyeqve+fn52ajwUL57pP+oq4acL+qglLZk9y4WEpUPnuRiOVVvkdWrjdI4OoNMOIqRrO4w79mEn1spNV36PVD/dtFjK+1dLkxzDKammlLL1p4NP5EmQ/vkvudQTLr5xnqjy8ef10vfBKJ28e2bpPwPOzQ5a20Pyjv0YexDqr7Roo526UE98qTqfrBbURjV4PtZEHy5q3YK0uX6YMuNi78UFF4fu5YeKc8+z5c9VpqbNBnyjfmp1M3NqmZz6DhvnkSeJRpu9V3/2ejh3/WkSkmPppRijj2O1Xen8vKNDFJKvtOcGY26g9r3S5wysPmONUY3jRR3kUVGW2pNSF6Nuqa1AxxmYWclQmdOi+NjZ8utxfmHA8XONLLeZc7IZJMEy1aT/QWiyR1sptK5ZQefYruXtz9XuPUNePmQ9aWjfB/0qgwp8yCbdBgKAk1Xo8VXd7PZNlJYiMlx+lcvAxz3jNm9nQqa6tpvJh52vNYyPxtIcfH+HObj8qQZ8afLESnbUUE8apdHcqjEvAIFbi+d8fthgrbjwo4vxrFl0xz1K5BsLQFZoJY++8wIzmm/mX3zHRmfpE4DxxGvmw77PhSJVWtPS57XjHnqNnOt/6ySLhTck9ucqlVqPwG/vH+Q4S6kySfqgcB63rbvs5Xy1mr4/MRZ3ZGMqgoAfc1U4Dsi51xWcK+W52SutFDorJxj4ZxVKShX8YTq4g+HLIUApCicqqngCe/UzWk3l+SpflfF76nJYnOlbG4Yza2KcniUX3AyD9HEsfscybTNRm5pPy/KrI5xX14qLgzaSnSu0mu1/sut3c99j/JOBV0ZO7HUaHtY9q3re/TDmbHN+d2Nv+YnuQLlLpT2S/9+61M2cvxt7JYHjyfJPUtfyM0sv8Z/CV20ty7pPnN9L9WzECDAJNGWCHcYn8cfkGdmqITQOnNViggKORt6dmXzy/AbszpyUBjegzEb7MHH7U6U7KyYmCYtV/zFq40Inmg2cvXw+pzWLbZQ3noYK0NpmpAWy0JPhZYVEavaAJeQfwXbskEVRiuRs64Xwq0POkLWWnBWw5v1/TYYYFeLVpOMBYWYytYpiYOfL8TSWh1j3CaH7zfBwLeSbaaucAF7IeRCxEfTgbd5A40x0uEQ8UG5f2filikhs/E9b9ZByhmu2zgVKaUDlAwonr+b6VvLQ74X2HavOqnkFdM279q3PMkzBDr8bhxoAQaeqW74JHi7LW3TsP2Pt5RRL8fhZUL8RdbhszWlN6L6v7gVLJks6qZ8LXT5N9Wfpd5Or0qXGmV/hZpSQd/Cd26aVdGr0mAfFc2qOp+/ENkDMEPSVQiTU8MaEhVGJYWGsLOwZikUbvJYBw9GRLy5MQsWE7cjXzZaPj4vv2GkB58P2PMWyEY2eGvYFnF3kr/u3t108l2mVV6jCqREvdmiXlQWnRYMBv1BuZHV+T82bJL3VQUuswBVH0UTuuJLPRpURga0uuJrIxcLVEcWJHhwqLKejI4Z4nyLv/tPi78S3+Gavg+Brf4+bFN1K+lVU59Ndn4pSeZvRrfXWooV53k/hkhEiQOuzadxmhzRf1isy14oS7XDOp2K4PdGkHF/dmNN1JrbVIeT3WlhFn6iawTK+fnVRkqFl7uvtt56eag/bh4NMRnSPR22OPZH1HB8cl7r7nuTSMtEzI/+62qUQJaQBbN1yz/jF74GAjhF8YJxVUSN/v0YGgzhLKwVBbkRy+zb5QLSh4R8VneAS7MrSk/DHzMHqrrd7W5ciXA2z96PdwAFrzXjx/Fr3eRGuFRDzF/UGfI9MkFxp6P8tO0XC+ttA4ea17gp5BM4gYCVlSe6uR21wxFvLx8sX2GQrkda1UzhybxwzuY7inQmQEOUdf+zSdYS0xTA6anzMZvFbU5k+2hNutnjruF4TvYCb9FKtWv6ceux8oajW7GPlu+blMMzU31UxnlfjqHS1/Ke6bWrdqxvDdHr4uPgj+SxZE+dHGDvZZ192WB6PE3FaUD5piQq72dmfL2+JtK6ufNGIbaPDt2+auD7bVQnlKGoefDvHZ61eEDX7c83eX9Wth/OrSLF0wemZ9f3SMI7sNVppx/w+hiK8HbY+Kch+Bk24ZlMdqis5eETmBfd7/oFech2+/ZskGE3Z4EI22/9KumpLb7215EXPSwfNS3G97V2R8nxSbTnjgbYOtEzh8sR/zjq5IIqOG4fQKg+9K18Jai1KLBSLMwkr03D1bcvRX9QlTsnYt15+ya9ssk3D6tWnWf2U3xvhcecyXRlqTWsszfdBhbOal0Y9DRi9cXt9paU9AH3Fj+nrFmm/lbKftWsWZtPhr6HgVBYPLIPJoqYp9ZzEEthB37K+1fu5NufvTUIuY0e961BsEyulG3X7YVOIy9V981KCxKSUKwFkrcfrdL3KeDObq83zLRLfpIrak7PzkxG6gUuJVy5Df9Avc1/W25L6RqFTbcdj2YqyUdDYjRv9i/sk9RXFqi8at5cC/Hr8Bjc0WkYxv+M/VH3z2j1Ybd4u532VUHz/6sCh+QPtSzN9DpNRTckrrCJZ3hw4i2gX+mK3Kqf7R/hj69dPTD1ShBQm1dz6/0f739D7RhJmoGTW+uoGjomJNbu/oGK8GBW2ZPDPmNW0gXOCvrTndUhW719w0LJbsNsKd3A5vNr8rmTukJk33GpOncPDh1qScx3rVcHc9jBUD0JkDhUghufZyfkrpNCOJePqyRyYtDlQkxYzQ/P657uD9bDrpraU7LZSaSIrMWLkVRNpQct2YjqAblXTdf0gdOb1mG8SHWuzSKB1QBWxgCzjKkxvkml8Olf/fGrn/3ecW4OFRS4vIFXjREC4VVnnQhM2BKgPqRyS7obnlFHn3IhmZ9oTGQuuc3mfR5INwPKtjz4g+xizK5JjuSz8nAFmmQVWgvZuo5llKazAmdX14f1IrRoMRAxET4pCw7Z12UTZAe+IWREXvVNEeDcj/ZW/DuKXAQFKpOorrlJqJLeLs5SHRQYYcAK/8yg2jT4MOzc3QyPQ8H8tiWQzQYg8d5hikPRKl/z985p0EfoT9USGFMT2EvBcrrs4UW6780HhuiYC8baGjm6rZ/74lbyntfYRUggD18EEenqTalXKBEGEc/+nI2lwHREC38I8PD/7mtQptN5FMQm9x81++kIdvk2o0VD35T5s0kcAQJMx9+tTg//pFyuu+IMIC+TH2er5y/ivce70cRb0cMcYr7wKjukVCPJCmupSaKtHdLEM4M1QgqoG7afStRH1WfTOUdQbigwxvKDwG9QtaU4nfagpRq3kHbGxFfQZxAxd6cNd40Bf0VDS5mT1ca5wY7TkZJfL8zgMzvp5GP4GIbxpdHQmXvW2p9GwTUB9Qa+B7o28rYTK3qf/bD0SQg/BOhjDwu8T8kcZ/9MwbAvARuYyQjcKnIMHYVpC7PrdAMb7RFPiqLYAX9MSq6tmnVAGwPg+8rz2Jp9Di/oY47I3j7ren3GgQb/ozVV+8FDuyzyVJIngGgf/0TSp1IitwmKzp/86fWuRv8tH1cWcBkKy3I4A1/Ppxb3AE7gzRHDu1q1zOPHdllX6l6R98tjc12vmj+YkkcD4FfOC6dbpPxXdv4z0ijWZ0agT1/dSicZsv22/WldYqXu+NLo9xNLwNTfrdKLfZnTw/E3npocqWqpd+A5PkjVaOZ01UrJRmW5T5Ubbz3qgtGsCJYD5A99+vtzGXrwcxxDOQ5kdRcMewJfO/ZWxDcw+rQDQbsc44lXiKVDhZ7+f1mmMdq2UdMZv1BlscLGIhgOBXz0EPrGdnsS8bznu6pFf7Vvr8sMtx/odxkF2KBGVyK+diSZyat9fQArb5owVWYYHDDukWH3h/yqA0cMj9R3S6WjyQNJHKYwi+xdijSzuoWbmbMDpKTVt1VYgmqTxXGpT0s7Ble+KQY3fgfRxznLdfQyjFx7HHFgrpFz97sjFQcaibdIinCywhxFKVSnF2N+peoe61s+Iu3DWZjQ+5Hx/8DIlhJ8UXlWrH6bsYBd36xZXwS657hCZ4ja8RNLR1SQxM9j9pfIZ9r692VWr2j8toWB2/vh+praWgR8QKsXx/vd3AsqSOyXa5MNDA9WL99uPSRp/d00QX5dMrqmpqar6uqQnxUbMdyUC0AatWc2sxRdL7e6EO1tp8x5J0HpD34E3aQPyQkt3cDVv1uOZcC7x0DIK1tQQjl4RAj6RsdfDDB8EYyeCJlO8/P2l4Dj/ZTzhmH1R//2kAn4Re/FldoP5g4dtvyXNvgid9N6hH9kvSbUlAVnPd413Vsp/u3OB5kKrmyA3WaMcPOiT63SkiZ0sOlOIcV2DXL8rqxFRP/EwjbO/AqBcRDOTh55tC8Umcmtf/kgt8h9xibkYb/i2PPv1kgB3IgwBTmUwCk1btOMHYInJsoYvGb4hEki8QBpgrWf7tAgtZAqWh86vptyd/doFaZBrDBVxExGU47BDYRkmnQ9JUCG58VbYtXImcsmI+/cBKj5j6DdtiwA0o5xKLPeONYXSecq+lbqoEveZGSlxc3XPGG7jEGc2wglGjIVCI2m+xT03nHrCfxYa78GuDdKOOHZpPRb3yHFqfRAzbqAN2J7bL8yHmkapwbFX35fAL0m9ZF2W8ZWDJWAAOgpRMZuC1cYDY/S+a//+1CwqgQ8R4CYkYWjwZkbYQXKBmZbXRUI7Hh8DOP/hg7iqvxMvLG+27Ly6jWG6IuReKn+bR1B0jlOll+TPJg6z3fDcOsf89yCOAU0lgzJefIsSQKXW1teEKTEZ/5nXixI2Mt4SaAcFO5RNEJ5xb+dZHWzsuc3FI1CVxbGxsyL0WDM3X8MRiXdAcnC2gl27TEAoVYp/3WVnoYq/Pzt1duFg7dHFWsD4I3/XdyckMqbPt07I/mOI3iYVMmh4FNCqWHgnqwUyUX9Sk3XSO60yp33wdWdHaLGLvMwbLnyAMaHmm86cEPqEBV4JsTCmKBv7WfDQQVf3/kcLXWvG10CTnWdIfnACiqQD2l0vdjg0kyaiBm0IPGqnKDOT8eGOi6LUeTvFxqYjdi24IKCGjeqljTIcfaUEiWBr6Uj10hLlU4GyiI1c8czm/4Z0bk+StoVYclcHB5RU2MCe1q2lwAg8sEt8FQDOhgxeTu1+n/J5ZMmElz+Qii0xzUEHYzWBZyQF+UB4LSKxIEUBTQBkMJO2PSi3kn25bOBJIGvzuq3DRIupc5bsU8/jADvfiy2Nnup8ltHSV3ZjAkAnm1lEcVkjAnWY2RzTY7vYGntBDWOe3lIEzsxuaieNXFqtv7yFlZwjl8zliT93CHyENfsa52damLUKKiqL3DLjDpbIba5cSnvkhAi95vB+loVHSQxdVKaF0FRoZ/AYE3oJ2bPLEerAvQ1NbyjS5cC0ItdUULJgKUb51FzMbVy7z+MTR93GvKRLbx3fnoXERH0b5RmD4JBWPpbKahEPNaOvmh2RWw14CfUpNbjznpzulhGpCGdnETWeR5B/lROnzeA921CghiXFbylsWt0BsdG1oUjobEW4W43vxyBHyHXzxteu1YF7kdzp4BIiI+taEF0IkYAyqyRZiskCqGk+Rw1ogcSIHRC87r6k+8AKsYG4DJL8x2aKh1ZSH67WErFOaEiOvynQx/vjklFBpT2TX7NNcwf8y6/VfHMl/Q+ZIIIss4Ggz793fIPpqBZQP23Gfi0Z/PlbABAReAZPDE9X/6xdd4r59O8p+KGCDotYd4v0A5UssBGcNgPLuyHlJC24gEdqwECNk8f0ku6EfxsN4GzFVhFk4mCwnHigDIdA66D7UMVwLP/1TLE+En7EI17hnJ+ePkNpkdKZypwsBSxzYh5luqeEZJ/gP4yxHyqugOsUbcdS/i4qd8vmLdRzfv9M/mxR0ykCOWYClCDAV/8z/uMpIX7wP+0Kq/99cRFq+TsXelHqcgvHdfwzEIDI6pv0G2s2a2usMw+vpdfwmlDfjk5Ota478nMS6bYaHh28t3mtfjJ1hZvC3MeB1hGN/HB4eGji6byT8rif2ZhKiDJTUSCOEl6YmVJ8aE6qR2uD+o2bGOvILwDEZ4v6fq8nJxfX8ecodZwRQGXkWsNdA+FA5XeD83wYadLL58tsNLITykD4c5rylmxr4CyksA2E33/ynUDTjC8Wj4GeufwsFfxYZaPc/SpJIEIiLa0MLTyz/hg81aAn8F2TILIDt2yU939r/TgisgS9hwXDHJED5m/FFqYTUsDTi767ZY3ZROZQJKoIBAoiKKw4RWyTBF+nEI2PpfoMGwAQ70jRWSq5Q8R5eq9++lIdyI1PFX+JkA9OvtXhoqt9LCZLhi6pw+D24+r9OPR4a5PEgnt8pBaLpu5eZxTYs4xmS69xmpY7zoZK66sYEAtBgEo8GS1Iggk8OeNrZwrfXm9LvCyKnmnGzUGfDJqrRZ8IFqha4ALJJ0QD8+JQ/3PjEuSYfILVO4Sqidsbfh4lBaSIAHYWbX5zH89NPxgtxf1RdFnkuCZ6TvhQOqf9LU68n4PGbY/SpGq0NGwTpaZzw++30QAJhh7RfC8aGkuvq2b19LGTd/dNTre9StdxD4HY04xtJUU1fR7Qh1wMDv01Wa8SlRe+0HIf2HXIyyPjGVH29/PE9fvJSWQlyoFTU9/nuy8e0NnXkY04zoMoZPqaeLyLpMknsHVxVnRjdans2y69Q6MCzOhfbgaTAa50yF1jowvgU6znl+spDTgwYf7QnJqReuhyi989kL/ySu+2EiqExnY5k9i98WgqGVDtutlBSx481Qp29KbJ46TVTvVSBOr5Byhv9pJ4bWOr+F6PPACBNQlEUMqnA1VdKflf+E5esqEyDkoy3r+8UKxjVmJTe/8ihOVnsW/dt0tO8oys0DgAO5heLHXA4WulRmoFnv2y9KWvyUXgUFAh741d2yvZQ3nhtQD0elYbcx7eEgfQqyBfmVTUAEf+yyuNjNb5XWtc2HDBxNBWJf6BpMI7AYrjUUhYxvs3fKGkS0nkV/nyMpgAcggTzX1QVi0U1kM4Hbju6sQGXFliHvzQDNEEFgW49d1tgX6gzcFBJXNbjx4/551PZ0qxKYsXYUFQ0J/CqDcKgGODClYFKU0J0Pg1QGLtykt6eHCXYyBMjudrqtZuV+6/LFtYd/OUONtdRxHQrB+kiPkTD/FvNG8YHjHvPQflA+cI/ecXY1rzPU5tftFpZesaJv5raGf94KLLtm+jvMcbLzf6mTKQ17Djyjhe6jIjcyhswlNBj555QCjqTm/wqZxFBNQiGgL+0+4uQTmyr5DrIcjYafJeB2HRbWYxDPQX9yv06fEmYHs+6V240/W+se/XWEYJKmzRfaEGBgdMKz7nbMvF9NU+S/4FzP79tfPX2z87IDKcBw56o1EMpfiRWBhMyZIUhvLFE3v9/Y9yy8FcfDSkZJMyQOygmP5Jx4iA7GnI4gJ0C3nt2ObwbXwSAu+u8NdMWeBApA1RfOrhznHCesrmO4NyNNpE61My0LGKXNgB9J+liYuCOuck9TU6x2VLrhWJ3iiwN9mWbviP7/opnv3Y+KL9k5f3AlSPtf2B99HRTbwf7GKazXCCllT6lkbfhQ9NNFADj9lKlVdaDGIE34vS6eVifPbv3YHUM2ocCPSMl9RR+PoK4Lm7/cRtpeWd3uBC2uzqW/PDjo8s6H67pQa5n4vNYav6pNq0NPyTtQyBFXtU7HiPDyje8vixSH52GChueiN1+ci78GyzHPC4a0pKYoSs4lYmIiGTNJGvB5HaFY3KwKP8wq08pV/WWZiS18vaIVqcxi9SWzOnAycmIfXtlBvC4wLuWoOTjFnWGaSlaprZIJt5i0bk3d5veNicGWghdWu9pGyz6p6WiLA95VaSvKhJdHff++4PDty3B/gYHRQ4skoXkuTTgy5yNzN9NwVMi7xBPG6Gsc/xcuF7W5EvmFxkM5uNcZKRRVLsy9NxlxoEW4M3SMKrZeA1GKKdHoxGRwUoVECbTGfwrzBlJ//L69TnIPzOOOKqxCYru3128IQ/HxP+B2Net6jH/S+v/fwj8pvVXQZeStqeiZ2TN+6WRz1n7mtWFzs/nEer/NPfsL2aC5gHy56JsHy2WxfSiEEUr0nDAt+JsD03rLWdzJGD87B4uTB3ozuuI6Pkz039k+hnf64qrr69PkXKZmZqawsOmLhsfX5fWFfGWRZkcS2vrL7nyl/MBQe1y4kjXpyz2G5pbr6UAtHZXxro5hFPDKw0aMD3zDzhrmkiyMcDtoZObR/GZAlBOCIeLwvPuWCQenDt5Pfz2HKEM7CIiIsorGfPIkURiqCZxJNKJPMtGBHiXCzo/1xqUldnAMhzwTwPPRflvA49jy7UlQJxrSzZkIuUQ8q3DgTjF4jnBmPOcpG4i3br1f8HtW46vL/qKPpi7kSrwYlqIFn+hlZiUjUy8P9SvDyJVWNL7DS/IRf/2t9xfTbchImPF9pr6yYkJe4hCzEyxHTd4EUBgo7ii2/nlAKZnydKgZ88iMmrDbn64OU5zLrQn3fAtoUcsuEh8rRLyuVA1KTDPRKYYJEy0/PQ+8MU68YH2PfbCNo3M7ltdLy16OpKr1Q8+vv/41njlPELmZ6/vFDro5j3Im1h8WrAcPn2C/wwq8NOAftvwr8n0EyEhQfME6UtG3jvfppxJCu9TKb4lHxLIKgOtePQa4C3QSFr09xpqCXeCB7oP7rGxyuXw0a2ebTk0Pxh8zw3OJQO5EuJjXByyV/mM7+jzTksuthnsT13tcdn+/ZBC3ODl686HNO1tMzlW8WaPkrjAGj9lLphqzJ1dlbqMqWefmaesPynfYA+BsvwMJ2OId0FLofpCjFI8dKV/xFP0U7EwG5S2NeC+koITAa7ZmHWjMTpiWZFBNE0nNUcCicRFUmqO8stAOe2QpiEcke0H4khODlLOGWz89m98s6sHsTPYCBXFZE8scFJm6hi/AnakyW55SmNykvNr1ZP0n+1RcN8qoNCD7GEU0227xPdzFfdvOfnGNLTt3jfYq92a5SEdPZ2zmO12npdWGkiC8iVtTptc0jSRlTMOTef3lSklqA1iKOJTE8PqG6Tm8JoumzYbFvGpIt5JWSCnhAANxqzEhD0SLMcuEYzSZ3bH7AmmHZgmIQmD+5PjUB+ULmHCRDiuX8yl8eea2jJhnp8pCPsKC8iVrze5BhqTAsVTIJgBjVnS1rAYYG6Kxjen89mI76b6ghBplJOxwc6aqBAqp4PM/mEE3YJFGj0C52+U2SLqEI8kZNVjrMz1A9iVZPfzsVyORXQWqyNB7TSgve4ZPQFEbdDMuJhbjrw349Lz5vsPn9xdHf21VhFDvHtX/brs+IhrTtAJ4fuiKV++gszK5NPMCnT1fGZqktp0EaPSY9rHRu5hDMQ4YVr+QTRf/QbAKyTRomX5qFzwaURD2JHw3d+6sHUDM6CIfzR5OJKbPtTipyLcV7S1hULrYKehEEbgRfv7WRZlUkjJIeY1cdriWeWTKrfbrsuDt0OXV/IxeaN+bnTjfu2rn53FvQKMz+ajEHx76MmwOlZCznOGgU+HSYEueCd1sVfnJT4PDCBGP3YU2ztn3J1nCuolz8kya7vTTzrB/g+uftM/dbJ8kEcuzAnZYUHP9b95Qm2Z9Zr/Xs23tIfJRUX36STeMXHgSs+zrCHajz1pWfodVmbNdwNr+UOeDgOaqhiU67rN7WGZdSiw0bMSY7MV1PbVb9vrzXYD/XVHTNgjdvdqe719guqoGPvI9C+SaEv/kCNTlLT7Vq3Zx2NbfaNtbnBnFEE1YffYE93lHPJcAfBQ6H7nY4fedPBabo6b+MvmcwcfA3s3e+MvoDBx200iGHQqo5Zuab7xMGagbTLoTmfvGN5FdnYw/ZQpGHuXfEx4jcNr5UeLQSPMd+PSVYurjdnnU3t8bh7BjGPb0vOLzmb7J47DTHr4p05vp+N9URdY9xEPmLsR/UDJePbtcKkrmYahcFKRxFY7wsWDRcvv2vmb5w4E6mdh2yLJp50TZxMSXh/Jv+pCbqTC7V/aXwjxaCgtFXxFZocI8JElTb2BdFwNU1Hy4ZiSCI7PBc8i0uCfk+MpWi/rt7HjfafUqAA7pJNPzie0SNnlF40Ah+GCpZmO9B+E4EcTOfmV6axSOqWDRh5M0c3WsevpHUSFJDwTOuAbiRollqkP97vN3s9KrglsR0wb1L9+41O2/Od/AXCtG42NN+2MRP4iSdj/1DKpys7OLpCGnX2TUxzuq8RA5ZriNR0ovbKB0vBz+uJs2+jJSqJ1nPygn2E6K1hNFq5OG/nM5j6GD290+BH9UZa7gdH8u/wH5TW3cArizjWkaQsR+GeK6b47X56fLOnECfvS4odE02LVSM2/Chsf4iJn5n8V5Wn3BAUdqD0PD41rcZAnSIBdnXNBc+dGBbAxlDQBRBsAMAIVeDfBj7GUI09p/NNB9tdAlJoq0PySYZKwBu5voohABiB9+9wI0MNoE3RbIFWwcFlczO/hzA0yKMoCDOY6mMGLRiyWXya6pvX3eObvA8hu+f88QXm/ZjSngIgEO8/kxrnMo8VJ9QjVt/oCic9CI+nlMMrYtQ2RtfdYDb9IRvMGn32/o+GNdpHnYwvREaYuyvB/5zj7C/GkTug3fikCm3d+GoTAcXYyRpMVyg4EAZJ4ozJWh9xno5NY+KC7zpU/6szJZWzGmDPYmbD2ROCQNfGS41kXzNdI5HCTVbXrmwdEmLykqSqAEtBYjODyE3+FddapgKhMoQxGqXlcuXO48SVFfpvGrHLis+8P7TR+5k0CkIc2BXKsgA4vTtsFkEdn+X0d942jcfXsp/kcEorH66X+KwZTBBz+Xxoyh2c4eaZqy/zNHp8WkX9VgOxclFGzc/30RTB2ivxk+41/L26fSRN1ush/E3MzLyK6Msyo+mhApn8N6j/qf2itVGq7M5uzW+rPrA9A4lM80/QKwstBevCngqssc0xaMpw15FyXbxVB045/Vny+V9uTvDiT2Rrnpwc5wb+bicdLE/98eBDByVbt05UHuhWdpCoujAfeh1vJ7h45WV8LCoJ0IdcDApVtbW1vvaJkgYVVzFKa8V0uvZp4YkZR7Xn5CMJAok8CyOTEaeRbRpJqbH0BGx2JpODmR1kEPkjqWOvbL21RAuEWT88iWRIcexvgUO/gKKMTHfAfaocNBujeesY1WTE0mDrh97884P1L65q01cS6Gbsb2A8UadKwl2noev1x3gfbDeHERHwMrtz2RUPqFY8rdnaw+YU+CQmTh/Ma8NMrnogathptxjO6/3Nyq47uPy5ujXGq9+iDf0yX+u3cCpIfzuENpJlyeNMF7jv8P9u3vqDf9m3CHgGeXcVGS/Pu6fQSzCftEQCQ5hiIiSTVerODYNCRxP5A7qSl0VtQqKdRpMJR6x6+bV2XrMFgKrEbIjYynPng2YiXL0/UWzDZOchjEBC+7WBi1P1S0y8AWIKl2JugOYb/bRBLpdHWaOXLpwG4e0KmXrHnHaS3carf0wa9GxATEdIEjgNoH3FWxBe88zIMCtVd6CgU/I1TmW0pEXY+fb/9h9sgjx+/xwHpIMhhJr69ur0aj/bzQueMi2knfXk6F3VpbQghQcWSIgQ0dwr/deunwP5Do8bGKw0CeYKWzClA8b3gnENpyWxPG/v8Xa55E7Jrkm6+GHbyzpzn5pwu3j0o7L/+UpFWmRKcxFOT6++/Kx//9U71dRHM+WtZ3YdGd8O2mQACtf9lMY6b0JLZFw8ER9/Tb+hBfGgfVZpwMLijZSp3X9jWbLmIh3ecBSuRRzcdLvOduzyTbUp9j59VLBwRfKlGW0AM1hcpzjI8pseGXmZ89ABGxP3QQt+FzvB45Fvi6/TF5sOYKrO049dm/FejBWNfkSs8LDZiiHBXaq1/eZVn6CUzG4qsWD7lytS1YMr6aqQ/Je6qrEpvoBO9ZqPhVEwV7OawTBiV0eY/9T/Em2XsxX/cwfrC8kZc4gi91+WJtqZxVQTOe4M0mmwIhyQE138uv3A+XQ7csllTupoWXKJdlsRIFBmMojfwklIi03c9pqPuOkLNwWLjdTVSC1NkfXGAgLox9D0xSE/pw2L7Zz0dT8vOGpaY9FOaT2eOrumVSLHY4bzMG7A9BbzxGnm1+zI+ovHFDuxU1Kl9brFjuVAYc/6WXse8U3HsEYgeq3b9PkKqh3qT2o13yocKgXfvrRvJae4FRd9qqsnRrkhxQnDngxNIlefxll6JE1wGQR/2zKANwnGD22gMb00kkB7McEiS5kfRu5Urpv6Gu5i9mT1ULNJMMBCKFyQjrT3aJTDFe8CKzGbrLSnh5yi3kMGQfLBgQhZrK3Gdayz5rChdLGcOH/EwOsK3yOPw4vuXD7XdrODVLm1qicMj/pfpcsE0cXYneksd8RTqV2j46ncITJJG8AM9mhgmE6bfdfz+kjw6vDJksQMubWeCekalwvidT752JrXVfHrabzfMQ4QskMWW9Wosr6hoPviDmi0rM4I4H3x1Dm+FjuVyDsdClOQrmQpjXzfiQsD10DP8Exl3o1Ui9L/WwfEPhycqoWbhsnFIM5Ng6NAUJYpqlq2CtiDYDu+3dEd4IRFsdI/IWWK+BszapRwFdmX5dDOxyXuScp5cFMslWokeQYP91ybOelTbOzA5KGIJ9dXRryIP3m9z477+f7y9BVSUbdc2TCiglIoIMpQiIC0NI6WIIN0l3UjDANIgJTkoCEh3h5R0l9TQ0tLdzVDzXYN4309+77v+tb5/rVm6OK/pfZ57H8fex96z2pXl7cgoyfF4rIHFI/jmVe7XNhuA1WlcorWgVBYJHoulxxqmliuVWK7Wlwwd4axEb+pu4XAJbKngu2Hd2TCVBwh5hA4K/nn8/gFvEeZlQv2OU2DpEpEvHarGi6fopj88j/I4zDl883DEcAwUzN4NP1LpjJS3Pd+jYLPF0r3rSXRGOVSYU/yQryNanpGB2DApEji6Q7bJop1CJbozBVaXoo7pchbWvR59H19ssXhIXuqzUKxb7udie3K8e+Ub44CzwuRZasZ2Mw/nB461TGpXOogb030X383SpMcQIUZTa9vKcCo+3WaBVvbz5cM+BhSm0tmy8U4h+mczdaPLD5dbBYa/bA1ejk78Rm2dQs3sSfz5s7zK2EAAnmlGcYdr863pBX63fUzOxmZ6kJ562CogwLdKZ5iUTPxRxxu1pYdqtiQtzafM95mAj/lN6tUBqpfTSejJNNenFGAMQJg5fsnzANn+jowy3BXP6AEyIVqBu0Hp+vE6f/CItw16viHGY3GMOmOt14w3FCUxgypADEQjeaeVmxh4mQbNNA6qA15HnuTMjSxNMzegWOx5UMRcq8DH/Wn5Il/bYtuV7P2EFKbY73cunQj+CEgWZk8+n7x2Cja9suvfyGsy9XfuoDJM6xbtn+zBGAoUizAKWdedtW7KlCXy0cagvpUppf0dmT24azVJD7a1sbG5I33MR9/UqvWYIBM9hGYTnfJ3Qh/JmfB4Cu/PW4X7z6PFP8RGObcJQOp2YdiiUUeoM+ajpbyu7esS7oez5jo6L0UXMHCzC55uE62vtX4IfWyx80cwkk3OBYQZ8IyChsbzzr9jy05zQIdH2L5eYn5/vmUcgkEFc3r/n6QjkjFI6cjO/O1lSfEr9cjb7L6of1eODPmjSG0PyCCVHtZAbEGcf1puJz8B8RqkCeyMzOTlApgoEdUNoyGRiD8XCCxvjCiPuqbVbhiIFE/yIHoMG7dxZ7ufb2OlqFhqvSwRZEyXWs9d+d57Jteu1fiWz65KMQMJ3pHYyMIpv5yt7yPOTBQiKROzPyYdVvJ4Mf+WzNl8Y/1pmsWwXNVxA++QBcL1tDmOf9URsS5dP7EJGlduFUTIUbxVf/3QBIazgpbEW17wU+qhGCrGt47em72aM3w7Ki+cbo4d9MZ+2IZPNLgttDgp9j16KVibdXSBUvH4mUUc4i2+BeQduWESNbIcETm6pm5smPSJWChOkPxkQrEoYC++fLHW5m3PgnZdUcySa0u95QT+XI1M14oHgbjpTnj10N1eRF4ejPeqgbAReA+U0aAvSISIhfy+PstFVcSeV9hcRoVk5AUz0kp+XtmOcpVJ5Duiyo8jIP28anw69gh7Shvx/LDApd5jECxgPiKTmKCAoc6GXoUvJLQ+0aI0vrAbVMtBc/fU7w4+c7B44N0vC1hkH8iYHYjf88TJ1zB1rgug6CTTeNmFymT33AHfQLafpSG3f+CtKHNKPnxdZ6wv4pIni1Apd8k4BLtwwqNtdA9+aGeab6JF8od4BrWCGZcT9MRpj0rfn4gZc3KO7xIaD73Zx6vot2oaDQZlPmHWYFLrU4nhBc88kjyvMWjO+O3uLiAUnuPKGIkz8RIa2v4yqwIo5Ed9csAiyUvgH/X2hgYf1IjYg7AJlshhk6KAsVahwjOUoWc/Zc9QACs7wY7qRmfxY89rnHty7V9eSFf9GrNacDaerSzysVDYFkChOLdJx65BVGpKhVcnBmVGVidqnt4/WSsO+PQ9oGzJVeYYNcviMxAmpzCTImV+69MRw88a1BMOxQDk+5LNs4MC/1zXykPjyMN5Z9CexmPEp9Jk5rzidKfxy1CyMRfZTtf3vMOb2hFyr40YTo8uTqbEBd1svd8W9PbfNHcipAdMyIcsvvBRXImE1jpQ5s53u4DnXDRaGhhlaAtMdjd7oDjAtzxb28GSR71c8EIn52nuWMF5Q0mRWMNa8pHVLoeGO+34hU9Kwx49siRWhbNC6FkVt/IQ2YbxAmOSx8N8e10DtYEFEcXs8ibuXDdFeJtxzcX80ebFvt67no2PF4sSZpqv2rXqCQ43EimEVYpMQV8K8lJqvkvOqloi+49QWdCrUHaKk9uRzZ3NyMMT9klYLxwRzdBWIXM2QQ+eXZLW+wlPHPE5sd/4ShjluBrcY6nYyJR4FPS88GqYi5GQFeoPUlM84C1h6TRM/RoZmy10NXMyd9H2OR2mZYopDtCq2XCZTokdUXqugqUup9t+MoXe/utbhG2gusIIv51r1+ob5GORSerGyaxXwDdFjGFd6bFSIqyPfx5S65qHQPcIxitcmXJRkV2FKpXaWvYoXGipnZGrRbzd0FjixscHvzS/CGJAWQ4JK6R8b79ahOyGe3bnTbCN5HmC2P0FDXIQ79zjylHHN5jP3hDrCHc2d6TvNeiEixajklfwhfl+AgXH4YEELDrtnChLPtk5Nc1xEPOAaPDN4dQ/GrO1N6q5PdjMffp1x6eC6HZHMwpeEE92IODJPgwovF95OcYBotZcGZAY7SFqTuomhZmTEyyfWTfGyp+FKY9ItjAdj7Yv2udxROgiUhP8FxXqw/fmYMJLUvcwFr7sjf6ghllAKGVk2U6vtLojlKDX5G1+jXde+zf2h/HgJgIUgU+DX8qJGTrLgOX/hFvDhksy4RXobktK7IeTCK5g1qL1gMBuxxZPW7AseQBIiJrsae8Yyur7THAUctBa1RXPuvnNNuhRh98T0dgmvHbutcezrxjpo56QfibdvBl8RbA6MA4+BD1zlVjUFV4TvSdd4vTxb5FEg73bScuCR+M0S1JgPmEEWUfGsp/A8jmJI0RJ6UnT/bOPRMbb+u+Y511YX+BNr1Rp03yxmny9rmyece7wkGM+kRo2/CEfsjwpoKis3PkT3YJmk2W2VS7zHuzWm1AudGmWo1f8n/F+j6Jr6NC9YbTTCnD/6KKoK1aGj5qZ1vw6lNUfBAIuN9p8Ca5jLPPHJXrBQzUa0f7kzqfw5mkJ4VRbJyEMWn8mAOOByOJww0LxEr8QurLIF6MCfw010+BnohPcpVG2tP0CoHTIk1iHRWQTF3OfHA/ZPekITGnTpCXzpFRr25kbFBrkyDZcrid11CBiIV3M2Jkb1MixJclE7u2LuyVeRGcaB3MiimT5GRvd5txD+IkvOmT8uQsTa8Yrbz8zPFA/XW0lmxS57YrGsx09DiXnMREoG+MNr9qMFT1mnuxyY04q/kxhZ35KxzMtx1xz+GxeipjMXeO9xXW1+NsIMbKT9/DRzNXED4jYldSm7GoOA7LsJId/VSSXQ9GRQRZsP+JU4J/zJruPLmBGXSTyI9RyZQXlimiblw6VDs59p6uN8UiGYFSPxhc9+U4H99zmkFUdoYdZgkiZhKCjuvxDE1acmWJE2F3hyQ4KGfh8giDzGe3zxbIQHhXZj84UakLvWwuPk+sNJuTdOoUkjJGV9Cn00stOzwZP+CJUREj5SVftBoUp2W7tMktW8T3t/dfbqnU9sfr+zklXcVyoGSVcMitHASOGAD36AJPX5vKkBc+DRfuSTXfPKrg0H3wScNL2rZDragbJ/MHCEizlp8xDEz8cnQFo0H756ISgGkdwrfutfO2qL+6sk054QtO74YzqSlC2O+2eWqgfU4WR+msU4iyHX912qysSpnaLRutEZ0/CK6Xd7OIya1I51ntOuBHUg35r8hjqYPRPAwPy3HForyjzUtJ8tq0XFxcto8Xu2Bgacro4YHbcVUYlw9r/O1sRdJWteG9VE2CM22Qa7t8cYA3yWbS2Wb0Wnr4Qw4kCYmes8LbRshHTe3zxK/HxzMvHybHNAY7FqXUaaH8Xop5Z81CDdjRXgMvU18rjJnxaAOiyrd2jL2TMnNzQcD7faDe1U5uK3Z2568aIJH7tW+jsptDBcQJDpYG2EJCbTUzVzdVntkipAak2F1eDcLQA5ZinjKcdMiX/YXzS2lfmH5rPBKcP5oBtFIC983GHr8MjndX3P2Uw4vQe0//JYOzt7Z2/KrJDdoHfd7kl8lud3N6qJX5VmiqzIhO+bDEI98/F/KNPRiGg9d9v6r/QE6o9nIVco02tqaXoP9qRJP9Brz0RkXMG3X45Wzs73CfTMjC/K+mI2GHDP+jQPOsd8C4CeHrQPRY0QN0dnckjBU4hjx9eOt6eRQ6Y0qGOe2HLjqO7Iyz4piXA9ldLxxUCjcP9jUCF3HhAHWtdcbj/kz7t+IUgAObUFs+LHclPXLbOl9FoEW1JZnDutyZHZ4aKZdQn73T1Z2RBnUJ2V+LEMsMyrc4GvGczbjP4NpfzBaRRB1LuDiihsYep/ir9RL0FmVIM9qjfUidN3Q0b/FARZi5c1JxQ4GRyYGM8g+PbOpw0G1sqM4dUQRRTfU9tbnyfGhvrfP5h+VGFBr0CUozElqSCArzxSAWkcHj054jkZcUGtIMiMTsM47KIIf5CjOjdIPvJ7IBiQzZxhBdEd2bSrpea78TqMzMKDsFpNk7pM4ydoCtAzQV/OPHR5pIp+F1+AVvYPRwdVOshOdHWx5MHD4TEOYI+yGKQRKHWDG8LWlwK+5vzt7gV9WZ1yS9H1evOrZdz9bUYK8uVo5dkTWz3qF1Nhjtls176alz1aMKf3Pmp2+lq5tBagstWgHQ8eHi43P3C+d3Ed6ZckC6uxihJU2XFDnWE6EfMJoYPn9P0e2I4h5KF/ZPmLLEIXxpV7DzpwHhZlBbGQQf1IdhNnnmEQyTP34C+YW1Bg8yT4D6yeTR74/WTgp4ncZ9xV6LWHq+GdTa3VdmvZ2lWn3CPzda8+YxbquZ9Wmg3ZaVcrVZqeyj+TmHID2/U2nPb0OOWMdspmzGVmaFV51tXFKX7TKlM1R6WK7n7HatCRh3cFntiXHd+Be9+ALEDoWwSCTqt4Sq3lSJOPv9W6MBEyiyZy/xDEDSh8cgRfbVH411TpS27uTEchp337YhPVluqMTopR7/pojM4zmWlhnp/+0tE7CMST2zIBTVoHnv1gQk/OUK06VZwQyc1smg1MhMngexvSp4Yy/Nlf3TqS7BOZo9xU9nfZ3BsBXlnAhXY0qeTv+YcouADUfietHESPmOdzuGHFc92jFeOtQYjKz9zdGDUIFgakf0DVC8HSmRi8WrK4cuxsFO6CDvXF8yzGt/N1VYAR/Bttetz6FiJcSrY/B0QJWfj+JzYgu6+NNUAPuo/9QUTfYa62KT04pHyPAObj37LEkCXre4boDJMEnnWVyIs42zQR1W6XeUWywGrFwnL/6KIIf8RVUOKmM0ACzlzFbJde6xXz4LMlfcNEDc89jyYPY38uaE+/QDWwLscJ3xxephavVXVZfiSAs18Fi95+HrMkslDnEnrZlxIN2t7C32uCoqe4ooXW58r8Qx5i4VAmWR8zMwFzWpD5YH7xyFdDv4wa/zUTD3cmnDBnfRDNEp1tdk6Ohwr7OZp5eXDZOQ8mJ/9UYP1b9Q+yi0TftPuPQmk8Tgs2wePDtUK4MSxZ7hLU1ycsdMXimvGmayYKN3s0azIk5uoqXKlYjp3R80fl/2V5QBqItbSGy4TRu0503UTzpKJgX2Iw5KK36N2Y12vjvhV+e1ZNaFUzAFvvnp7fRmnacuQwidQ2f38NBlBeGkCCUmdKsSMGlnA48JUyjg5n28XU2dBSahPY7k48EGFHI0ZV3sUXhh/0SZTmJ/8zjRUuSkA6bcfycQgxw+M6at+tzZTc7h5Oi0oBj3fNVqe/fa9tKCmM5yBEy+3i1u2MAVriu38jfNlf3W847BxX879+OXv5uNCvZO2DzfdT8lCXmutFmff72DKOEMcPCOrfcusebGyHmJU/Qjny5e8jFh3d7ZNh45p0Es3eUhB4hl48f65TwLFIdsUeIFnzWaFNyJQUESrchnnXfVIEcda6/JKg+Zt0z7tX1LiF6CXHY2YFJLNmPhJqVDBqkdDtCLLZuUUMpgU+0Gj5+xzJelv+ieixQdOH3qKCjJkZmZKJTU73v6lkMTasWA+ZTe1OrTfYFlat2U2XKkX8thsuo5fidnN3xYP7EtTwxoZVx6joJaIo15S+WGct95cnvrRd/cOD4EkEX5Ysm8h1FaX6Vch7dT+OKmz+Y+tbDS756Hia8PZTPmjC66b3KX936aKt5M5Dgo8+rDrBYN+dtcHg494vyZUzOJkSg2yWmTMOcSbeVI4aXUqz/hdGg/chLpK9k6/J1WyOP3Ehg2Wa/1pAcOQyD48W9ubvghqQ1iRkdW/fXvqICQ4vt7Ya7rRk5nxpJCBhl8KPzQx8vMoThZRHWNyfS1/kPRGTEIt9uHHGbcJo3kMV/CuT4Z7XeJ7iqqO7SjU0wop4x08G3O0Rgv38Cwqhi4nuo+0m98nLLfKy95y9d+NfwdWMzj2KrVsE6Qrzh9w8+DhonvDuCw3u+XXMnsDg7Zm4G0azEYv4ZVLw60C5e5LIv4T8Zv5cRCyTY9+MRX8920uMK8XTXdfh4Ib8JoEwxpc5mrFGVGKwZsvujJw09KnXaW0XS8xas/cWEN6em0mOgPKDSdInRtbQ0vpOt+H794q+RRIAUs8X6Ad+t5DrsgG5yg3cRpjO9XdfVxVoeLi6E2hM8N17+Uj81G0aSzxvsbdkL2ZgSHCIxyj1BKVQi08q3wTJ7WIJZnz54eVawQqBqeRBugnM9wBmwYH84MymFvbAiWRJnOtmVVR1BP2m4sLkysrgfIiTlPuQc7EkgAWGL0S1Ln7VtnFSYB5nLS0EiYw2yddUXEvO2+pW/O8uGxpWRXSClbjcuOVKzx2CcPVIa3M80n5etPyyIl9NxecyKZeblvk3N5JSdYW+kRmoysj7phdDNMaXG3Ac5xGv7ofQLC6tO4hpTJ16tWFfVrO/cJ/12rbyY7UoPlHj2r4KEswaeIL8OTlfU0pt9DawKUCNfGpcYhdEBxOd5oSpV3oDvPrMH5OBIdzn45D193Twt2L3T6N7N9QR/ScRRDmf2w2OGU7FZe9Pde04K8hziWw1a57I1T3Hj09/5MPfL/YMvLgr3x5yaxv8z1r5o1WES19h63p2Cxlz5xtPz6vA549IVjpnOZ/fXb8NCnZvzSdnbkqvA1aORJ/kkxl85KLeHe1bJbfKtSLRVBw2vXw4mBp7DtctOoyQ3ukcWx/h2D1Z5D/QF5rbpb92PMiZ2+q2JuePgqf9fs8U+e16OUPT+GQi3z8kexeAeMpcRWyEe8ljwtCqXh4fImb+q0euA0B2UD+j0UvM7un3mQFU9RxAdieOUrsK36CwVne+z5dWyns+CuvCgXd0TUdpvQhk1/qk7tOtxeW0F/jobPp2vb29OSnjsSFWdNEanOe6Z82T2uxt+onCnvYU5D9zJXaeI6Hw9WBgSG708GEokr2MN+MxrDXMtZ907/O5Fkr6tSlb9nsgpc2auo0iunVsO1+796k3TcDs5uZWm4B7YeLj5/Xf/xuqSYzgsufSE5xTBcCQgSUlKSU9B0F53+jno1YCHDjvVsfghsYOVgvGCuUrigZjmANVdd0kSejEHPQ/vUorp8ebL+danriYped25z8xt1ikNju7gNi9+/xqJ/xeE+sna27E7M96o8gh7uCyfrSdXFG7qdxXhvuKVK7REbMHx574JRlmLtOyoiuaUzn50tMyuWKcU2drmuPPDRQF9nLmrAnat8TOFpwG0KEPizaDnHrA5czTdlOWETMPoR46cKNPaiMIUne1eWO7k2Kz1m0xl6PuTGnBjpxqqCRgIz9SlLqFo8m6IpmVTwuBC4sFPGrts6DVkcbiRSMIOhTh00GztvltNxcr0a00vS1q4Ybv37S4a+3AJ3C69QOtkbemhqaUtvTm4xzj5Rtu925OD0SsJEZ1EOxKFfQRdtrOwN79JOAD8JCFTW15khaZ5KyZ4k/Wk4IOflurg6Zbafm7JXHaSOy7Venh86rQNtpZxSXVBdH5Y79myNxUl83jyGveklsxF0tp0oTfp05sLqd77utRG4KZhbbYZs2R8rf1vrOffMAi/9ozmxrIZTiJr4gV9wpHtngfaGXY5JxmC/pzKKb8rNdjrdscuKZyLN1Jtg6zYf74Ya1MacOsFRQAfHrs3gePmbtVHfTL2/4HaY4YU+OSrYEF2Tu21+Gfe8m25xEex+55D7GEV6PnaLEdThpSHTcd+4vYfrVKXcVys3/YjnhknpRJaqvZTexf2aZnFONBFFiT34o1hCgHSdBUz8neXnW/+GmvukgVMJ0FgpbqxDmHs/OMn2A9ilPb61OwoZUcKe3abIlXUIjEuOMbWvPHdJrZZnxiAKazHqMl1rQriEFGzaRwv3RwBKPaSTDv333tCsE3WT3eHrelvhwZayrcnzE8/V68qnhcvIYfyj/1wLei2fOZ4tEpLia68npukEGDyCQQvdGcfsx/Pt0G68JYQuJ77KrDlw6WdZeaDOripzxqLyxqnQgPhy9SFH0ppt705Oeg+LyM0/tvmk5bS1Yj2ez0kw00GB5MFizIEXOo07q6VjFXGH9JRimkysWaudLuaUo7gpfyy3M6rKpT8bGM5F5Trgk7x4LU3w/ULgSQ1WVeUMAHY2FTe1kKgo2YupztrG0fbwN41aTn9tHDv6cNSG0T9A2bcm9byzI6w22jJPkVfsKOfwBNZkzw2/9lrb+Yj2rJPpdono9okbjMWyuh5W79bxIZfuX4qcPs9mDmfUy5fE2PiXPOpLYL3dIelQ5ut6GRIAT2I1D9xvk16RlHvpOreazwjaXteKfUjFECXSSiFGEOj1J5ylFDcdQQfmQazNuWK4wBJeIihg8izrYXNDOINXY74MJF7pKZvNp6ujxwwPb6dbO0yXW7Lt1AqVBkybWbw0X+trZPUw4XHAtsFQyYaf7YoHgKDUFztlIGXBaAmaCJIwhZQCVnpdlNjn4fEpfBV3+7kGAd0Fq1n4UjM4+SnW1O0dPBSMmGjWW1j+WOOWLE59PmSvcNrdG6ustYhtf+uHBc3i2fRnjMPzzDq8VU0RxWqkGEQrdFN9taWT3iMZ7xQuOd9lLDSgpFhzUUFTmGhYnZRJphRu5yuyqMwRRWhU+ErZxBPBWEudDr2crMwmSkRAhCqRmdeUm0WvqMSGN9EqOrdZt/uuZ0x8ifteFNRJ8yPiE2qmvepWDtON+l50aOV3RkJ2qAe+NVm0Z38+A85FqmI2Ght/jpr2DUiMhSfxJ/hI7Qc94GqlOVaPrGIfR6/zrkfUXuj1r1zc8cuR8Pqvon3FPbvAs/5a4nASVbvsbNAfH6p6Xsy49/ov670ud3F+7Z8RcvNrFhZz8Jq+qCirwCciq4LkMWGpBvUG4G0R3tDQPSWKXOEMRIrpWpRZFIetLpfieUuQF5Cvd0Ttpb6DX9SXuMxTKeu25PyWm5+/+ki/4E/xn+YIX6v3NNlySz3gnEdqTVsFYyyh/xsl4fvnXcTJpt/5pnMw/iU4/hViNn1RQSNfeRIGOpP1pcgCYMAS1r4lEIzf3z5InsquM4h+6ypDisIbXI/QMf5aqUGQp/uqJUKdE9q15hRQoKf9ZeoxVz2YTevLqB25UWGKXPpWoyc3LD1ciMByag8digjAJriTt283FPQQ+6It357zqkdmeBROFhU8vy3gEqEENzItFn5ppK3jejb7BgH1ztVv0klsTvUe0o72jc0oNoUH+gIi/wChWM20nGtXy23t4IN4TLG+FYlTmNL7PpMTIn6zgclMrflTmxFbx7AzlaszfYbLBjI8MZ5tGy25YsPPoihB5vcD9XQetsEe4lj96gdejgo4wNziFp6E78eK8JooSLfb6PXga2hh+YmpuDUaKbWC0i6z2b1qM6yYi8RXwZ0p3BWF+s4MzR2cRa3pUouQ6uOLy/E9AxOdNa2HmKC6qID+rP11o5DCWtL0GyzX5aWdPiY7QJ1HPoluVo6vELUJPJXioK7/uNM3cwCuo+WTn9CTuaZCg9b2IJ2QYorJOpV8jkD+XkTxNA6VQ+dQSHjCteZCAoFzeG/1RZIoqWpWGbs2Dg8H8929q6O88cScBxeGF0xblwZk7ZolP7JvGPJmlmceEnDpQqck6VFHEiCX5m4cIDdN7AoPJXE+pA7BxSE6wg8IZ8DUGDChnOWx3yCMkE57tGcziNHnuf94WO3y2+Inl6dgaFr2V6+Ng3ACJ0Uhg5yFbbr81mvLgz0kU6y58IwQeQ8kLo8KPSttVuXo3t/wT+SqWbTnlv/5DUw0nsFvE1pcUdv7qNbpA6uUv3g/K/7Wizp6E7VlqKXCz8K8RfVBkpR3KzZT+f30i5DziNi7tFyev37y5lZtdUzM5tfP4qYKamlrYD35oD1aSN3ZtGETo7m/fU/4JaoC+7cYNUr+h9OOqY033nzrWfk3Bfvueaqf9pWgK+Ngn6w9msX0UeLS/OkJPbW9ddax1iv1bx9rvWff1WN4EfClQ39Uuef4ZlxNs5CF9jN8e5UOw2IjelpD0ZDVs4e9e2IZQ9/sZ5EaTOPP1IaxEXu15eYz9/8UBEe+rSL5qbW0t2nbHRA4DTkh86hA66AWLi4/v+Ywn8s/DERITewkdMU3o8I3R48do2sIHH9mgLvumLi8sfLOf/avQ3cYDIo5+5mn7T9PsQz/HxoLOthwb/sxHWPh7PoI24f37J0osDnh/T0hQ/YcJCcV/l7n/mpDwn2vcb1F23kv93Xamj0yW5zWWVf3V24xU1J16sfXhPAyzHORspjf7cjUj7WHfzQa8exljX7wgOjM+zaf9kRCUHGKkNL4DuZ3C5Qr9vngtsH87EKQYtM8kfYNIAENP3n9cmFAOQaoVtZDF4tpxiX56dEb09Qceoi4Je3adRnbZslfdcMshSQqZ0iGPYbLIekEaJrpmtUO6gNvpEGrQM1afsSWjSPmJbunSbvjlxcZYqfPhGtjJZo67Ho475bJdW8d2NXiJyRPVuO9jOp1CRkuoN17c+5G/muoDUkPBIKq/E2I2Q3K0r7/2UI1Zcw3p/Iefi/JB/uTCrwobKuTPRUH0alNhWih5eLcyS/pvviV/aVzUtfQJqRlyHUh7k2zuiblzNyLnZ2/78AXYw0XT2kPqU35uigUp77RCar8XOibGV+teCQ2eoJAsgmVScX8D9Lpyg47Pslxp7rcXYX4myeaMb78zQgCz/cmLlRz6fPvT1PNY+t/TYnLIcXnBpZ6cocpyD02UcQxKV98QwpZ48u+ZDYtvr1DLpyDlm/+gchN5aCJCCEt3txzhtXOiFGFI+oace3KTuC8Rh64P6/aLh6NyGN/q3+u80zxAbMZbBGu2O01D0xg+V08S8dq8e5oVQ/Mtj/jDfcVQXG6sBY7b4/dfmpYKfSJqfiF74C13wySHnlm8lO7O/RxUWulBGqbADIrTz0nwgaEX/U+y0GUfWARC7pfniD//4X+Dw+QJ3kKgRiDL56RN+JSzclSX3SXi7rKtk9O2h2dXtYvxktNlr1XCV/auKd7Ms8zEc4NMq0u+8eeEj+6gMbykRKB/e4we/0zlY8OUaFXi+9NVwts8yX5CWJM4Xuf4DT2yEol3VK3RLihuMlU0/cD2ulMxdaMYtWKK6g4CeDzyZg3CkjwluDEoCn0wvecY5EY1nstVdshnZJPHAzrI7ZcdNJJEd1NCpCRnhOw16rN8wW1PyfhVRlcUvJq3uj1e4UwcRhDY+ngFYnPfjx7POHWofDSpcA0Km8s0+YXYWivaHi/9GUXceeYe4jq9pk/OFej3NpuAorX3jn8Mj0kqx8/RLOAl2latD/MFvkFxtfR9uaT9c1gEFGN/8rxrC+YSJqHhM6tGmYnfVkpOWCaAEiasF5hHgMtQrbl2NDdwVillE5rzAvdn14o0B1+vpuHlrzmtwfxzWfRz3hMEDWa31fj5Do4lRuWPozyNh4kGlPhyXyp5uN5fcuVXjJO2jap5ur9nZM5DMdoiii3jKqtoyUvGebjwOkj6dK8GovQO8hQ3WC1XOjN5QQnbCpLDo/Oq9t08pMnrpXsIoXFqC/hnYm9R8k/qZW9n7qh4XMhTZUzOKCvl5YWd2XJLEnE9Mt4UhEP0CyWL09iJH86eswrUNiG46agrNXhkYJ/qxjnXPm14g4kUyNB7gcpL6d7z+IKEAEX2AvFMkJ1qmGen3+AbD0iVrzpc53GOM6RJuMwGrOnhyvMLhon91H2r9lirGzeSBEQvHfzZMSdk/0WciWJQVsnsWge/bYqIcQ+fr1MQPkfxr9JRa8Z3S8Uls21xNTwfKzcgVbw+0uXMvBSj1iMpMZwiIbyV7bmL6gu24TzvbF3FjC8xUJzprvfGNz1gs6gWocpUXC+geJmjfMb8c/UD8mrjC5kf1wthuEK2lDdeXt2Uw4S4UBh4Vm7xXi/QXJT4HcPMbPNSPC6L0gYE0V2dJQ3yPlXfpGun/9bICi2aM++TWX2CFyqSpa8IE0XkraxW/mINcdUC7LpQgdtuO8RKnzk4bAfBEiQNkBNZheiHTnVFsWwv9odRZIanMi/4gO8FaKaXeAhU8D2X8eDJsquWzlyH4kczxMMwK4Nw7Z8MjOnwfUu97dSfQhDaXRRZuTFHefy+bCXva2apXE2c9GmzV5kxsOkLEuI8K2O6YH4gu1J2qPqcDrOC+t4xZYGI9cL7hC7titwU2/C4x5GRZnI9GgUdu4/iw8dEA20fjf/6pVwkja6mrR1lEYFLU/LcJcDGsHkYPT4K7LJqHlYjCiVPHzPJNk2d02Gza7Qj4jdEt8lJWKDu94d5wjAygHtYdnfEj2nElGnWKQuoRkd/fNffEFr8ApGfnz9nuRMfUPz9O5eqioCiwLvTVOGxBmiOMarh5FDKrHmYAggEok5rdNxgIINhOvnilnnZc2nwH32lZQKbLZnisqR7BdHUvF3y1GctrMybE2p+3L0oUvlgOI8ucJmvQISzUi18Xhi8z4DCJkD9x3rBSOtxzdzivl6gBGy7919tW4wqhucjhIW8Ndy+9wND5yYm/yP/6wUffanFbtp2SR5Yyu32ExNu5SwV0UqeAQWY63Qe7ildfEvzLnZoZAhhJgFzwVCpHLi//S6+fXDnKuQwmNB9JjlwIACfc2nUV0SClJoW+HNw9dQOkr9wFkbRPnjYnhErCsPU98XlhMbhtnlTSeY0TO4qkkR3kQ4quMfBYpz3ypFHwR4wr2opDHXZu+gJYkjDgWsyej6VMzTm+jSYZ9l3wK0ZsbHLM7Im40bKxHqKUsIdlfwg0ka9r+Zb4t0yQ8AEsBkGUAPkGHA8WPJwnQ878seoJodyQgxk0rRBGHyKH3iB03vrZUr1GUoWz2Q0AW1AgTlPPOPeiRwRaXs4qwrJSi+NIdcy3Pw7T3yB4PdJPluBLOPJiopQQmPvy9CbeUqgZe/LLQEUmF+uPzX3S9P+W8LDic7tXasLNAwLX/bmzSucOYqVsyo3zCtY/VtfdBtmDvY8ggwsKWsHvft5n3Npv+JB5aTPTZFV4tzeQpSMlT0ZrN8WyXsAnFllsxoCi+sFyVsG6GGvsLzuXN0In4ThNtzGc36Jd70AHHD0/90Bb/Bh/77Y9rRGuwryRm+G5m76hazwRrbdZtwbFuM6/9P9ryUvQOklSxvp+y/4l1mUgK0PsaVAgS3mgFaGVhRrAlAF46qVrg6yJsSW2TGvgidQ4H1WCmROuKiFmRvzkWBbIOABS3NUdbcgOazCFHaSU5gNq6RVq+2X0QGpIssvFlgq/RMfxI9mfe06g3e7b78vKWDNoGeJw1XRZgf8MfdSS3BIJA/g/AbV7Jj0mMlbyn7k09Fzd2nQ+LcH6r+oZeDkFUizY3Ohuok2NToaVmfCPmHsoY+mzKU6MxXiGYfrSk9xsJxueWPB9RnqZAjhu/Gn8U9uBvlVEz8FMeM9CIkBodJRWHY/xD9/CpI2Hrtfheu516iXtNZ/y8yeijIT2J2YTskQHgrGpB1FEIoxOpk4+uE3ZYs4/L3OKawfY9aHC+TYbkGwGKe197zxRCxacQVBTOB4aMxvGzcjbWwnB5xPs0CQnZ5FJWDftDqadn/IkgpKn8StP/ajAuwXiPu+Cff3AioWYF0+LC/qJzPubmU+j2jTgOXyN5Qgz/JbjD3COp9uuW5uWEl5UPzYlq0/fO7+rXrvNDaIkXscMKklFWB9S6rzCjwv+L33+S4jzTIOgohEIORq1vXxH4VGVVEg2uv5tImKHgCbgwO5OWKSJSkZZGbKBLe2eIt/XiyI1pXHkv5aSbjsgDqwXPZXfmm/hwLsnOYbW2C7pW7LHt7m2S04x42XM0m3as6mtre26hGXoKmjdvzthF9BuuVnc1YI9eNR/EuxYytEV9DU9OlO85RdL7+j8Nl8O9Sxl2+/Gzdx2Rv5uR/1JcInGTJlEn+dOW8EXWJQ6SsN9mo7LRKqTU9/0ykXfEs3wQOilY4LuzizG6q+eGg41gZfsYqhPFNYHeeJN3mVSy+dmevP2a1kA56K1q8WxMzX2ADHX/DfdThzQRUDSS/7YhMZeNCf9RSYz0OqLokFnfYMwvF5GvTzai4PBWfwlLhryM9XaB2m328uTtWfaE7/GnFb37JKuPRmJRjshZI9p1knC7p0vbT0aKvs8Nt/sRX/3LGgwz82k3h6bMx0z+byHL7EO6M61G1oPL8xVsImTY+vx/w8PlM6vhg8VbXgOqQRNBkz9zPNSSNIwxKtc/EkWtptooxw7AZ1952KcCzJX7tbqz3qlkPnRghMz2MIBTxAe1UAIehx0Lbmeg95WtkB+0R7fWIRnpF+1IdQ7LOwtNSiPT7cpz3b3t7S8oBzV1GcJsvIrP+CX+I71h9AuNBHvDq5gQcqoI2hPLn7cgbvSQOT+HFU9TScqIrFA75hT0Exdrm0XteLvNMjWWCjyYot8N/YskY7Jv21Qe68kmx6f9JDg19/Dg+fc0+Y8saSEGU7yjIKK81LnX6sHci3COQOlQXelixEngjtt0+RENdp0lc9U7gxuPNepZcp4+6mUmyMM7nmDzMk/lQgOdmW0coBAJoKVJ+J5N148GoDJkIq9zL3JzI+0FlUcObRB2bue2mOZGQkxBGfTh3lbTmtmC/mn6YO91C1p/DhVXR3jZ10jcAw4u4zbDSol2Xver3+8hb2WorfrlIkcx8Ve6RtOLW/zDSfxWm4aSIEzqcaVn+fM4os6N76UN8PQcAf02bizamn8dyJGVxdMzPvMzLBkzIKv0VsaZNa+05BQSZWrf5jLBMMw+KLn55EmY3ASu4pES6nRWDCxXarf/EspMr5cy+AlzdGVOoKN6fvYYLN9mTgT/mCtpIdiSSk+C36wMoapDDMIAoWwtwSueYUJeNdr8iiOcjPoZFSU13fo3tdPW2oXcrx0sIdAZGqU9EEHmaknHZjcCH/wJKhXC6Y6CxuOACe7pPpX3QtAxGfY6j04cCzMcTs5bYEymK883BGhjvb1x83h8aEggae5Za51iKm4NaHo7kWJdOO8LqDZbe8CC7utzJ6O4MBe14WStscfdPUSSVyJSY8wGuuTt9jcOvQVlF3HTJHIZVRgbo7slpzifNPtoubDsoLssTrk5xMLoYGpR5kTqlml7H+rDdCn3HliTxx4NJw6wFC7TvIgyX9mjioPLSwSH48eHif10JDDnCCKtocTynSe9zj59kdPbpcKdD+BJ2//Mr1gqwM4LyclU0f/tkxwsDf6Q87b9y7XjgEHpBl/CXyvz+BTpNCK8/3F7O1APTcmEJCz/kPRa8v7ALaczIS7djRzLNSiWAYfhHSczMrsVVvoPuN+YpduExtaDnImOWNGWoJYFLAXSyv3HEv0KzmetB2YJc+GMKLwaxdVxMlKhwKwIsR0gs58Nmy9/bU0Dsx9MRG3RU/4vRaAK3oYy6JZwLnX0+YjBvCRYy/kfVSq3TFh8ouz5C5mucHvxYSpdFZ3GfJDL+0fmq2esfRuHk4nOCZit3YmLV+qOWSnASp4ZR5x7R3xYPG3b5kYS6P2YLATinbQFV1IDrYVSJjgrQRBV2nr4mLy9tWMK6Bm5vmdzmu6G5/fIrIPtYCADsk7etZK8kJiP1KFYXmbA3eHC+TE4ju5jTmA/jcBwCEQkl5H+sD17W4lxZg5sy4vAoABFUVGk6Z1UtH/Pz5M2y9tjf6qZmZmcs0vlsy5OeBqH+xl4xSL1pnRbZI8xhlIRcXl1/TxBnOmbj//gstXWELnXhcFajb0hsgwFVqNEuYzqTc2o4EHF63GXMuT983LDun37FtjhdqJidl+XPhJAx2ftszWOZ/imz/X3DLGUrwglBPZWcMl7AxS5BEaJ1NizBZqu2ysya/qP9OKgwvVXS12w2l8yP8DKvv3lJgJog9FE9XmCwAwvUOhlfNPgsT3ULfUdb69sDaIG/xvGZeSnpBAMWK1XH98/zXMxSIJvKqgOax0l+vkTqJmRS2khF7/4GsfX9q2ZyanVH5Emn9ntwgBVmnXxDVMep7BhZbCs9y1W+o0e38/BqWS8gnLFRTuN87iTnDFAmnvwZaL3VeYC1YoTwwoHxMdb2E9xkA4Jhvntf9fQ8D9G+U6Hcf/iH1YcA3wCM2vkt4fUawXwKH5ubDj2+pvetJpf98xRc3aDQojVjurZEyxJvmjg/9tKvR+0+0/xPIsLkoK9vRQpxRH2x2Ws1VQI6x4KNXhAXWJPVNDmwee3rIxAX4vdSSE7qY4/pq541qrvxDzFl9ckbxuEAn05JEO3jd97pWzDpFjzUz+z6HrL76YQwHqzLpDaJ8YZv7uNec3+mbLZc449l8nPWDeNj9X3CtL7YObnm43ERyXOLx4yg8oBNmBz8ISux+FBASCsgnOC0ipE9TlQA0BTg0Q0lCKz05cPa8n0gOD1lBQsk2JwglfRD+1Ob7Tlv9GhTa5BocNfhmK9J33sgCiCit/0D5d0z632g7egSQABiNpEN3OGMYyaO+eMD89Tveg0EF5Hj395AByRJPlwj0ThPA7iqMQficGaVyhwDl18zTB9MCZy3DUj36h2goKd00Xh17Vo5UTwj9TvaDEBKQOHqfExkYYvIMQwAisFLoqpOaEVL8Yu9TA31AvLRK7/vh4NyB1TbSB0+SgHgxQBFvmawsEP0xhzyekTEnBHAh8827DgLvLG8o6gu3BkjvSQNR5mM46DBLZOgKRfL2p5QYJTTVtGKSqFmpm/M8WNAys6xkpWXB02OuVtWMW4Gsf973wnVkylFTGIoOd6kB/AYzd0DqFZi0QkmK+uNJGV4Du4QQtOMDul648wOD0glwua91yhFW+2WHh9RB9LMAh+iOBTbgPg2JyyNr0K3zWCjivaMjovriDcJ5QVhwUtOZv13Iat66xR6LkQM4mA+c3z9mvvELw4u7btxtYxlhd5wuWDta63Aof2wscGY2VX/WvV2Pi25KAhxdceTRTe+XzpMQ1tlh2daurz7z9TgKmhKermG9wF5lcCtd96eRILBDA/ZwNItXNSyOt8QFvrcIVw8pA3b029cz/aIUm1zcXFUXiMpej4vpQQqmi8NEt6ozkOChnSb32AbMbVOmOmTL9nguiOVX0vLJ7vwUey+fXfMyHHmSRLAOjCWOO1cHM4/ryM/mhlW0yXJj3C87GQp/nDKoFBx4pB5kmM+bx/HYvDza7pXaOjIWMgQtzw+kSxQLfikIGQx4XZozB6k6nbUSrUBBIb6FnAVw6brCfI0zz+ZagzpbCwbHEsunNnspJP4HnDm9PO7RvD4hsBdQy9xdJHMxql8NRWwmjp2ZI4jaQ8kmEP32088fqBRqTVHXK9oApL/ScWuSd7FpKfJfceazDEvm5yDj3oQvNvJOd+zGVvns1yKGV39SL6cuqS/8vF9O0zOvvmBLIMvc3o5kGAzF/RLRXcI6e6HHklYmQ71n678SL4pcqstC8k2TJsUWewXPZesliOsXj/buvHkCbBVuYcp6ypqH71YJ0Mq/CiEU1MvLDhEsRyKd5fy7RVtruR71Lr2p01O1+yuaU9nLHuVn8USINMVlfqQDK0C6enuePSmsyYc3TlYRfra9zghBNCuH+mNfFojKmSDWecI2fAsUs4CE/neeAnvk6ZM9cjSJFzMFW0XHHjqnOjHpWLQvgJ05yYYGB3m+9ePyE545uruCENo+4oUuA496/gHYso4tz0QMrnlwd1RDmuDtouNJI/fCUYr8MPKTDr2dY+j0ii//oZgKeeGPeVw4w+LvtNQC6yc+8k/t6qGEy97OKkgiuDK/G+fNbQ/32yeJD+TW96VAMtcq1nZm8Nu5p7zJ5CHBFAADXL6bvNCualOj6bFR5QTQSOVR2fk1c9nCJWKwbsF7MEW3ITlm3jky5zRJ5fEWMtdMOF7agoptXtqNqqekL42zkBjTFZoJsnMtbZMDKK00t38f0mTNRYyMkzEptuFFerWmtQ+yJxnBTiqGYssGFhYWhfjSoeT8VccXfUSlakT4JLo7X3B1qoyfPQetyAt4wiT4vSuHRXVTITzxjCQ7sreWkPgUX1C49wMzK0+8ZfcqlKK1adXgw0w7k2cnpn7BWn8HGn/svivX2zjpEYq3luS8FEMlKLY2ZAwwjAsybFUAgGJe9GcxO/9OXNCb/ZW42OgNktkXhyN3JYkaw++sS/McJcDGV5fugE3tcguR4Z9k/4rgNH1IgQQPL6ugKEDgtH8yDEGAxXI5fwVwXS9kIBNNj/9bogm4c8d1XLuj+gLgMHkzSbQL1wuAH+O988fLiek0PXwLc5N/v53IwMhqbec74J716Gio0sGyThbmin+4zG7szSVOe8rj0QE/Lppzle2FSyj9NrpGAVhUM1Vhlwbj3HpLmRAJAh36Oyoni+bmnp7+Tla13que5POZJGaQpr3oz9HzrRzMbsNxW/hWhe85yAsZhF6wEogGwbx/aNviCzKBiwC7Hu9/vfi1a6SpL6mkxQNaUXRC2hr//036gkBLyjlzcDgKO6DK6eP8alvcJaM0vXJmifntKQwhOKWyElRaNReEyhtfICjwAzu9Ga6D6puJDL0rsiRZ/mXv5l2t++240WyRvqb95TQy8IC1pGmhxkaa5v2rrmF/cCuev3d0l/RyigYyPQV8lHFbJtvKN25Sofir3YX/i5Dzx1TIq6QHPqTXCwAzxfwDem7eAkDPB++EMOHrheKSkhLHyM9U+Hwj6oShC9YGZnk48ZxxFl393ObW5rpXhBIc4J1KtEb8Zb+xstM2wV7SWAqUThBKyjD/o0vVBgzAQeed5t1HShHS3Euji3n0k+2Bg7fMbF+SBbR8H+gYc5xTuom3JlmAjMkg3D+wPuB/hPUmX9FhmLyodpxXaakF7P0s6tBiy+8LLJfD0amD+RKhN0l5E5b3RvLpxGl5RRiZRlIA0IzMytLGxsZ+i3vA/GpiYmJxotERTe4YNSm8uDGp8gywSQH5Ss3QmISx94L9MWqWaarqs6+VM5VBCcWzc5RwJKx9DuLUl97qXrrHdUxEiIlNQCDHRDKqwLVSX1ZWNr1pP80uN+SGEgeNuVyRRWaekM+5MNRXOa1GHGwWJP1A1jwLNw9gvUUyl9+ReXLOmggM5pgdy6/ooTeJHkeXwwwJMUke6DKzdJZlGUnJ8S3bMOV09MSQUgzwXqbSfg279xepi0CG9U5dKdnrhSTAdknwtD8nOxIJaCV48279Ab1I2MCROOR3fdblkWedzp/zhe/vBU+aZ7MnSrzLqrCvOcylwH+q2dnLJjcLr7y16VNRlraN/cp9L7dQymz5CaNujAveuUYAv8UXJOh7ltqIlt2XWUbfLJCS4uf10lJaSi75ws2odAl4aNUIfU9DLbxda66ZNFXdEtK7ZMW9xpwCfVm2XWNZ2cYCuarvzLMuOYvAFCYQC/V1Ca2yHQd9AfwHiwIo75ifP6BWPggmvGJvjk1lP1pREwEuYKfDFivZIBtEFWpSCCDsXCQP9x8Q4he5ThWccVpk6GFc0glE4ALnsvlr6xv40yIcsTj3BdeIy07/kkY9uXCPTkx7AA1Ki1rUdmoT+e83zrVLtM6ZD30zGVZzTTIEPhjnIaeFumpybYttiU1DsVLmoKxtCrUycO6mSUQBFi7QlzzHeJgTsi/3vtkAOAci7aqIPacgmDpuZ3s0rONS+qyeBTXCcU0pdtz55tQYMXaGCeM2E9iK+kDNjfnnz1O2WULs9kHTpPr5RfMhRhJp4PlPGR5I0x5uGYm04m8Lt8Kzxq72oyY1tx/H3tcSubVZAWXx2MhTp1XSwYZlZkcN0nTBzTbx/JCiSwCWAxGspbxkds1sqbhELrXkGcX8omuEs2vIlwGjvIRe0/lAwM9UOLs+f4BMZQjYrX76X+ykLTE0EMi/Oqq+gGLDXhTjQT9w1Q4JLCZl22SxGPG8KC67KFbIXc7PiERlI4BdxEiKxKaWYo+f3hjHct3lds84O1iumwxBmCrEXhwAvFIVD/A3jO08NIGEVM8c7s2GmHlMuq00I3LOpnsRZ4RUn60Bz/OL+W8QWj7BsgksWRJ44bvveA8mbjkew+ur622GL47z7RE7Kh6FZ6PCxtX+iFPaKkHnVd+tksSLjV89dUbVpILuNS5LX42NBDtHgPed+ypJ8G2VinSiQO3x4ShRYd3pmurwwb6KlpPIT+1KW8NyQVr6v9Kbz+U2taodpL7bCDg4OABQkkWz6snEXynObSu0Umw7pTikQ6t7v5tiW7QOOXLMHt/cNqK24IknYtYIu8je4K4uRhIN/fOFjnApgRkUaNFeG0zFoR4RQOuQOHoMPqKS2XAIF64WvICGbgsc9lOMaiIyHDcvD2g97Lu7t4r021+vZ1kzFmjVVkm4Z8+ZVzjbbk9W2nX6Dq8u0bVNuTqCexMFp/kw4/B/SnFHxYeUGPeKRyqSt2g4gGXiwYb2VfYd8KzKB5WcrflvgtyZYvIGjGiL/BrY7lQYz/c61O9/7Z280/ml0CP9rKn90voU/qXECoLYbancCsDndxaudv3H0OFI6XVJG27p4nIGXjSbPhA6PpnHlxR0ObwUXP3pscs+Nb/tcVLFBadB+pnbXud3dR59Or6j+hhNlOI8dnAawSIndosaiCYwZP6UvS8a20uEsg/xy+WuwSthwBWhEQCuyAuTduA6DM0o3/2BMaP8KWTqD/bQQQu/PJFonwfR1ey3i+QciHpUXUYPmBAca0MvdjPwQGt3g2EYsC/ScytQWvvjo+WU9sD9foq7X8eJ13uXS1ehTlk08rSkTq+ooLakDEgL9g5ItucK2vIs6Z1WzgE8EX1SImWhP4wqkyCUE4qiBD7Tf+G9PGV82OBFinzAOlTzuugZUhR7MdP8u+wJH/m3sqdl3fP4A3/hLlJrgR+3onOv6g602tGiTDYRnEXUHruFbaTxlpq2jDPvV/oexdZbq4isGuDdk+gCnD3zCDJf07G0tKRlrvHxCGG4q7XIXUcZdnh4eEGhfVwLRCgyAVfVl+3RxbNrXrCgRiUR+ExTTVFRcpUM0UqqKLSTqxMzvHJJrbcTA6bE3+GfSCUaOlkKi2GldTSyM/3xdaCjqQUt5wEuEcNF+ffvfhrbbm2v1N6+nThzX8FOL8kQBEfguuY4P5B9QFZmwgPihAe5XWGVVr/xBiRfP1uwCpK27ewZJNESRyeS+hmETD0B4MQqvosej4OxSUMrwkyuJyblPfNpmcm5qAzbI+q/0i10SFwR0RyTdL0gBFhcaP/Nn/oZIzXgNiJtVQkUrxdMAC9hIpB9nabySkdmpuTxzD7cvs5bSek09X8eswtdO6fNMPmXLEuGkvceQPmj5pk7/fTlFLbHh1YUSVa3Aee3EU4Ag+Y4CojHMXG9c70vWINPiFRXrKQVrkEL5ozF9KX4j95aVFxizqbSAPhtiGEpvadnj/nirevM20OWwxyp/uQ5ABK43ifTFzipGca8LAIiYF+GiPFusjnAJpo/ILUVcfXMhFiz3fn0ixunqqANOlC6tAUKgByJnguDP/r/S5olPtVBQMPtoUJrrGqctJJ2c9MqIzcafeXHnuQUCA/FaCqucV17YImXB6fWKFO31U9k0V1Lv3UVcjjqa/9/TbXAPf+kWr63vfAIjOmSzowVHm1/3aXh8CfLsuPXnJ/QFboMBKuNHtvoH1s0BDdSx4cms5qmsLLa/AZQIFolRag633konhTteTG5VnDlJARIouU+cQAw3bzSfxNVsN1VAiELdf7zm65hLidy3Gg5Rclzr2idMsd5KFpWAeyM9GTQOJN+NXYZpBfr/+E8mDIIp6RVt+sK4XWyAZh069rcq8y3UXH45hVXW6rNtb9QBA5CKab8f1FNNzA2Nt6uoWmnZ9aooMzWqHW2R8oo+tMlotqQuew8NBDpPjGSmyGT2Se+4NeLrjuTjuZZL2Ntl+zz4gc6bEChi1l2tdfZbJ57qYHI1FN1FIv2r67UR4LStKQdt+yf3KsWp+01oNVSmkmCbYIBMlDXnwQBUyyQcQ19EjP23rWGS/w72s2O3QY4z2zDOLtyHK6KjHSDRo+hcXMI7/jXPy4Ff079hmioZdW5vVlgp9QSZVEpwA+WbJGsIM7c0nJrkYEoZXNzs3bb/7cvSc6p7b0WUzz6LaYAWFRJo6+Jvr63UarvDUHngxUTfkJ6PBp290kRmTFrATGB86ySHNjrnujw5yWVPj0XIq2NTS2vFIMwjiJwT60sLDYdXaaOJSMAALwuMIEI+Z2QxbhU5Md+GhSqzhEXEnPeWjcw6/rNU+9PEvu3iCB/Np3Dut+g1gSJfnlR93/+q5YiRTto7/8fEcVQpdhGdqoG6SodcOTOf9Ma/3kiUiYQAwzzkBRVzperAMlWTz6BXysBbNUobwcd/zenETFG6igmP4khsxPESh3Hj8hDy9ehUtzQOPy226B3GvG4a5qXTRq7EIFrFQVgya1P9Uzg3dnbFupznje51JdQinJYaWl3UpeKAQZJKp1vM3StR0CGcR1mhkhf37cKipZ6iYdGcoUi4WPs/ZIojgXxXGsCKDh1xaiSjhQoxt5ecsh0xBRmA5wS9Qs+h6wGSKnXwM9f/BkAg8NdT2yM/b8Ab9i6AVo8C2F7sZP6QKg99kpH8W/5iHAPDQ3A88iF139TronDP63+zY0iJ7HSRWhxlDzWyAASFpsXptIJBTBnP9Y1Dfv2lxxqfHLtH44uIcrXiP8xIyrmiSCXQubRUUUoQSgitPaPrDPvriDwPtRnzmxrvx+Bvq4a/BD0C4gATkhbJnDoM9ywZPfycDcfyD/Gz5pr5Ki5XA/Sru4JnBMAYo7cLckHSqoQgN7UuKHVn09vbwkgyhimzuoXE8qnefmR+yVQ6BJTJ93veSMztlcRVQEiEyGDgLpPU0154wuMup1ka6tHWhOIAnuFidJ6RIKltrQLl/BpXLnSMr1p80M5+xJsGH+O3Wacob1H0oK5RhXE9HQ6duUhAwy6PzvXUmg2DwGHnFgD3pIrjiJzsDUAf6RtFUKyy4ofD+vlg9jq33HkUu25q6cErl//dOxQf9aEqJo9PgNtn0ZudWyfSUz1XvpPb112Nh87TMPJtlfdDxLPy7eWE9yLXWZ9/8/EDapjK1+tXL7Lv/ZV8PVas99nbSJAg157FvXuUnfq5d2aJlic/e/Td7GelZ+M+edq6XcA6/K14d8FZ7Bz5+x4BIqjv2Fi6+ztYuJ2mGbKWc1tE5s7uV2+bsa/y0XPzsx8V/u+Ov5n2Ye7uUI7FI99+/YtId5OhRE0BtKbPr827JXKzorev1Na+J4vTkuInWQ119ScqQjo0ZOPgeV7yS7D9ufnzvlt37nz5PTmoy/X7BSccLny46OjednZ2Rfqf07Y3bwSMhxSZ/d95/mTf0rsdnhydV++4fr+bPb/7dXfdj7f9bf++f6vc/vLn5+rTZXSevxP+/+n96X137f7fb1vb7ni+b7f76T/1GQt+H/85RO/rHduDqA55e9xdT+fvmtkl9PYlFdRbnbH5uiC2e/PHQ/f/+tq5r/tZxifN89p2LmrwVo6/tel6bu+7/u+19xgzt8HoOGzn38+nz99x2U9sM7pn5U9MTr4zqwtL64v+sN7IOfZ89cCO5g57tT+kL1r/+PH8fyAqqf89Zutt2y7ubcsNf3X5BfL/3/d/8Luw/W6h21lP5FLpXcOTM72d9fU/o5/O+/PPfdvdUe/PZf72fiu2/5r+bsDL5SAxdQDBY4U5gKBKY+VWfg6Dqz/33++fHr83Zdx1X/Pvtv3+1i9aLlD0SNgszqhEdSCOtAsnKlYYMxyrT+vrDzUhvF9Ws5yf2Dy/CRwwvDAU8nTLkJM1Sw/Hi/GtSpgIn8fT2DF8/7ukIcvWNy3HeM6nhny0EPK7MpeYAvljfbcIwf9HZKPr3jxvqbgzum/Iq7SfTkrF5VcmyMLqvLd25GGcTz8TsJqtnK7uJdPdoh/XMXEXlQAGhCwL4NVbBVlvyMKzST4p+9/4trtZ2F0POrZpu7mO4+vgsa4dv18V3JJnk+3GFhaqSesg6ynrEmW23ok8NqRaX+O7pcqBhV+GxuU3Y/G/ttlit5S/s8e+YCjOlFxxR0GIPB09XNZ55TQBABQSwcIUIv1pzOLAABnmAAAUEsBAhQAFAAICAgAqrL+TAatzOTLBAAADyYAABcAAAAAAAAAAAAAAAAAAAAAAGdlb2dlYnJhX2RlZmF1bHRzMmQueG1sUEsBAhQAFAAICAgAqrL+TMQHrm/gAgAAwgwAABcAAAAAAAAAAAAAAAAAEAUAAGdlb2dlYnJhX2RlZmF1bHRzM2QueG1sUEsBAhQAFAAICAgAqrL+TEXM3l0aAAAAGAAAABYAAAAAAAAAAAAAAAAANQgAAGdlb2dlYnJhX2phdmFzY3JpcHQuanNQSwECFAAUAAgICACqsv5M50AIx7IIAABPHQAADAAAAAAAAAAAAAAAAACTCAAAZ2VvZ2VicmEueG1sUEsBAhQAFAAICAgAqrL+TFCL9acziwAAZ5gAABYAAAAAAAAAAAAAAAAAfxEAAGdlb2dlYnJhX3RodW1ibmFpbC5wbmdQSwUGAAAAAAUABQBMAQAA9pwAAAAA"
};
$result;
}

ENDDOCUMENT();        
