##DESCRIPTION
##  Polynomials: Multiplicity of zeroes
##ENDDESCRIPTION

##KEYWORDS('algebra', 'polynomials', 'zeroes', 'multiplicities')

## DBsubject('Algebra')
## DBchapter('Polynomial and Rational Functions')
## DBsection('Polynomial Functions')
## Date('Spring 2012')
## Author1('David Covert')
## Author2('Teri Christiansen')
## Author3('Rick Lynch')
## Institution('University of Missouri')
## TitleText1('College Algebra')
## EditionText1('4')
## AuthorText1('Beecher, Penna, Bittinger')
## Section1('4.1')


########################################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PeriodicRerandomization.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
PeriodicRerandomization("3");

#######################################
#  Set-up

$a = random(-1,1,2)*random(2,10,1);
$b = random(2,10,1);
do {$c = random(2,10,1)} until ($c != $b);
$n1 = random(2,12,1);
do {$n2 = random(2,12,1)} until ($n2 != $n1);
do {$n3 = random(2,12,1)} until ($n3 != $n1 && $n3 != $n2);


$f = "$a (x- $b)^{$n1} (x + $c)^{$n2} x^{$n3}";

## Answers
$ans = List($b,-$c,0);
$multians = MultiAnswer($b, $n1, -$c, $n2, 0, $n3)->with(
  singleResult => 0,
  allowBlankAnswers => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ( $e1, $e2, $e3, $e4, $e5, $e6) = @{$student};
      my ( $f1, $f2, $f3, $f4, $f5, $f6) = @{$correct};
      if ( 

($e1 == $f1 && $e2 == $f2 && $e3 == $f3 && $e4 == $f4 && $e5 == $f5 && $e6 == $f6) ||
($e1 == $f1 && $e2 == $f2 && $e3 == $f5 && $e4 == $f6 && $e5 == $f3 && $e6 == $f4) ||
($e1 == $f3 && $e2 == $f4 && $e3 == $f1 && $e4 == $f2 && $e5 == $f5 && $e6 == $f6) ||
($e1 == $f3 && $e2 == $f4 && $e3 == $f5 && $e4 == $f6 && $e5 == $f1 && $e6 == $f2) ||
($e1 == $f5 && $e2 == $f6 && $e3 == $f1 && $e4 == $f2 && $e5 == $f3 && $e6 == $f4) ||
($e1 == $f5 && $e2 == $f6 && $e3 == $f3 && $e4 == $f4 && $e5 == $f1 && $e6 == $f2)

   )
{return [1,1,1,1,1,1];}  

return [0,0,0,0,0,0]; 

} );



################################################
#  Main text

BEGIN_TEXT
Find the zeroes and multiplicities of the polynomial \(f(x) = $f\).
$PAR
$SPACE The zeroes are \(x = \) \{ans_rule(20)\} $SPACE $SPACE (separate your answers by commas).
$PAR
$SPACE The zero \(x = \)  \{$multians->ans_rule(5)\}  has multiplicity \{ $multians->ans_rule(5)\}.
$BR
$SPACE The zero \(x = \)  \{$multians->ans_rule(5)\}  has multiplicity \{ $multians->ans_rule(5)\}.
$BR
$SPACE The zero \(x = \)  \{$multians->ans_rule(5)\}  has multiplicity \{ $multians->ans_rule(5)\}.


END_TEXT

Context()->texStrings;
if ($attempts_modp == 0 && $actualAttempts != 0) {
    BEGIN_TEXT
    $PAR
    $BBOLD Answer: $EBOLD
    \(x = $ans\) are zeros with multiplicities \($n1, $n2\) and \($n3\), respectively
    $PAR
    END_TEXT
} else {
    BEGIN_TEXT
    $BR
    END_TEXT
}

Context()->normalStrings;

PeriodicStatus(); 




################################################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

ANS( Compute($ans)->cmp() );
ANS( Compute($multians)->cmp() );


COMMENT('MathObject version');
ENDDOCUMENT();
